Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.real
% (61769)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 151
% Passive clauses: 2159
% Generated clauses: 24430
% Final active clauses: 126
% Final passive clauses: 1756
% Input formulas: 23
% Initial clauses: 38
% 
% Duplicate literals: 1130
% Trivial inequalities: 351
% Fw subsumption resolutions: 3981
% Fw demodulations: 3808
% Bw demodulations: 20
% 
% Simple tautologies: 394
% Equational tautologies: 54
% Forward subsumptions: 11502
% 
% Binary resolution: 365
% Factoring: 3
% Forward superposition: 7923
% Backward superposition: 7012
% Self superposition: 262
% Equality factoring: 149
% Equality resolution: 27
% 
% Split clauses: 510
% Split components: 1033
% SAT solver clauses: 1283
% SAT solver unit clauses: 3
% SAT solver binary clauses: 80
% 
% Sat splits: 510
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1283
% TWLsolver calls for satisfiability: 160
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (61787)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 548
% Generated clauses: 13890
% Final active clauses: 97
% Final passive clauses: 407
% Input formulas: 23
% Initial clauses: 38
% 
% Duplicate literals: 1010
% Trivial inequalities: 166
% Fw subsumption resolutions: 1599
% Fw demodulations: 1218
% Bw demodulations: 42
% Condensations: 1595
% Global subsumptions: 5
% 
% Simple tautologies: 242
% Equational tautologies: 55
% Forward subsumptions: 6392
% 
% Binary resolution: 287
% Factoring: 5
% Forward superposition: 4489
% Backward superposition: 2898
% Self superposition: 178
% Equality factoring: 135
% Equality resolution: 20
% 
% Split clauses: 185
% Split components: 371
% SAT solver clauses: 1145
% SAT solver unit clauses: 55
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 185
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1124
% TWLsolver calls for satisfiability: 208
% 
% Memory used [KB]: 8187
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f27397,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f128,f133,f131,f154,f158,f156,f167,f172,f302,f366,f454,f459,f471,f472,f458,f456,f518,f519,f646,f664,f693,f700,f695,f702,f1502,f1616,f1618,f1620,f1622,f1624,f1626,f1628,f1630,f1632,f1634,f1636,f1637,f1639,f1606,f1781,f1784,f1787,f1790,f1793,f1796,f1799,f1802,f1805,f1814,f1819,f1886,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f2463,f2496,f2494,f2526,f2566,f2564,f2634,f2636,f2619,f2620,f2621,f2622,f2644,f2623,f2643,f2645,f2646,f2640,f2796,f2869,f2801,f2873,f2812,f2876,f2868,f2882,f2885,f2884,f2841,f2894,f2871,f2896,f2897,f2898,f1516,f2899,f2860,f2902,f2900,f2861,f2905,f2970,f2984,f3980,f3981,f3991,f3994,f2903,f4066,f4070,f4074,f4075,f1518,f4209,f4214,f4213,f4433,f4438,f4437,f4435,f759,f318,f316,f2214,f1179,f387,f385,f4547,f4553,f4564,f4572,f4616,f4618,f4549,f4589,f4634,f4642,f4716,f4717,f4689,f4620,f4548,f4722,f4723,f4724,f4734,f4767,f4772,f4778,f4786,f4792,f4799,f4805,f4811,f4815,f4822,f4827,f4833,f4835,f4837,f4839,f4841,f4842,f4844,f4845,f4847,f4848,f4850,f4851,f4853,f4854,f4856,f4858,f4860,f4862,f4864,f4866,f4868,f4870,f4871,f4872,f4873,f4874,f4875,f4877,f4878,f4879,f4881,f4882,f4883,f4884,f4885,f6295,f6365,f6363,f7240,f3993,f7371,f3996,f7506,f7567,f7565,f7543,f7593,f7591,f7546,f7597,f7595,f7549,f7550,f7601,f7551,f7596,f7602,f7552,f7592,f7599,f7637,f7638,f7640,f7645,f7648,f7647,f7600,f7723,f7736,f7771,f7772,f7904,f7907,f7913,f7919,f7923,f7927,f7931,f7935,f7948,f7977,f8008,f8010,f8012,f8015,f8021,f8023,f8025,f8027,f8028,f8030,f8033,f8034,f8035,f8056,f8067,f8059,f8117,f8130,f8133,f8134,f8135,f8229,f8230,f8231,f8232,f8233,f8234,f8235,f8239,f8256,f8260,f8277,f8282,f8317,f8324,f8335,f8343,f8358,f8373,f8381,f8389,f8393,f8450,f8451,f8452,f8453,f8454,f8455,f8456,f8457,f8458,f8459,f8460,f8461,f8462,f8463,f8465,f8466,f8468,f8469,f8512,f8514,f8516,f7542,f8518,f4086,f7541,f8517,f8519,f6925,f8520,f6799,f8521,f5846,f8522,f8640,f8642,f8765,f8792,f8817,f8820,f8865,f8872,f8879,f8891,f8894,f6596,f8896,f6595,f8898,f8058,f8858,f8903,f8901,f8924,f8926,f8979,f8992,f8994,f9000,f9005,f9009,f4211,f1643,f7515,f9010,f7516,f9011,f9016,f9017,f1644,f9244,f9287,f9290,f9292,f9302,f9520,f10067,f2642,f10068,f10069,f10070,f10077,f10073,f10071,f10074,f15974,f15979,f15978,f15976,f16863,f16909,f17831,f17856,f17857,f17833,f21907,f21937,f699,f25632,f25609,f25685,f25594,f25585,f25655,f25513,f25868,f25508,f25503,f25502,f25494,f25493,f25488,f25880,f25482,f25883,f25480,f25474,f25472,f25911,f25924,f25928,f25879,f25934,f25871,f25935,f25937,f25886,f25949,f25950,f25962,f25970,f25976,f25878,f25982,f25984,f25900,f25992,f25997,f25870,f26004,f25897,f26012,f25869,f25408,f26135,f26134,f26138,f10076,f2639,f2637,f26132,f25682,f155,f26252,f26253,f26202,f26205,f26260,f26256,f26210,f26255,f26261,f26859,f26860,f662,f27386])).
fof(f27386,plain,(
  $false | $spl17),
  inference(subsumption_resolution,[],[f27385,f81])).
fof(f81,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.REAL_INV_POS',aHL_TRUTH)).
fof(f27385,plain,(
  ~p(s(bool,'T_0')) | $spl17),
  inference(forward_demodulation,[],[f27384,f129])).
fof(f129,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f125,f84])).
fof(f84,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.REAL_INV_POS',aHL_FALSITY)).
fof(f125,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))),
  inference(superposition,[],[f79,f86])).
fof(f86,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.REAL_INV_POS',aHL_BOOL_CASES)).
fof(f79,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f50])).
fof(f50,plain,(
  ? [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6)))))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_INV_POS')).
fof(f27384,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) | $spl17),
  inference(subsumption_resolution,[],[f27354,f81])).
fof(f27354,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) | $spl17),
  inference(superposition,[],[f2892,f4886])).
fof(f4886,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(subsumption_resolution,[],[f2259,f455])).
fof(f455,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl17),
  inference(cnf_transformation,[],[f455_D])).
fof(f455_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f2259,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f2258])).
fof(f2258,plain,(
  ( ! [X4] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2197,f123])).
fof(f123,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f84,f86])).
fof(f2197,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f465,f269])).
fof(f269,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f266,f84])).
fof(f266,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f157,f142])).
fof(f142,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))),
  inference(unit_resulting_resolution,[],[f84,f80,f111])).
fof(f111,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.real.REAL_INV_POS','thm.bool.IMP_ANTISYM_AX')).
fof(f80,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),
  inference(cnf_transformation,[],[f66])).
fof(f157,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f148,f84])).
fof(f148,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f111,f86])).
fof(f465,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real',X4) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(subsumption_resolution,[],[f464,f84])).
fof(f464,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real',X4) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(forward_demodulation,[],[f436,f142])).
fof(f436,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real',X4) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f108,f269])).
fof(f108,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X11] : (s('type.realax.real',X6) = s('type.realax.real',X11) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X11)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real',X6)))))),
  file('thm.real.REAL_INV_POS','thm.realax.REAL_LT_TOTAL')).
fof(f2892,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(forward_demodulation,[],[f2891,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_NEG_GT|48|')).
fof(f2891,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) )),
  inference(subsumption_resolution,[],[f2890,f84])).
fof(f2890,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) )),
  inference(forward_demodulation,[],[f2889,f1234])).
fof(f1234,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f84,f251,f111])).
fof(f251,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f83,f115])).
fof(f115,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X11] : ~(p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X11)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real',X6)))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_LT_ANTISYM')).
fof(f83,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_LT_|48||49|')).
fof(f2889,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) )),
  inference(forward_demodulation,[],[f2888,f85])).
fof(f2888,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) )),
  inference(subsumption_resolution,[],[f2838,f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => ~s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X11)))) => ~s('type.realax.real',X6) = s('type.realax.real',X11))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_LT_IMP_NE')).
fof(f2838,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f110,f647])).
fof(f647,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f107,f88])).
fof(f88,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_MUL_LINV')).
fof(f107,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X11] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X11)))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_NEG_LMUL')).
fof(f110,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X11] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X11)))))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_LT_MUL')).
fof(f662,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f662_D])).
fof(f662_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f26860,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl222),
  inference(cnf_transformation,[],[f26860_D])).
fof(f26860_D,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f26859,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl30 | $spl33)),
  inference(subsumption_resolution,[],[f26832,f1516])).
fof(f26832,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f5832,f88])).
fof(f5832,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5804,f5783])).
fof(f5783,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5782,f4900])).
fof(f4900,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f455,f88])).
fof(f5782,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f5781,f84])).
fof(f5781,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5780,f1234])).
fof(f5780,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5779,f4900])).
fof(f5779,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f5759,f1518])).
fof(f5759,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl17),
  inference(superposition,[],[f592,f4900])).
fof(f592,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f108,f85])).
fof(f5804,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f2817])).
fof(f2817,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f647])).
fof(f26261,plain,(
  ( ! [X26,X27,X25] : (s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,sK1(X25,X26,bool,bool)) | s(bool,X27) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl220),
  inference(cnf_transformation,[],[f26261_D])).
fof(f26261_D,plain,(
  ( ! [X26,X27,X25] : (s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,sK1(X25,X26,bool,bool)) | s(bool,X27) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f26255,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl216),
  inference(cnf_transformation,[],[f26255_D])).
fof(f26255_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f26210,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X27) = s(bool,sK1(X25,X26,bool,bool))) ) | $spl46),
  inference(superposition,[],[f2637,f24215])).
fof(f24215,plain,(
  ( ! [X187,X188] : (s(bool,X187) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X187) = s(bool,X188) | s(bool,'T_0') = s(bool,X188)) )),
  inference(subsumption_resolution,[],[f24116,f81])).
fof(f24116,plain,(
  ( ! [X187,X188] : (~p(s(bool,'T_0')) | s(bool,X187) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X187) = s(bool,X188) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f1115,f129])).
fof(f1115,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f927,f81])).
fof(f927,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f114,f121])).
fof(f121,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f86,f86])).
fof(f114,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f63])).
fof(f26256,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool))) ) | $spl218),
  inference(cnf_transformation,[],[f26256_D])).
fof(f26256_D,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f26260,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK1(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24)) ) | $spl46),
  inference(forward_demodulation,[],[f26209,f129])).
fof(f26209,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,g(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK1(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24)) ) | $spl46),
  inference(superposition,[],[f2637,f24215])).
fof(f26205,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,X15) = s(bool,sK1(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) ) | $spl46),
  inference(superposition,[],[f2637,f121])).
fof(f26202,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl46),
  inference(superposition,[],[f2637,f121])).
fof(f26253,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl214),
  inference(cnf_transformation,[],[f26253_D])).
fof(f26253_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f26252,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | ($spl46 | $spl166 | $spl170)),
  inference(subsumption_resolution,[],[f26251,f10075])).
fof(f10075,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl170),
  inference(cnf_transformation,[],[f10075_D])).
fof(f10075_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f26251,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | ($spl46 | $spl166)),
  inference(forward_demodulation,[],[f26250,f10073])).
fof(f26250,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | $spl46),
  inference(duplicate_literal_removal,[],[f26200])).
fof(f26200,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | $spl46),
  inference(superposition,[],[f2637,f621])).
fof(f621,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK1(X0,X1,X2,bool)) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f116,f86])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.real.REAL_INV_POS',aHL_EXT)).
fof(f155,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f25682,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl186),
  inference(cnf_transformation,[],[f25682_D])).
fof(f25682_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f26132,plain,(
  ( ! [X428,X427] : (s(bool,'F_0') = s(bool,X428) | s(bool,'T_0') = s(bool,X427) | s(bool,X428) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) )),
  inference(forward_demodulation,[],[f25440,f236])).
fof(f236,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f84,f232,f111])).
fof(f232,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f228,f81])).
fof(f228,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f115,f129])).
fof(f25440,plain,(
  ( ! [X428,X427] : (s(bool,'T_0') = s(bool,X427) | s(bool,X428) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X428) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(duplicate_literal_removal,[],[f25404])).
fof(f25404,plain,(
  ( ! [X428,X427] : (s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,X427) | s(bool,X428) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X428) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f271,f24215])).
fof(f271,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f268,f84])).
fof(f268,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f157,f236])).
fof(f2637,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl46),
  inference(cnf_transformation,[],[f2637_D])).
fof(f2637_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2639,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl49),
  inference(cnf_transformation,[],[f2639_D])).
fof(f2639_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f10076,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl171),
  inference(cnf_transformation,[],[f10076_D])).
fof(f10076_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f26138,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool))))) ) | $spl166),
  inference(forward_demodulation,[],[f10025,f10073])).
fof(f10025,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f10009])).
fof(f10009,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f622,f621])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK1(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f116,f86])).
fof(f26134,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl210),
  inference(cnf_transformation,[],[f26134_D])).
fof(f26134_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f26135,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl212),
  inference(cnf_transformation,[],[f26135_D])).
fof(f26135_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f25408,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f24215])).
fof(f25869,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,X52) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl190),
  inference(cnf_transformation,[],[f25869_D])).
fof(f25869_D,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,X52) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f26012,plain,(
  ( ! [X329,X328] : (s(bool,'F_0') = s(bool,f(s(bool,X329))) | s(bool,'T_0') = s(bool,X328) | s(bool,X329) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25101,f15074])).
fof(f15074,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl171)),
  inference(condensation,[],[f15073])).
fof(f15073,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl171)),
  inference(duplicate_literal_removal,[],[f14972])).
fof(f14972,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f14938,f10726])).
fof(f10726,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,f(s(bool,X1))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl171),
  inference(condensation,[],[f10725])).
fof(f10725,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,f(s(bool,f(s(bool,X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl171),
  inference(duplicate_literal_removal,[],[f10573])).
fof(f10573,plain,(
  ( ! [X19,X20,X18] : (s(bool,X20) = s(bool,f(s(bool,f(s(bool,X19))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl171),
  inference(superposition,[],[f10301,f10301])).
fof(f10301,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | $spl171),
  inference(duplicate_literal_removal,[],[f10277])).
fof(f10277,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl171),
  inference(superposition,[],[f10213,f121])).
fof(f10213,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl171),
  inference(superposition,[],[f10187,f86])).
fof(f10187,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl171),
  inference(forward_demodulation,[],[f10168,f236])).
fof(f10168,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl171),
  inference(unit_resulting_resolution,[],[f10076,f271])).
fof(f14938,plain,(
  ( ! [X222] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X222))))) | s(bool,'T_0') = s(bool,X222)) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f14618,f387])).
fof(f14618,plain,(
  ( ! [X222] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X222))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X222)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f387,f10726])).
fof(f25101,plain,(
  ( ! [X329,X328] : (s(bool,'F_0') = s(bool,f(s(bool,X329))) | s(bool,'T_0') = s(bool,X328) | s(bool,X329) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X328)))))))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f12995,f24215])).
fof(f12995,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl171)),
  inference(condensation,[],[f12994])).
fof(f12994,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f12826])).
fof(f12826,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11863,f10301])).
fof(f11863,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl171)),
  inference(condensation,[],[f11862])).
fof(f11862,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f11743])).
fof(f11743,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11165,f10301])).
fof(f11165,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl171)),
  inference(condensation,[],[f11164])).
fof(f11164,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f11083])).
fof(f11083,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10701,f10301])).
fof(f10701,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl171)),
  inference(condensation,[],[f10700])).
fof(f10700,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X40))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f10629])).
fof(f10629,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X40))))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10316,f10301])).
fof(f10316,plain,(
  ( ! [X18] : (s(bool,'F_0') = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl6 | $spl171)),
  inference(subsumption_resolution,[],[f10294,f123])).
fof(f10294,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f156,f10213])).
fof(f25897,plain,(
  ( ! [X118] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X118))))) | s(bool,X118) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl206),
  inference(cnf_transformation,[],[f25897_D])).
fof(f25897_D,plain,(
  ( ! [X118] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X118))))) | s(bool,X118) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f26004,plain,(
  ( ! [X316,X317] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X317))))) | s(bool,'T_0') = s(bool,X316) | s(bool,X317) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25095,f15076])).
fof(f15076,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl171)),
  inference(condensation,[],[f15075])).
fof(f15075,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl15 | $spl171)),
  inference(duplicate_literal_removal,[],[f14971])).
fof(f14971,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f14938,f10301])).
fof(f25095,plain,(
  ( ! [X316,X317] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X317))))) | s(bool,'T_0') = s(bool,X316) | s(bool,X317) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X316)))))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15325,f24215])).
fof(f15325,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl15 | $spl171)),
  inference(condensation,[],[f15324])).
fof(f15324,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl171)),
  inference(duplicate_literal_removal,[],[f15196])).
fof(f15196,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X14))))))))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15076,f10726])).
fof(f25870,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X60))))) | s(bool,X60) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl192),
  inference(cnf_transformation,[],[f25870_D])).
fof(f25870_D,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X60))))) | s(bool,X60) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f25997,plain,(
  ( ! [X306,X307] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X307))))) | s(bool,'T_0') = s(bool,X306) | s(bool,X307) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25090,f15076])).
fof(f25090,plain,(
  ( ! [X306,X307] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X307))))) | s(bool,'T_0') = s(bool,X306) | s(bool,X307) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X306)))))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f12995,f24215])).
fof(f25992,plain,(
  ( ! [X298,X299] : (s(bool,'F_0') = s(bool,f(s(bool,X299))) | s(bool,'T_0') = s(bool,X298) | s(bool,X299) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25086,f15076])).
fof(f25086,plain,(
  ( ! [X298,X299] : (s(bool,'F_0') = s(bool,f(s(bool,X299))) | s(bool,'T_0') = s(bool,X298) | s(bool,X299) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X298)))))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11863,f24215])).
fof(f25900,plain,(
  ( ! [X122] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X122))))))) | s(bool,X122) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl208),
  inference(cnf_transformation,[],[f25900_D])).
fof(f25900_D,plain,(
  ( ! [X122] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X122))))))) | s(bool,X122) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f25984,plain,(
  ( ! [X282,X281] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X282))))))) | s(bool,'T_0') = s(bool,X281) | s(bool,X282) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25078,f14938])).
fof(f25078,plain,(
  ( ! [X282,X281] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X282))))))) | s(bool,'T_0') = s(bool,X281) | s(bool,X282) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X281)))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15325,f24215])).
fof(f25982,plain,(
  ( ! [X277,X278] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X278))))) | s(bool,'T_0') = s(bool,X277) | s(bool,X278) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25076,f14938])).
fof(f25076,plain,(
  ( ! [X277,X278] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X278))))) | s(bool,'T_0') = s(bool,X277) | s(bool,X278) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X277)))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15074,f24215])).
fof(f25878,plain,(
  ( ! [X76] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X76))))))) | s(bool,X76) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl196),
  inference(cnf_transformation,[],[f25878_D])).
fof(f25878_D,plain,(
  ( ! [X76] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X76))))))) | s(bool,X76) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f25976,plain,(
  ( ! [X269,X270] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X270))))))) | s(bool,'T_0') = s(bool,X269) | s(bool,X270) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25072,f14938])).
fof(f25072,plain,(
  ( ! [X269,X270] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X270))))))) | s(bool,'T_0') = s(bool,X269) | s(bool,X270) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X269)))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f12995,f24215])).
fof(f25970,plain,(
  ( ! [X261,X262] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X262))))) | s(bool,'T_0') = s(bool,X261) | s(bool,X262) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25068,f14938])).
fof(f25068,plain,(
  ( ! [X261,X262] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X262))))) | s(bool,'T_0') = s(bool,X261) | s(bool,X262) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X261)))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11863,f24215])).
fof(f25962,plain,(
  ( ! [X246,X245] : (s(bool,'F_0') = s(bool,f(s(bool,X246))) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25061,f14938])).
fof(f25061,plain,(
  ( ! [X246,X245] : (s(bool,'F_0') = s(bool,f(s(bool,X246))) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X245)))))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11165,f24215])).
fof(f25950,plain,(
  ( ! [X229,X228] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X229))))) | s(bool,'T_0') = s(bool,X228) | s(bool,X229) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25053,f10174])).
fof(f10174,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl171),
  inference(superposition,[],[f10076,f86])).
fof(f25053,plain,(
  ( ! [X229,X228] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X229))))) | s(bool,'T_0') = s(bool,X228) | s(bool,X229) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X228)))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15076,f24215])).
fof(f25949,plain,(
  ( ! [X227,X226] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X227))))))) | s(bool,'T_0') = s(bool,X226) | s(bool,X227) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f25052,f10174])).
fof(f25052,plain,(
  ( ! [X227,X226] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X227))))))) | s(bool,'T_0') = s(bool,X226) | s(bool,X227) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X226)))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15074,f24215])).
fof(f25886,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,f(s(bool,f(s(bool,f(s(bool,X103))))))) | s(bool,X103) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl204),
  inference(cnf_transformation,[],[f25886_D])).
fof(f25886_D,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,f(s(bool,f(s(bool,f(s(bool,X103))))))) | s(bool,X103) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f25937,plain,(
  ( ! [X215,X214] : (s(bool,X215) = s(bool,f(s(bool,f(s(bool,f(s(bool,X215))))))) | s(bool,'T_0') = s(bool,X214) | s(bool,X215) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(subsumption_resolution,[],[f25046,f10174])).
fof(f25046,plain,(
  ( ! [X215,X214] : (s(bool,X215) = s(bool,f(s(bool,f(s(bool,f(s(bool,X215))))))) | s(bool,'T_0') = s(bool,X214) | s(bool,X215) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X214)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f12114,f24215])).
fof(f12114,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl11 | $spl171)),
  inference(condensation,[],[f12113])).
fof(f12113,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl11 | $spl171)),
  inference(duplicate_literal_removal,[],[f11977])).
fof(f11977,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11893,f10301])).
fof(f11893,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl11 | $spl171)),
  inference(subsumption_resolution,[],[f11746,f4976])).
fof(f4976,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(superposition,[],[f318,f236])).
fof(f11746,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11165,f86])).
fof(f25935,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X211))))))) | s(bool,'T_0') = s(bool,X210) | s(bool,X211) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(subsumption_resolution,[],[f25044,f10174])).
fof(f25044,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X211))))))) | s(bool,'T_0') = s(bool,X210) | s(bool,X211) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X210)))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11863,f24215])).
fof(f25871,plain,(
  ( ! [X62] : (s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,X62) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl194),
  inference(cnf_transformation,[],[f25871_D])).
fof(f25871_D,plain,(
  ( ! [X62] : (s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,X62) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f25934,plain,(
  ( ! [X208,X209] : (s(bool,f(s(bool,X209))) = s(bool,f(s(bool,f(s(bool,X209))))) | s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl171),
  inference(subsumption_resolution,[],[f25043,f10174])).
fof(f25043,plain,(
  ( ! [X208,X209] : (s(bool,f(s(bool,X209))) = s(bool,f(s(bool,f(s(bool,X209))))) | s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X208)))) ) | $spl171),
  inference(superposition,[],[f11687,f24215])).
fof(f11687,plain,(
  ( ! [X0] : (s(bool,f(s(bool,f(s(bool,X0))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl171),
  inference(condensation,[],[f11686])).
fof(f11686,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,f(s(bool,X12))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | $spl171),
  inference(duplicate_literal_removal,[],[f11582])).
fof(f11582,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,f(s(bool,X12))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | $spl171),
  inference(superposition,[],[f10703,f10301])).
fof(f10703,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl171),
  inference(condensation,[],[f10702])).
fof(f10702,plain,(
  ( ! [X35,X34] : (s(bool,f(s(bool,X35))) = s(bool,f(s(bool,f(s(bool,X35))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | $spl171),
  inference(duplicate_literal_removal,[],[f10627])).
fof(f10627,plain,(
  ( ! [X35,X34] : (s(bool,f(s(bool,X35))) = s(bool,f(s(bool,f(s(bool,X35))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl171),
  inference(superposition,[],[f10213,f10301])).
fof(f25879,plain,(
  ( ! [X78] : (s(bool,X78) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,X78) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl198),
  inference(cnf_transformation,[],[f25879_D])).
fof(f25879_D,plain,(
  ( ! [X78] : (s(bool,X78) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,X78) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f25928,plain,(
  ( ! [X202,X201] : (s(bool,X202) = s(bool,f(s(bool,f(s(bool,X202))))) | s(bool,'T_0') = s(bool,X201) | s(bool,X202) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(subsumption_resolution,[],[f25040,f10174])).
fof(f25040,plain,(
  ( ! [X202,X201] : (s(bool,X202) = s(bool,f(s(bool,f(s(bool,X202))))) | s(bool,'T_0') = s(bool,X201) | s(bool,X202) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X201)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11326,f24215])).
fof(f11326,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl11 | $spl171)),
  inference(condensation,[],[f11325])).
fof(f11325,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl11 | $spl171)),
  inference(duplicate_literal_removal,[],[f11230])).
fof(f11230,plain,(
  ( ! [X12,X11] : (s(bool,f(s(bool,X12))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11187,f10301])).
fof(f11187,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl11 | $spl171)),
  inference(subsumption_resolution,[],[f11086,f4976])).
fof(f11086,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10701,f86])).
fof(f25924,plain,(
  ( ! [X191,X192] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X192))))) | s(bool,'T_0') = s(bool,X191) | s(bool,X192) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(subsumption_resolution,[],[f25036,f10174])).
fof(f25036,plain,(
  ( ! [X191,X192] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X192))))) | s(bool,'T_0') = s(bool,X191) | s(bool,X192) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X191)))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11165,f24215])).
fof(f25911,plain,(
  ( ! [X178,X179] : (s(bool,'F_0') = s(bool,f(s(bool,X179))) | s(bool,'T_0') = s(bool,X178) | s(bool,X179) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(subsumption_resolution,[],[f25030,f10174])).
fof(f25030,plain,(
  ( ! [X178,X179] : (s(bool,'F_0') = s(bool,f(s(bool,X179))) | s(bool,'T_0') = s(bool,X178) | s(bool,X179) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,f(s(bool,X178)))) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10701,f24215])).
fof(f25472,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X122))))))) | s(bool,'T_0') = s(bool,X121) | s(bool,X122) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(duplicate_literal_removal,[],[f25010])).
fof(f25010,plain,(
  ( ! [X121,X122] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X122))))))) | s(bool,'T_0') = s(bool,X121) | s(bool,X122) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X121)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15076,f24215])).
fof(f25474,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X118))))) | s(bool,'T_0') = s(bool,X117) | s(bool,X118) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl15 | $spl171)),
  inference(duplicate_literal_removal,[],[f25008])).
fof(f25008,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X118))))) | s(bool,'T_0') = s(bool,X117) | s(bool,X118) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X117)) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f14938,f24215])).
fof(f25480,plain,(
  ( ! [X103,X102] : (s(bool,X103) = s(bool,f(s(bool,f(s(bool,f(s(bool,X103))))))) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(duplicate_literal_removal,[],[f25001])).
fof(f25001,plain,(
  ( ! [X103,X102] : (s(bool,X103) = s(bool,f(s(bool,f(s(bool,f(s(bool,X103))))))) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11893,f24215])).
fof(f25883,plain,(
  ( ! [X99] : (s(bool,f(s(bool,f(s(bool,X99))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X99))))))) | s(bool,X99) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl202),
  inference(cnf_transformation,[],[f25883_D])).
fof(f25883_D,plain,(
  ( ! [X99] : (s(bool,f(s(bool,f(s(bool,X99))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X99))))))) | s(bool,X99) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f25482,plain,(
  ( ! [X99,X98] : (s(bool,f(s(bool,f(s(bool,X99))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X99))))))) | s(bool,'T_0') = s(bool,X98) | s(bool,X99) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl171),
  inference(duplicate_literal_removal,[],[f24999])).
fof(f24999,plain,(
  ( ! [X99,X98] : (s(bool,f(s(bool,f(s(bool,X99))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X99))))))) | s(bool,'T_0') = s(bool,X98) | s(bool,X99) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X98)) ) | $spl171),
  inference(superposition,[],[f11687,f24215])).
fof(f25880,plain,(
  ( ! [X89] : (s(bool,f(s(bool,X89))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X89))))))) | s(bool,X89) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl200),
  inference(cnf_transformation,[],[f25880_D])).
fof(f25880_D,plain,(
  ( ! [X89] : (s(bool,f(s(bool,X89))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X89))))))) | s(bool,X89) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f25488,plain,(
  ( ! [X88,X89] : (s(bool,f(s(bool,X89))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X89))))))) | s(bool,'T_0') = s(bool,X88) | s(bool,X89) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(duplicate_literal_removal,[],[f24995])).
fof(f24995,plain,(
  ( ! [X88,X89] : (s(bool,f(s(bool,X89))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X89))))))) | s(bool,'T_0') = s(bool,X88) | s(bool,X89) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X88)) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11326,f24215])).
fof(f25493,plain,(
  ( ! [X78,X77] : (s(bool,X78) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,'T_0') = s(bool,X77) | s(bool,X78) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl11 | $spl171)),
  inference(duplicate_literal_removal,[],[f24991])).
fof(f24991,plain,(
  ( ! [X78,X77] : (s(bool,X78) = s(bool,f(s(bool,f(s(bool,X78))))) | s(bool,'T_0') = s(bool,X77) | s(bool,X78) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X77)) ) | ($spl6 | $spl11 | $spl171)),
  inference(superposition,[],[f11187,f24215])).
fof(f25494,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X76))))))) | s(bool,'T_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f24990])).
fof(f24990,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X76))))))) | s(bool,'T_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f11165,f24215])).
fof(f25502,plain,(
  ( ! [X61,X62] : (s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl171),
  inference(duplicate_literal_removal,[],[f24984])).
fof(f24984,plain,(
  ( ! [X61,X62] : (s(bool,f(s(bool,X62))) = s(bool,f(s(bool,f(s(bool,X62))))) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X61)) ) | $spl171),
  inference(superposition,[],[f10703,f24215])).
fof(f25503,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X60))))) | s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f24983])).
fof(f24983,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X60))))) | s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10701,f24215])).
fof(f25508,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | ($spl6 | $spl171)),
  inference(duplicate_literal_removal,[],[f24980])).
fof(f24980,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') = s(bool,f(s(bool,X52))) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X51)) ) | ($spl6 | $spl171)),
  inference(superposition,[],[f10316,f24215])).
fof(f25868,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,f(s(bool,X41))) | s(bool,X41) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl188),
  inference(cnf_transformation,[],[f25868_D])).
fof(f25868_D,plain,(
  ( ! [X41] : (s(bool,X41) = s(bool,f(s(bool,X41))) | s(bool,X41) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f25513,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,f(s(bool,X41))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl171),
  inference(duplicate_literal_removal,[],[f24976])).
fof(f24976,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,f(s(bool,X41))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X40)) ) | $spl171),
  inference(superposition,[],[f10213,f24215])).
fof(f25655,plain,(
  ( ! [X34,X32] : (s(bool,X32) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X32) = s(bool,X34) | s(bool,X34) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) | $spl184),
  inference(cnf_transformation,[],[f25655_D])).
fof(f25655_D,plain,(
  ( ! [X34,X32] : (s(bool,X32) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X32) = s(bool,X34) | s(bool,X34) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f25585,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X29) = s(bool,X31) | s(bool,X29) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f24459])).
fof(f24459,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X29) = s(bool,X31) | s(bool,X29) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f24215,f24215])).
fof(f25594,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,X40) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f24383])).
fof(f24383,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,X40) | s(bool,X40) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f24215,f24215])).
fof(f25685,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') = s(bool,X17) | s(bool,X17) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f25603,f236])).
fof(f25603,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X17) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(duplicate_literal_removal,[],[f24375])).
fof(f24375,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X17) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f24215,f271])).
fof(f25609,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f24369])).
fof(f24369,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f24215,f86])).
fof(f25632,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X32) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X34) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X32) = s(bool,X34)) )),
  inference(duplicate_literal_removal,[],[f24311])).
fof(f24311,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X32) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,X32) = s(bool,X34)) )),
  inference(superposition,[],[f24215,f24215])).
fof(f699,plain,(
  ( ! [X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f21937,plain,(
  ( ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6)))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl182),
  inference(cnf_transformation,[],[f21937_D])).
fof(f21937_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6)))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f21907,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X6)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f5808,f88])).
fof(f5808,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))),s('type.realax.real',X35)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36)))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f2825])).
fof(f2825,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))),s('type.realax.real',X35)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X36)))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f115,f647])).
fof(f17833,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl181),
  inference(cnf_transformation,[],[f17833_D])).
fof(f17833_D,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f17857,plain,(
  ( ! [X113] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X113)) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f17489,f15074])).
fof(f17489,plain,(
  ( ! [X113] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X113)))))))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15074,f7237])).
fof(f7237,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7227,f81])).
fof(f7227,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f7210,f121])).
fof(f7210,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f7186,f81])).
fof(f7186,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f115,f7170])).
fof(f7170,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))),
  inference(subsumption_resolution,[],[f7145,f81])).
fof(f7145,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))),
  inference(superposition,[],[f4255,f129])).
fof(f4255,plain,(
  ( ! [X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X16)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X16)))))) )),
  inference(subsumption_resolution,[],[f4232,f81])).
fof(f4232,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X16)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X16)))))) )),
  inference(superposition,[],[f705,f129])).
fof(f705,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(subsumption_resolution,[],[f697,f84])).
fof(f697,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f110,f86])).
fof(f17856,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X112)) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f17488,f15076])).
fof(f17488,plain,(
  ( ! [X112] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X112)))))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f15076,f7237])).
fof(f17831,plain,(
  ( ! [X100] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X100)) ) | ($spl15 | $spl171)),
  inference(subsumption_resolution,[],[f17476,f14938])).
fof(f17476,plain,(
  ( ! [X100] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X100)))))) ) | ($spl15 | $spl171)),
  inference(superposition,[],[f14938,f7237])).
fof(f16909,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X5))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl178),
  inference(cnf_transformation,[],[f16909_D])).
fof(f16909_D,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X5))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f16863,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X5))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4)))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f5805,f88])).
fof(f5805,plain,(
  ( ! [X26,X25] : (s('type.realax.real',X26) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X25))))),s('type.realax.real',X25))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26)))) | s('type.realax.real',X25) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f2820])).
fof(f2820,plain,(
  ( ! [X26,X25] : (s('type.realax.real',X26) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X25))))),s('type.realax.real',X25))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X26)))) | s('type.realax.real',X25) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f647])).
fof(f15976,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl174),
  inference(cnf_transformation,[],[f15976_D])).
fof(f15976_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f15978,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl176),
  inference(cnf_transformation,[],[f15978_D])).
fof(f15978_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f15979,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(subsumption_resolution,[],[f15917,f81])).
fof(f15917,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(superposition,[],[f7585,f1280])).
fof(f1280,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s('type.realax.real',X4) = s('type.realax.real',X5)) )),
  inference(subsumption_resolution,[],[f1264,f84])).
fof(f1264,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s('type.realax.real',X4) = s('type.realax.real',X5) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f465,f86])).
fof(f7585,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl28 | $spl33)),
  inference(forward_demodulation,[],[f7584,f85])).
fof(f7584,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(subsumption_resolution,[],[f7583,f84])).
fof(f7583,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(forward_demodulation,[],[f7582,f213])).
fof(f213,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f84,f209,f111])).
fof(f209,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(equality_resolution,[],[f109])).
fof(f7582,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(forward_demodulation,[],[f7581,f85])).
fof(f7581,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(subsumption_resolution,[],[f7519,f252])).
fof(f252,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f83,f109])).
fof(f7519,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl17 | $spl28 | $spl33)),
  inference(superposition,[],[f702,f5813])).
fof(f5813,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f2839])).
fof(f2839,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f107,f647])).
fof(f15974,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(subsumption_resolution,[],[f15916,f81])).
fof(f15916,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl28 | $spl33)),
  inference(superposition,[],[f7585,f1280])).
fof(f10074,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl168),
  inference(cnf_transformation,[],[f10074_D])).
fof(f10074_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f10071,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X3) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl45),
  inference(superposition,[],[f2636,f121])).
fof(f10073,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl166),
  inference(cnf_transformation,[],[f10073_D])).
fof(f10073_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10077,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl172),
  inference(cnf_transformation,[],[f10077_D])).
fof(f10077_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f10070,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl45),
  inference(superposition,[],[f2636,f121])).
fof(f10069,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl45),
  inference(superposition,[],[f2636,f121])).
fof(f10068,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl45),
  inference(superposition,[],[f2636,f121])).
fof(f2642,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f2642_D])).
fof(f2642_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f10067,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl45),
  inference(superposition,[],[f2636,f86])).
fof(f9520,plain,(
  $false | ($spl34 | $spl98)),
  inference(subsumption_resolution,[],[f9519,f1643])).
fof(f9519,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl98),
  inference(subsumption_resolution,[],[f9500,f81])).
fof(f9500,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl98),
  inference(superposition,[],[f115,f4213])).
fof(f9302,plain,(
  $false | ($spl17 | $spl33 | $spl102 | $spl115 | $spl158)),
  inference(subsumption_resolution,[],[f9301,f9300])).
fof(f9300,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl115 | $spl158)),
  inference(forward_demodulation,[],[f6365,f9295])).
fof(f9295,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl33 | $spl158)),
  inference(forward_demodulation,[],[f6593,f8994])).
fof(f6593,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f107,f6048])).
fof(f6048,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f107,f5824])).
fof(f5824,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f5363])).
fof(f5363,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f455,f647])).
fof(f9301,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl102 | $spl158)),
  inference(forward_demodulation,[],[f4437,f9295])).
fof(f9292,plain,(
  $false | ($spl11 | $spl17 | $spl96 | $spl134)),
  inference(subsumption_resolution,[],[f9291,f4976])).
fof(f9291,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl17 | $spl96 | $spl134)),
  inference(backward_demodulation,[],[f9282,f8889])).
fof(f8889,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17)),
  inference(backward_demodulation,[],[f4900,f8532])).
fof(f8532,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f5746,f4976])).
fof(f5746,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(forward_demodulation,[],[f5745,f1234])).
fof(f5745,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(forward_demodulation,[],[f5715,f85])).
fof(f5715,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl17),
  inference(superposition,[],[f591,f4900])).
fof(f591,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f85,f86])).
fof(f9282,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl96 | $spl134)),
  inference(backward_demodulation,[],[f4211,f7647])).
fof(f9290,plain,(
  $false | ($spl11 | $spl96 | $spl123)),
  inference(subsumption_resolution,[],[f9289,f4976])).
fof(f9289,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl96 | $spl123)),
  inference(forward_demodulation,[],[f9280,f256])).
fof(f256,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f81,f83,f114])).
fof(f9280,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl96 | $spl123)),
  inference(backward_demodulation,[],[f4211,f7621])).
fof(f7621,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl123),
  inference(unit_resulting_resolution,[],[f84,f7592,f111])).
fof(f9287,plain,(
  $false | ($spl96 | $spl98 | $spl100 | $spl103)),
  inference(subsumption_resolution,[],[f9253,f9022])).
fof(f9022,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl100 | $spl103)),
  inference(backward_demodulation,[],[f4435,f4436])).
fof(f4436,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl103),
  inference(cnf_transformation,[],[f4436_D])).
fof(f4436_D,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f9253,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl96 | $spl98)),
  inference(backward_demodulation,[],[f4211,f4213])).
fof(f9244,plain,(
  $false | ($spl34 | $spl100)),
  inference(subsumption_resolution,[],[f9243,f1643])).
fof(f9243,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl100),
  inference(forward_demodulation,[],[f9228,f85])).
fof(f9228,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl100),
  inference(unit_resulting_resolution,[],[f4435,f109])).
fof(f1644,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl35),
  inference(cnf_transformation,[],[f1644_D])).
fof(f1644_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f9017,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | $spl164),
  inference(cnf_transformation,[],[f9017_D])).
fof(f9017_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f9016,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f6716,f6593])).
fof(f6716,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f701,f6048])).
fof(f701,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(forward_demodulation,[],[f694,f85])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(superposition,[],[f110,f107])).
fof(f9011,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))),s('type.realax.real',X3))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | $spl162),
  inference(cnf_transformation,[],[f9011_D])).
fof(f9011_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))),s('type.realax.real',X3))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7516,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))),s('type.realax.real',X3))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f2761,f5813])).
fof(f2761,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(backward_demodulation,[],[f2724,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(forward_demodulation,[],[f692,f85])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))) )),
  inference(superposition,[],[f110,f107])).
fof(f2724,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f590,f107])).
fof(f590,plain,(
  ( ! [X0,X1] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f85,f107])).
fof(f9010,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))),s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl160),
  inference(cnf_transformation,[],[f9010_D])).
fof(f9010_D,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))),s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f7515,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f701,f5813])).
fof(f1643,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl34),
  inference(cnf_transformation,[],[f1643_D])).
fof(f1643_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4211,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl96),
  inference(cnf_transformation,[],[f4211_D])).
fof(f4211_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9009,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl18 | $spl33)),
  inference(forward_demodulation,[],[f9008,f8805])).
fof(f8805,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl18 | $spl33)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f83,f209,f214,f232,f270,f302,f366,f454,f459,f470,f471,f472,f271,f252,f251,f129,f89,f579,f646,f693,f695,f831,f888,f1115,f947,f1280,f1502,f1552,f1515,f698,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1248,f1234,f1224,f782,f710,f708,f706,f426,f256,f88,f644,f643,f2463,f2526,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2778,f2777,f2796,f2801,f2812,f2882,f2841,f2896,f2898,f2860,f2861,f3442,f3678,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4400,f759,f2214,f1179,f2263,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f2238,f2236,f2223,f2221,f2259,f458,f652,f659,f1255,f3883,f4168,f653,f657,f658,f1253,f3884,f2816,f2825,f2828,f2846,f2850,f2856,f2817,f2821,f2824,f2827,f2845,f2857,f4166,f4398,f4382,f2807,f3581,f4309,f2779,f2839,f2863,f2862,f2834,f2833,f647,f2563,f4313,f3351,f656,f2849,f1518,f2820,f6841,f6846,f6825,f6966,f7210,f7366,f84,f621,f2603,f7237,f7170,f7208,f7364,f7442,f269,f1504,f590,f2724,f2892,f2761,f701,f110,f594,f2794,f6828,f7155,f107,f3354,f2744,f2743,f2735,f2734,f2731,f7165,f7365,f7209,f7171,f4410,f4409,f4338,f4337,f4256,f4255,f4243,f4241,f705,f7167,f7143,f4254,f4250,f1524,f1373,f1262,f846,f813,f592,f1561,f1500,f1116,f465,f108,f889,f948,f2633,f7168,f4253,f2613,f2580,f2530,f622,f591,f142,f213,f236,f2577,f2579,f2612,f2875,f7222,f7426,f109,f115,f85,f2765,f5381])).
fof(f5381,plain,(
  ( ! [X8] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8))))),s('type.realax.real',X8))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f89,f647])).
fof(f2765,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(forward_demodulation,[],[f2740,f85])).
fof(f2740,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) )),
  inference(superposition,[],[f110,f590])).
fof(f7426,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f84,f7366,f111])).
fof(f7222,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f84,f7210,f111])).
fof(f2875,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))),s('type.realax.real',X9))))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2810,f1234])).
fof(f2810,plain,(
  ( ! [X9] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))),s('type.realax.real',X9))))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f85,f647])).
fof(f2612,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,sK1(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f2593,f81])).
fof(f2593,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X14) = s(bool,sK1(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f114,f621])).
fof(f2579,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f2578,f85])).
fof(f2578,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | s(bool,X13) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f2545,f81])).
fof(f2545,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | s(bool,X13) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f114,f591])).
fof(f2577,plain,(
  ( ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f2541,f81])).
fof(f2541,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f115,f591])).
fof(f2530,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f591,f107])).
fof(f2580,plain,(
  ( ! [X28,X26,X27] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X27),s('type.realax.real',X28))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real',X27))) | s('type.realax.real',X27) = s('type.realax.real',X28) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f2553,f81])).
fof(f2553,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X27),s('type.realax.real',X28))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real',X27))) | s('type.realax.real',X27) = s('type.realax.real',X28) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1500,f591])).
fof(f2613,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,sK1(X45,X46,X47,bool)) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real',X48))) | s('type.realax.real',X48) = s('type.realax.real',X49) | s(fun(bool,X47),X45) = s(fun(bool,X47),X46) | s(X47,f(s(bool,'F_0'))) != s(X47,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f2601,f81])).
fof(f2601,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,sK1(X45,X46,X47,bool)) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real',X48))) | s('type.realax.real',X48) = s('type.realax.real',X49) | s(fun(bool,X47),X45) = s(fun(bool,X47),X46) | s(X47,f(s(bool,'F_0'))) != s(X47,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1500,f621])).
fof(f4253,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f4252,f590])).
fof(f4252,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))),s('type.realax.real',X13))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f4230,f81])).
fof(f4230,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))),s('type.realax.real',X13))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f705,f591])).
fof(f7168,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f7142,f81])).
fof(f7142,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4255,f591])).
fof(f2633,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f2626])).
fof(f2626,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f622,f86])).
fof(f948,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,'T_0') != s(bool,f(s(X54,sK1(X55,X56,bool,X54)))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,X57) = s(bool,g(s(X54,sK1(X55,X56,bool,X54)))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f116,f121])).
fof(f889,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X63) != s(bool,f(s(X60,sK1(X61,X62,bool,X60)))) | s(fun(X60,bool),X61) = s(fun(X60,bool),X62) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,g(s(X60,sK1(X61,X62,bool,X60))))) )),
  inference(superposition,[],[f116,f121])).
fof(f1116,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X33)))) | s(bool,X35) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real',X34))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f939,f81])).
fof(f939,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X33)))) | s(bool,X35) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real',X34))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f115,f121])).
fof(f1500,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X24)) | s(bool,X24) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',X23))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X22))) | s('type.realax.real',X22) = s('type.realax.real',X23)) )),
  inference(subsumption_resolution,[],[f1412,f81])).
fof(f1412,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X24) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',X23))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X22))) | s('type.realax.real',X22) = s('type.realax.real',X23)) )),
  inference(superposition,[],[f114,f1280])).
fof(f1561,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',X40))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real',X42))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real',X41))) | s('type.realax.real',X41) = s('type.realax.real',X42) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',X39))) | s('type.realax.real',X39) = s('type.realax.real',X40)) )),
  inference(subsumption_resolution,[],[f1542,f81])).
fof(f1542,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',X40))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real',X42))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real',X41))) | s('type.realax.real',X41) = s('type.realax.real',X42) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',X39))) | s('type.realax.real',X39) = s('type.realax.real',X40)) )),
  inference(superposition,[],[f1500,f1280])).
fof(f813,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X29),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29)))))) )),
  inference(superposition,[],[f121,f85])).
fof(f846,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X23) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f121,f85])).
fof(f1262,plain,(
  ( ! [X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f465,f85])).
fof(f1373,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f1280,f85])).
fof(f1524,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1508,f81])).
fof(f1508,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f594,f121])).
fof(f4250,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f4224,f81])).
fof(f4224,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f705,f1280])).
fof(f4254,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15))))) )),
  inference(forward_demodulation,[],[f4231,f590])).
fof(f4231,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X15)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))),s('type.realax.real',X15)))))) )),
  inference(superposition,[],[f705,f85])).
fof(f7143,plain,(
  ( ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))) )),
  inference(superposition,[],[f4255,f85])).
fof(f7167,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f7136,f81])).
fof(f7136,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f4255,f1280])).
fof(f4241,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(forward_demodulation,[],[f4240,f590])).
fof(f4240,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real',X2)))))) )),
  inference(forward_demodulation,[],[f4219,f590])).
fof(f4219,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f705,f107])).
fof(f4243,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f4242,f81])).
fof(f4242,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4220,f256])).
fof(f4220,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f705,f88])).
fof(f4256,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X20)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) | s(bool,X21) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X19))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4234,f81])).
fof(f4234,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X20)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X19),s('type.realax.real',X20))))) | s(bool,X21) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X19))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f705,f121])).
fof(f4337,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f4319,f81])).
fof(f4319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f2761,f1280])).
fof(f4338,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real',X11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) | s(bool,X12) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4322,f81])).
fof(f4322,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real',X11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) | s(bool,X12) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2761,f121])).
fof(f4409,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f4388,f81])).
fof(f4388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f701,f1280])).
fof(f4410,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X10)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) | s(bool,X12) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4391,f81])).
fof(f4391,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X10)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) | s(bool,X12) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f701,f121])).
fof(f7171,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X8))))) | s(bool,X9) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f7147,f81])).
fof(f7147,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X8))))) | s(bool,X9) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4255,f121])).
fof(f7209,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real',X0)))))) )),
  inference(subsumption_resolution,[],[f7182,f81])).
fof(f7182,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))),s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f705,f7170])).
fof(f7365,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real',X0)))))) )),
  inference(subsumption_resolution,[],[f7338,f81])).
fof(f7338,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f705,f7208])).
fof(f7165,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f7164,f107])).
fof(f7164,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))))))) )),
  inference(forward_demodulation,[],[f7133,f590])).
fof(f7133,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X2)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))))))) )),
  inference(superposition,[],[f4255,f107])).
fof(f2731,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f590,f1280])).
fof(f2734,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X10) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f590,f121])).
fof(f2735,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real',X12)))))) )),
  inference(superposition,[],[f590,f121])).
fof(f2743,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f465,f590])).
fof(f2744,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) )),
  inference(superposition,[],[f108,f590])).
fof(f3354,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f3346])).
fof(f3346,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f656,f88])).
fof(f7155,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f7154,f81])).
fof(f7154,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f7127,f256])).
fof(f7127,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f4255,f88])).
fof(f6828,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f6827,f84])).
fof(f6827,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6826,f213])).
fof(f6826,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f6777,f3354])).
fof(f6777,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f652,f88])).
fof(f2794,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f647,f647])).
fof(f594,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f115,f85])).
fof(f1504,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f594,f107])).
fof(f7442,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7431,f81])).
fof(f7431,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f7366,f121])).
fof(f7364,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))))))),
  inference(subsumption_resolution,[],[f7337,f81])).
fof(f7337,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))))))),
  inference(superposition,[],[f4255,f7208])).
fof(f7208,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))))),
  inference(subsumption_resolution,[],[f7181,f81])).
fof(f7181,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0)))))))),
  inference(superposition,[],[f4255,f7170])).
fof(f2603,plain,(
  ( ! [X52,X50,X51] : (s(X52,f(s(bool,'T_0'))) != s(X52,g(s(bool,'T_0'))) | s(fun(bool,X52),X50) = s(fun(bool,X52),X51) | s(X52,f(s(bool,'F_0'))) != s(X52,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f2602])).
fof(f2602,plain,(
  ( ! [X52,X50,X51] : (s(X52,f(s(bool,'T_0'))) != s(X52,g(s(bool,'T_0'))) | s(fun(bool,X52),X50) = s(fun(bool,X52),X51) | s(fun(bool,X52),X50) = s(fun(bool,X52),X51) | s(X52,f(s(bool,'F_0'))) != s(X52,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f116,f621])).
fof(f7366,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(subsumption_resolution,[],[f7342,f81])).
fof(f7342,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f115,f7208])).
fof(f6966,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6)))))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5)))) )),
  inference(subsumption_resolution,[],[f6907,f84])).
fof(f6907,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6)))))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f653,f86])).
fof(f6825,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(duplicate_literal_removal,[],[f6803])).
fof(f6803,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f652,f88])).
fof(f6846,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6796,f710])).
fof(f6796,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f652,f88])).
fof(f6841,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X6)))) | s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f6788,f84])).
fof(f6788,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X6)))) | s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f652,f86])).
fof(f2849,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X18))))),s('type.realax.real',X18))),s('type.realax.real',X19)))) | s('type.realax.real',X18) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f647])).
fof(f656,plain,(
  ( ! [X15,X16] : (s('type.realax.real',X16) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15))),s('type.realax.real',X15))),s('type.realax.real',X16)))) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f88])).
fof(f3351,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f656])).
fof(f4313,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2761,f88])).
fof(f2563,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real',X1)))))) )),
  inference(forward_demodulation,[],[f2562,f107])).
fof(f2562,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real',X1))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f2525,f107])).
fof(f2525,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f591,f107])).
fof(f2833,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X49))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f647])).
fof(f2834,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.realax.real',X50))),s('type.realax.real',X51))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f647])).
fof(f2862,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))),s('type.realax.real',X42))) | s('type.realax.real',X42) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))) | s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f647])).
fof(f2863,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X44))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))))) | s('type.realax.real',X44) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f647])).
fof(f2779,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f252,f647])).
fof(f4309,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4))),s('type.realax.real',X5))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2761,f88])).
fof(f3581,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f657])).
fof(f2807,plain,(
  ( ! [X6,X5] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X5))),s('type.realax.real',X6))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f590,f647])).
fof(f4382,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f701,f88])).
fof(f4398,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))),s('type.realax.real',X5)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4378,f107])).
fof(f4378,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4))))),s('type.realax.real',X5)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f701,f88])).
fof(f4166,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X2)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f4165,f81])).
fof(f4165,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X2)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f4153,f656])).
fof(f4153,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X2)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f659,f1280])).
fof(f2857,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f647])).
fof(f2845,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f647])).
fof(f2827,plain,(
  ( ! [X39,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))),s('type.realax.real',X38))),s('type.realax.real',X39)))) | s('type.realax.real',X39) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f647])).
fof(f2824,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f115,f647])).
fof(f2821,plain,(
  ( ! [X28,X27] : (s('type.realax.real',X28) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))))),s('type.realax.real',X27))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X27) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f647])).
fof(f2856,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X32)))) | s('type.realax.real',X32) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f647])).
fof(f2850,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X20))))),s('type.realax.real',X20)))))) | s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f647])).
fof(f2846,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X13)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12)))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f647])).
fof(f2828,plain,(
  ( ! [X41,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40)))))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f647])).
fof(f2816,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))))),s('type.realax.real',X17)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f647])).
fof(f3884,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8))),s('type.realax.real',X8)))))) | s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X9) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f3867,f81])).
fof(f3867,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8))),s('type.realax.real',X8)))))) | s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X9) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f658,f121])).
fof(f1253,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f88])).
fof(f658,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X20)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))),s('type.realax.real',X19)))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f115,f88])).
fof(f657,plain,(
  ( ! [X17,X18] : (s('type.realax.real',X18) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.realax.real',X17)))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f109,f88])).
fof(f653,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X10)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9)))))) | s('type.realax.real',X10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f88])).
fof(f4168,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))),s('type.realax.real',X8)))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X10) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f4156,f81])).
fof(f4156,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))),s('type.realax.real',X8)))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X10) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f659,f121])).
fof(f3883,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f3882,f81])).
fof(f3882,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f3864,f657])).
fof(f3864,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f658,f1280])).
fof(f1255,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f88])).
fof(f659,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X21))),s('type.realax.real',X21))),s('type.realax.real',X22)))) | s('type.realax.real',X21) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f115,f88])).
fof(f652,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f108,f88])).
fof(f2221,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f2220])).
fof(f2220,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2196])).
fof(f2196,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1280,f269])).
fof(f2223,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f2222])).
fof(f2222,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2195])).
fof(f2195,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1280,f269])).
fof(f2236,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f2235])).
fof(f2235,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2159])).
fof(f2159,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(superposition,[],[f269,f1280])).
fof(f2238,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f2237])).
fof(f2237,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2158])).
fof(f2158,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(superposition,[],[f269,f1280])).
fof(f4865,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(subsumption_resolution,[],[f1920,f84])).
fof(f1920,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f465,f142])).
fof(f1455,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f1454])).
fof(f1454,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f1425])).
fof(f1425,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(superposition,[],[f269,f1280])).
fof(f1468,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f1467])).
fof(f1467,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f1401])).
fof(f1401,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(superposition,[],[f269,f1280])).
fof(f1477,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f1476])).
fof(f1476,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f1388])).
fof(f1388,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1280,f269])).
fof(f1485,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f1484])).
fof(f1484,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) )),
  inference(duplicate_literal_removal,[],[f1371])).
fof(f1371,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1280,f269])).
fof(f4843,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(subsumption_resolution,[],[f1261,f84])).
fof(f1261,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f465,f142])).
fof(f1279,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(condensation,[],[f1278])).
fof(f1278,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1260,f123])).
fof(f1260,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f465,f269])).
fof(f2263,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2199,f123])).
fof(f2199,plain,(
  ( ! [X6] : (p(s(bool,X6)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f108,f269])).
fof(f4400,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4399,f4089])).
fof(f4399,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real',X7)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl18),
  inference(forward_demodulation,[],[f4379,f107])).
fof(f4379,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real',X7)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl18),
  inference(superposition,[],[f701,f643])).
fof(f4135,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X44))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))))) | s('type.realax.real',X44) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2863])).
fof(f4134,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))),s('type.realax.real',X42))) | s('type.realax.real',X42) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))) | s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2862])).
fof(f4133,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2857])).
fof(f4132,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X32)))) | s('type.realax.real',X32) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2856])).
fof(f4131,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X20))))),s('type.realax.real',X20)))))) | s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2850])).
fof(f4130,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X18))))),s('type.realax.real',X18))),s('type.realax.real',X19)))) | s('type.realax.real',X18) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2849])).
fof(f4129,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12)))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2846])).
fof(f4128,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2845])).
fof(f4127,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2839])).
fof(f4144,plain,(
  ( ! [X50,X51] : (s('type.realax.real',X51) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.realax.real',X50))),s('type.realax.real',X51))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4126,f4087])).
fof(f4126,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.realax.real',X50))),s('type.realax.real',X51))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2834])).
fof(f4143,plain,(
  ( ! [X48,X49] : (s('type.realax.real',X49) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4125,f4087])).
fof(f4125,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2833])).
fof(f4142,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40)))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4124,f4087])).
fof(f4124,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40)))))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2828])).
fof(f4141,plain,(
  ( ! [X39,X38] : (s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))),s('type.realax.real',X38))),s('type.realax.real',X39)))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4123,f4087])).
fof(f4123,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))),s('type.realax.real',X38))),s('type.realax.real',X39)))) | s('type.realax.real',X39) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2827])).
fof(f4122,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))),s('type.realax.real',X35)))))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2825])).
fof(f4121,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2824])).
fof(f4120,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X28) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))))),s('type.realax.real',X27))) | s('type.realax.real',X27) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2821])).
fof(f4119,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26)))) | s('type.realax.real',X26) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X25))))),s('type.realax.real',X25))) | s('type.realax.real',X25) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2820])).
fof(f4140,plain,(
  ( ! [X19,X20] : (s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4118,f4087])).
fof(f4118,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2817])).
fof(f4139,plain,(
  ( ! [X17,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))))),s('type.realax.real',X17)))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4117,f4087])).
fof(f4117,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))))),s('type.realax.real',X17)))))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2816])).
fof(f4111,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2788])).
fof(f4110,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2787])).
fof(f4109,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2786])).
fof(f4108,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2785])).
fof(f4107,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2784])).
fof(f4106,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2783])).
fof(f4105,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2782])).
fof(f4104,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2779])).
fof(f4138,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4093,f4087])).
fof(f4093,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X14))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2521])).
fof(f4091,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f2460])).
fof(f4090,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f1702])).
fof(f4089,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4087,f1284])).
fof(f4087,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl18 | $spl33)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1518,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f2463,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2794,f2796,f2801,f2807,f2875,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f647,f1284,f2863,f2862,f2857,f2856,f2850,f2849,f2846,f2845,f2839,f2834,f2833,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2779,f2521,f2460,f1702,f4075,f4071])).
fof(f4071,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(forward_demodulation,[],[f4000,f1702])).
fof(f1702,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(superposition,[],[f107,f644])).
fof(f2460,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl18),
  inference(superposition,[],[f107,f1284])).
fof(f2521,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl18),
  inference(forward_demodulation,[],[f2480,f1284])).
fof(f2480,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real',X14))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl18),
  inference(superposition,[],[f1280,f1284])).
fof(f2782,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f623,f647])).
fof(f2783,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f707,f647])).
fof(f2784,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f709,f647])).
fof(f2785,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f783,f647])).
fof(f2786,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1225,f647])).
fof(f2787,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1359,f647])).
fof(f2788,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1528,f647])).
fof(f1284,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl18),
  inference(superposition,[],[f107,f643])).
fof(f4053,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X44))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))))) | s('type.realax.real',X44) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2863])).
fof(f4052,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X42),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))),s('type.realax.real',X42))) | s('type.realax.real',X42) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X41))))),s('type.realax.real',X41))) | s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2862])).
fof(f4051,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2857])).
fof(f4050,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X32)))) | s('type.realax.real',X32) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2856])).
fof(f4049,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X20))))),s('type.realax.real',X20)))))) | s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2850])).
fof(f4048,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X18))))),s('type.realax.real',X18))),s('type.realax.real',X19)))) | s('type.realax.real',X18) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2849])).
fof(f4047,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12)))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2846])).
fof(f4046,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2845])).
fof(f4045,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2839])).
fof(f4063,plain,(
  ( ! [X50,X51] : (s('type.realax.real',X51) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.realax.real',X50))),s('type.realax.real',X51))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4044,f4008])).
fof(f4044,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.realax.real',X50))),s('type.realax.real',X51))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2834])).
fof(f4062,plain,(
  ( ! [X48,X49] : (s('type.realax.real',X49) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4043,f4008])).
fof(f4043,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2833])).
fof(f4061,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40)))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4042,f4008])).
fof(f4042,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X41))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40)))))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2828])).
fof(f4060,plain,(
  ( ! [X39,X38] : (s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))),s('type.realax.real',X38))),s('type.realax.real',X39)))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4041,f4008])).
fof(f4041,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))),s('type.realax.real',X38))),s('type.realax.real',X39)))) | s('type.realax.real',X39) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2827])).
fof(f4040,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))),s('type.realax.real',X35)))))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2825])).
fof(f4039,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2824])).
fof(f4038,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X28) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))))),s('type.realax.real',X27))) | s('type.realax.real',X27) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2821])).
fof(f4037,plain,(
  ( ! [X26,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26)))) | s('type.realax.real',X26) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X25))))),s('type.realax.real',X25))) | s('type.realax.real',X25) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2820])).
fof(f4059,plain,(
  ( ! [X19,X20] : (s('type.realax.real',X20) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4036,f4008])).
fof(f4036,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2817])).
fof(f4058,plain,(
  ( ! [X17,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))))),s('type.realax.real',X17)))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4035,f4008])).
fof(f4035,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))))),s('type.realax.real',X17)))))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2816])).
fof(f4029,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2788])).
fof(f4028,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2787])).
fof(f4027,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2786])).
fof(f4026,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2785])).
fof(f4025,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2784])).
fof(f4024,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2783])).
fof(f4023,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2782])).
fof(f4022,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2779])).
fof(f4057,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f4011,f4008])).
fof(f4011,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X14))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2521])).
fof(f4009,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl33)),
  inference(backward_demodulation,[],[f4008,f2460])).
fof(f4008,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl18 | $spl33)),
  inference(forward_demodulation,[],[f1702,f4001])).
fof(f4006,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl33)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1702,f1527,f1518,f1490,f1366,f1363,f1358,f1284,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f647,f426,f256,f88,f644,f643,f2460,f2463,f2521,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2779,f2794,f2796,f2801,f2807,f2875,f2812,f2816,f2817,f2820,f2821,f2824,f2825,f2827,f2828,f2882,f2833,f2834,f2892,f2839,f2841,f2845,f2846,f2849,f2850,f2896,f2898,f2856,f2857,f2860,f2861,f2862,f2863,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005])).
fof(f4005,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(forward_demodulation,[],[f4004,f107])).
fof(f4004,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(subsumption_resolution,[],[f4003,f84])).
fof(f4003,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(forward_demodulation,[],[f4002,f1234])).
fof(f4002,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(forward_demodulation,[],[f3101,f643])).
fof(f3101,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(superposition,[],[f592,f643])).
fof(f4001,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl33)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1702,f1527,f1518,f1490,f1366,f1363,f1358,f1284,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f647,f426,f256,f88,f644,f643,f2460,f2463,f2521,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2779,f2794,f2796,f2801,f2807,f2875,f2812,f2816,f2817,f2820,f2821,f2824,f2825,f2827,f2828,f2882,f2833,f2834,f2892,f2839,f2841,f2845,f2846,f2849,f2850,f2896,f2898,f2856,f2857,f2860,f2861,f2862,f2863,f3354,f3442,f3678,f3883,f3884,f4000])).
fof(f4000,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(forward_demodulation,[],[f3999,f107])).
fof(f3999,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(subsumption_resolution,[],[f3998,f84])).
fof(f3998,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(forward_demodulation,[],[f3997,f1234])).
fof(f3997,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(forward_demodulation,[],[f3102,f644])).
fof(f3102,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(superposition,[],[f592,f644])).
fof(f3678,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f3583,f643])).
fof(f3583,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f3576])).
fof(f3576,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(superposition,[],[f657,f643])).
fof(f3442,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f3353,f643])).
fof(f3353,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f3347])).
fof(f3347,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(superposition,[],[f656,f643])).
fof(f2777,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1528,f88])).
fof(f2778,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1528,f89])).
fof(f643,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl18),
  inference(unit_resulting_resolution,[],[f579,f88])).
fof(f644,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f608,f88])).
fof(f426,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f252,f89])).
fof(f706,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f623,f88])).
fof(f708,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f707,f88])).
fof(f710,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f252,f88])).
fof(f782,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f709,f88])).
fof(f1224,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f783,f88])).
fof(f1248,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1237,f81])).
fof(f1237,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f251,f121])).
fof(f1358,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1225,f88])).
fof(f1363,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f88])).
fof(f1366,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1280,f88])).
fof(f1490,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl18),
  inference(forward_demodulation,[],[f1364,f643])).
fof(f1364,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real',X5) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl18),
  inference(superposition,[],[f1280,f643])).
fof(f1527,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1359,f88])).
fof(f1749,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) ) | $spl18),
  inference(forward_demodulation,[],[f1722,f644])).
fof(f1722,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real',X14))) | s('type.realax.real',X14) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) ) | $spl18),
  inference(superposition,[],[f1280,f644])).
fof(f623,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f608,f89])).
fof(f707,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f623,f89])).
fof(f709,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f707,f89])).
fof(f783,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f709,f89])).
fof(f1225,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f783,f89])).
fof(f1359,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1225,f89])).
fof(f1528,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | $spl18),
  inference(unit_resulting_resolution,[],[f1359,f89])).
fof(f583,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl18),
  inference(unit_resulting_resolution,[],[f81,f458,f114])).
fof(f608,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl18),
  inference(unit_resulting_resolution,[],[f579,f89])).
fof(f1515,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f1514,f107])).
fof(f1514,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f1501,f107])).
fof(f1501,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f594,f107])).
fof(f1552,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real',X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1532,f81])).
fof(f1532,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real',X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1500,f121])).
fof(f947,plain,(
  ( ! [X52,X50,X53,X51] : (s(X52,f(s(bool,'T_0'))) != s(X52,g(s(bool,'T_0'))) | s(fun(bool,X52),X50) = s(fun(bool,X52),X51) | s(bool,X53) = s(bool,sK1(X50,X51,X52,bool)) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f116,f121])).
fof(f888,plain,(
  ( ! [X59,X57,X58,X56] : (s(X58,f(s(bool,X59))) != s(X58,g(s(bool,X59))) | s(fun(bool,X58),X56) = s(fun(bool,X58),X57) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,sK1(X56,X57,X58,bool))) )),
  inference(superposition,[],[f116,f121])).
fof(f831,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f121,f121])).
fof(f579,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl18),
  inference(unit_resulting_resolution,[],[f458,f109])).
fof(f89,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),
  file('thm.real.REAL_INV_POS','thm.real.REAL_INV_NZ')).
fof(f470,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f442,f123])).
fof(f442,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f269])).
fof(f270,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f267,f84])).
fof(f267,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f157,f213])).
fof(f214,plain,(
  ( ! [X0,X1] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f209,f209,f111])).
fof(f9008,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33)),
  inference(subsumption_resolution,[],[f9007,f84])).
fof(f9007,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33)),
  inference(forward_demodulation,[],[f9006,f213])).
fof(f9006,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33)),
  inference(forward_demodulation,[],[f5929,f8805])).
fof(f5929,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f592,f5783])).
fof(f9005,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f9004,f6048])).
fof(f9004,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f9003,f107])).
fof(f9003,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f9002,f84])).
fof(f9002,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f9001,f213])).
fof(f9001,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f6070,f5824])).
fof(f6070,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f592,f5824])).
fof(f9000,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f8999,f6458])).
fof(f6458,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f107,f5825])).
fof(f5825,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f5364])).
fof(f5364,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f4901,f647])).
fof(f4901,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f455,f89])).
fof(f8999,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f8998,f107])).
fof(f8998,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f8997,f84])).
fof(f8997,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f8996,f213])).
fof(f8996,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f6480,f5825])).
fof(f6480,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f592,f5825])).
fof(f8994,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl158),
  inference(cnf_transformation,[],[f8994_D])).
fof(f8994_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f8992,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl141)),
  inference(forward_demodulation,[],[f8991,f107])).
fof(f8991,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33 | $spl141)),
  inference(forward_demodulation,[],[f8990,f6048])).
fof(f8990,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33 | $spl141)),
  inference(subsumption_resolution,[],[f6615,f8518])).
fof(f6615,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f592,f6048])).
fof(f8979,plain,(
  $false | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98 | $spl103)),
  inference(subsumption_resolution,[],[f8978,f8920])).
fof(f8920,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl103)),
  inference(forward_demodulation,[],[f8909,f8908])).
fof(f8908,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(forward_demodulation,[],[f8907,f8805])).
fof(f8907,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(subsumption_resolution,[],[f8906,f84])).
fof(f8906,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(forward_demodulation,[],[f8905,f213])).
fof(f8905,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(forward_demodulation,[],[f8904,f8805])).
fof(f8904,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(subsumption_resolution,[],[f5929,f8862])).
fof(f8862,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl33 | $spl85)),
  inference(forward_demodulation,[],[f2905,f8805])).
fof(f8909,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl103)),
  inference(backward_demodulation,[],[f8908,f4436])).
fof(f8978,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98)),
  inference(forward_demodulation,[],[f4213,f8908])).
fof(f8926,plain,(
  $false | ($spl11 | $spl17 | $spl18 | $spl33 | $spl85 | $spl134)),
  inference(subsumption_resolution,[],[f8925,f4976])).
fof(f8925,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl17 | $spl18 | $spl33 | $spl85 | $spl134)),
  inference(backward_demodulation,[],[f8917,f8889])).
fof(f8917,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl134)),
  inference(backward_demodulation,[],[f8908,f7647])).
fof(f8924,plain,(
  $false | ($spl11 | $spl17 | $spl18 | $spl33 | $spl85 | $spl123)),
  inference(subsumption_resolution,[],[f8923,f4976])).
fof(f8923,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl17 | $spl18 | $spl33 | $spl85 | $spl123)),
  inference(forward_demodulation,[],[f8915,f256])).
fof(f8915,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl123)),
  inference(backward_demodulation,[],[f8908,f7621])).
fof(f8901,plain,(
  ( ! [X43] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))))),s('type.realax.real',X43))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl154),
  inference(cnf_transformation,[],[f8901_D])).
fof(f8901_D,plain,(
  ( ! [X43] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))))),s('type.realax.real',X43))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f8903,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl156),
  inference(cnf_transformation,[],[f8903_D])).
fof(f8903_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f8858,plain,(
  ( ! [X43] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))))),s('type.realax.real',X43))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f7540,f85])).
fof(f7540,plain,(
  ( ! [X43] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))))),s('type.realax.real',X43))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f591,f5813])).
fof(f8058,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl136),
  inference(cnf_transformation,[],[f8058_D])).
fof(f8058_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f8898,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl152),
  inference(cnf_transformation,[],[f8898_D])).
fof(f8898_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6595,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f88,f6048])).
fof(f8896,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl151),
  inference(cnf_transformation,[],[f8896_D])).
fof(f8896_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f6596,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f89,f6048])).
fof(f8894,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl134 | $spl136)),
  inference(subsumption_resolution,[],[f8893,f4976])).
fof(f8893,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl11 | $spl17 | $spl33 | $spl134 | $spl136)),
  inference(backward_demodulation,[],[f8892,f8889])).
fof(f8892,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl134 | $spl136)),
  inference(forward_demodulation,[],[f7647,f8811])).
fof(f8811,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl136)),
  inference(forward_demodulation,[],[f6048,f8058])).
fof(f8891,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl123 | $spl136)),
  inference(subsumption_resolution,[],[f8890,f4976])).
fof(f8890,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl17 | $spl33 | $spl123 | $spl136)),
  inference(backward_demodulation,[],[f256,f8887])).
fof(f8887,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl17 | $spl33 | $spl123 | $spl136)),
  inference(forward_demodulation,[],[f7621,f8811])).
fof(f8879,plain,(
  $false | ($spl17 | $spl33 | $spl98 | $spl103 | $spl136)),
  inference(subsumption_resolution,[],[f8878,f8813])).
fof(f8813,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(backward_demodulation,[],[f8811,f4213])).
fof(f8878,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl103 | $spl136)),
  inference(forward_demodulation,[],[f8877,f8811])).
fof(f8877,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl33 | $spl103 | $spl136)),
  inference(forward_demodulation,[],[f4436,f8811])).
fof(f8872,plain,(
  $false | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(subsumption_resolution,[],[f8871,f81])).
fof(f8871,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(forward_demodulation,[],[f8870,f8813])).
fof(f8870,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(forward_demodulation,[],[f8869,f8058])).
fof(f8869,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(subsumption_resolution,[],[f8868,f81])).
fof(f8868,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl98 | $spl136)),
  inference(forward_demodulation,[],[f8867,f8813])).
fof(f8867,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl136)),
  inference(forward_demodulation,[],[f8866,f8811])).
fof(f8866,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl136)),
  inference(forward_demodulation,[],[f6616,f8811])).
fof(f6616,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f594,f6048])).
fof(f8865,plain,(
  $false | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98 | $spl136)),
  inference(subsumption_resolution,[],[f8864,f81])).
fof(f8864,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98 | $spl136)),
  inference(forward_demodulation,[],[f8863,f8813])).
fof(f8863,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl136)),
  inference(forward_demodulation,[],[f8862,f8811])).
fof(f8820,plain,(
  $false | ($spl17 | $spl33 | $spl98 | $spl136 | $spl141)),
  inference(subsumption_resolution,[],[f8819,f81])).
fof(f8819,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl33 | $spl98 | $spl136 | $spl141)),
  inference(forward_demodulation,[],[f8818,f8813])).
fof(f8818,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl136 | $spl141)),
  inference(forward_demodulation,[],[f8815,f8811])).
fof(f8815,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl136 | $spl141)),
  inference(backward_demodulation,[],[f8811,f8518])).
fof(f8817,plain,(
  $false | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98 | $spl136)),
  inference(subsumption_resolution,[],[f8816,f81])).
fof(f8816,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl98 | $spl136)),
  inference(forward_demodulation,[],[f8814,f8813])).
fof(f8814,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl18 | $spl33 | $spl85 | $spl136)),
  inference(backward_demodulation,[],[f8811,f8810])).
fof(f8810,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl18 | $spl33 | $spl85)),
  inference(global_subsumption,[],[f6716,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f83,f209,f214,f232,f270,f302,f366,f454,f459,f470,f471,f472,f271,f252,f251,f129,f89,f579,f646,f693,f695,f831,f888,f1115,f947,f1280,f1502,f1552,f1515,f698,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1248,f1234,f1224,f782,f710,f708,f706,f426,f256,f88,f644,f643,f2463,f2526,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2778,f2777,f2796,f2801,f2812,f2882,f2841,f2896,f2898,f2860,f2861,f3442,f3678,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f4056,f1284,f2905,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4137,f4195,f4209,f4214,f4215,f4400,f759,f2214,f1179,f2263,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f2238,f2236,f2223,f2221,f2259,f458,f652,f659,f1255,f3883,f4168,f653,f657,f658,f1253,f3884,f2816,f2825,f2828,f2846,f2850,f2856,f2817,f2821,f2824,f2827,f2845,f2857,f4166,f4398,f4382,f2807,f3581,f4309,f2779,f2839,f2863,f2862,f2834,f2833,f647,f2563,f4313,f3351,f656,f2849,f1518,f2820,f6841,f6846,f6825,f6966,f7210,f7366,f84,f621,f2603,f7237,f7170,f7208,f7364,f7442,f269,f1504,f590,f2724,f2892,f2761,f701,f110,f594,f2794,f6828,f7155,f107,f3354,f2744,f2743,f2735,f2734,f2731,f7165,f7365,f7209,f7171,f4410,f4409,f4338,f4337,f4256,f4255,f4243,f4241,f705,f7167,f7143,f4254,f4250,f1524,f1373,f1262,f846,f813,f592,f1561,f1500,f1116,f465,f108,f889,f948,f2633,f7168,f4253,f2613,f2580,f2530,f622,f591,f142,f213,f236,f2577,f2579,f2612,f2875,f7222,f7426,f109,f115,f85,f2765,f5381,f8805])).
fof(f4215,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl33 | $spl85)),
  inference(subsumption_resolution,[],[f4200,f81])).
fof(f4200,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl33 | $spl85)),
  inference(superposition,[],[f4137,f121])).
fof(f4195,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl85)),
  inference(unit_resulting_resolution,[],[f84,f4137,f111])).
fof(f4137,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl33 | $spl85)),
  inference(backward_demodulation,[],[f4087,f2905])).
fof(f4056,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl33 | $spl85)),
  inference(backward_demodulation,[],[f4008,f2905])).
fof(f8792,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl98 | $spl103 | $spl130)),
  inference(subsumption_resolution,[],[f8791,f8523])).
fof(f8523,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl103 | $spl130)),
  inference(backward_demodulation,[],[f7690,f4436])).
fof(f7690,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl33 | $spl130)),
  inference(backward_demodulation,[],[f7600,f5783])).
fof(f8791,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl98 | $spl130)),
  inference(forward_demodulation,[],[f4213,f8689])).
fof(f8689,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f8688,f5824])).
fof(f8688,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(subsumption_resolution,[],[f8687,f84])).
fof(f8687,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f8686,f8537])).
fof(f8537,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f8536,f7690])).
fof(f8536,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f8535,f7600])).
fof(f8535,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33)),
  inference(subsumption_resolution,[],[f5965,f4976])).
fof(f5965,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5964,f213])).
fof(f5964,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5928,f85])).
fof(f5928,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f591,f5783])).
fof(f8686,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f8685,f5824])).
fof(f8685,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33 | $spl130)),
  inference(subsumption_resolution,[],[f8684,f84])).
fof(f8684,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f6070,f7661])).
fof(f7661,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl130),
  inference(backward_demodulation,[],[f7600,f1234])).
fof(f8765,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(subsumption_resolution,[],[f8764,f252])).
fof(f8764,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl11 | $spl17 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f7600,f8689])).
fof(f8642,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(subsumption_resolution,[],[f8583,f252])).
fof(f8583,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(backward_demodulation,[],[f8553,f7600])).
fof(f8553,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(forward_demodulation,[],[f8552,f5802])).
fof(f5802,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2793])).
fof(f2793,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2496,f647])).
fof(f8552,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(subsumption_resolution,[],[f8551,f84])).
fof(f8551,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(forward_demodulation,[],[f8550,f8537])).
fof(f8550,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(forward_demodulation,[],[f8549,f5802])).
fof(f8549,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(subsumption_resolution,[],[f8548,f84])).
fof(f8548,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl130)),
  inference(forward_demodulation,[],[f7025,f8533])).
fof(f8533,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl130)),
  inference(forward_demodulation,[],[f8532,f7682])).
fof(f7682,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl17 | $spl130)),
  inference(backward_demodulation,[],[f7600,f4900])).
fof(f7025,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(superposition,[],[f592,f5802])).
fof(f8640,plain,(
  $false | ($spl11 | $spl17 | $spl33 | $spl39 | $spl98 | $spl103 | $spl130)),
  inference(subsumption_resolution,[],[f8554,f8523])).
fof(f8554,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl11 | $spl17 | $spl33 | $spl39 | $spl98 | $spl130)),
  inference(backward_demodulation,[],[f8553,f4213])).
fof(f8522,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl148),
  inference(cnf_transformation,[],[f8522_D])).
fof(f8522_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f5846,plain,(
  ( ! [X4] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5845,f5783])).
fof(f5845,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f5844,f84])).
fof(f5844,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5843,f213])).
fof(f5843,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f5842,f5783])).
fof(f5842,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))),s('type.realax.real',X4))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f5761,f107])).
fof(f5761,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f592,f647])).
fof(f8521,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl146),
  inference(cnf_transformation,[],[f8521_D])).
fof(f8521_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f6799,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X4)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f652,f6048])).
fof(f8520,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4))))) ) | $spl144),
  inference(cnf_transformation,[],[f8520_D])).
fof(f8520_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f6925,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f653,f6048])).
fof(f8519,plain,(
  ( ! [X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X44) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44)))))) ) | $spl142),
  inference(cnf_transformation,[],[f8519_D])).
fof(f8519_D,plain,(
  ( ! [X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X44) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44)))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f8517,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl140),
  inference(cnf_transformation,[],[f8517_D])).
fof(f8517_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7541,plain,(
  ( ! [X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))))),s('type.realax.real',X44))))) | s('type.realax.real',X44) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f592,f5813])).
fof(f4086,plain,(
  ( ! [X47] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))))))),s('type.realax.real',X47))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f4086_D])).
fof(f4086_D,plain,(
  ( ! [X47] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))))))),s('type.realax.real',X47))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8518,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl141),
  inference(cnf_transformation,[],[f8518_D])).
fof(f8518_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f7542,plain,(
  ( ! [X45] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X45))))))),s('type.realax.real',X45))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f594,f5813])).
fof(f8516,plain,(
  $false | ($spl0 | $spl45)),
  inference(subsumption_resolution,[],[f8515,f131])).
fof(f8515,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl0 | $spl45)),
  inference(forward_demodulation,[],[f2636,f131])).
fof(f8514,plain,(
  $false | ($spl0 | $spl130)),
  inference(subsumption_resolution,[],[f7651,f8425])).
fof(f8425,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f8424,f84])).
fof(f8424,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl0),
  inference(forward_demodulation,[],[f8423,f131])).
fof(f8423,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f8422,f84])).
fof(f8422,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl0),
  inference(forward_demodulation,[],[f108,f131])).
fof(f7651,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl130),
  inference(backward_demodulation,[],[f7600,f426])).
fof(f8512,plain,(
  $false | ($spl0 | $spl86 | $spl130)),
  inference(subsumption_resolution,[],[f8511,f84])).
fof(f8511,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl86 | $spl130)),
  inference(forward_demodulation,[],[f7725,f131])).
fof(f7725,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl86 | $spl130)),
  inference(forward_demodulation,[],[f7669,f590])).
fof(f7669,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl86 | $spl130)),
  inference(backward_demodulation,[],[f7600,f2984])).
fof(f8469,plain,(
  $false | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f4976,f131])).
fof(f8468,plain,(
  $false | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f8467,f318])).
fof(f8467,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f4979,f131])).
fof(f4979,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl11),
  inference(superposition,[],[f318,f86])).
fof(f8466,plain,(
  $false | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f5089,f131])).
fof(f5089,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(superposition,[],[f387,f142])).
fof(f8465,plain,(
  $false | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f8464,f387])).
fof(f8464,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f5092,f131])).
fof(f5092,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl15),
  inference(superposition,[],[f387,f86])).
fof(f8463,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f5788,f8425])).
fof(f5788,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2582])).
fof(f2582,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2496,f89])).
fof(f8462,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f5790,f8425])).
fof(f5790,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2584])).
fof(f2584,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2582,f89])).
fof(f8461,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f5792,f8425])).
fof(f5792,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2615])).
fof(f2615,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2584,f89])).
fof(f8460,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f5794,f8425])).
fof(f5794,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2617])).
fof(f2617,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2615,f89])).
fof(f8459,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f5796,f8425])).
fof(f5796,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(backward_demodulation,[],[f5783,f2776])).
fof(f2776,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f2617,f89])).
fof(f8458,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39)),
  inference(subsumption_resolution,[],[f6199,f8425])).
fof(f6199,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))))))) | ($spl17 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f5796,f89])).
fof(f8457,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f6161,f8425])).
fof(f6161,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f5659,f89])).
fof(f5659,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f5361,f89])).
fof(f5361,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f5185,f89])).
fof(f5185,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f5165,f89])).
fof(f5165,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f4957,f89])).
fof(f4957,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f4901,f89])).
fof(f8456,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f7118,f8425])).
fof(f7118,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f6161,f89])).
fof(f8455,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5659,f8425])).
fof(f8454,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5361,f8425])).
fof(f8453,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5185,f8425])).
fof(f8452,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f5165,f8425])).
fof(f8451,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f4957,f8425])).
fof(f8450,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f4901,f8425])).
fof(f8393,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl85 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8392,f81])).
fof(f8392,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl85 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8391,f8131])).
fof(f8131,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f2678,f131])).
fof(f2678,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X4)) ) | $spl49),
  inference(forward_demodulation,[],[f2672,f2657])).
fof(f2657,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl49),
  inference(unit_resulting_resolution,[],[f84,f2639,f157])).
fof(f2672,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl49),
  inference(duplicate_literal_removal,[],[f2661])).
fof(f2661,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl49),
  inference(superposition,[],[f2639,f121])).
fof(f8391,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl85 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8390,f131])).
fof(f8390,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl85 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f2905,f8036])).
fof(f8036,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6363,f7690])).
fof(f8389,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8388,f81])).
fof(f8388,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8387,f8131])).
fof(f8387,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8386,f131])).
fof(f8386,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8385,f8036])).
fof(f8385,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8384,f81])).
fof(f8384,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8383,f8131])).
fof(f8383,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8382,f131])).
fof(f8382,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f5930,f8036])).
fof(f5930,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f594,f5783])).
fof(f8381,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8380,f81])).
fof(f8380,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8379,f8131])).
fof(f8379,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8378,f131])).
fof(f8378,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8377,f8036])).
fof(f8377,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8376,f81])).
fof(f8376,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8375,f8131])).
fof(f8375,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8374,f131])).
fof(f8374,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6071,f8036])).
fof(f6071,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f594,f5824])).
fof(f8373,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8372,f81])).
fof(f8372,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8371,f8131])).
fof(f8371,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8370,f131])).
fof(f8370,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8369,f8036])).
fof(f8369,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8368,f81])).
fof(f8368,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8367,f8131])).
fof(f8367,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8366,f131])).
fof(f8366,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6481,f8036])).
fof(f6481,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33)),
  inference(superposition,[],[f594,f5825])).
fof(f8358,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8357,f81])).
fof(f8357,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8356,f8131])).
fof(f8356,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8355,f131])).
fof(f8355,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8354,f8036])).
fof(f8354,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8353,f81])).
fof(f8353,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8352,f8131])).
fof(f8352,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8351,f131])).
fof(f8351,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6616,f8036])).
fof(f8343,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8342,f81])).
fof(f8342,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8341,f8131])).
fof(f8341,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8340,f131])).
fof(f8340,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8339,f8036])).
fof(f8339,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8338,f81])).
fof(f8338,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8337,f8131])).
fof(f8337,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8336,f131])).
fof(f8336,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl39 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7026,f8036])).
fof(f7026,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl39)),
  inference(superposition,[],[f594,f5802])).
fof(f8335,plain,(
  $false | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f8334,f81])).
fof(f8334,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl49)),
  inference(forward_demodulation,[],[f84,f8131])).
fof(f8324,plain,(
  $false | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f8323,f81])).
fof(f8323,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl49)),
  inference(forward_demodulation,[],[f8322,f8131])).
fof(f8322,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl49)),
  inference(forward_demodulation,[],[f8321,f131])).
fof(f8321,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f8320,f81])).
fof(f8320,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | ($spl0 | $spl49)),
  inference(forward_demodulation,[],[f8319,f8131])).
fof(f8319,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl0),
  inference(forward_demodulation,[],[f115,f131])).
fof(f8317,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8316,f81])).
fof(f8316,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8315,f8131])).
fof(f8315,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8314,f131])).
fof(f8314,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(condensation,[],[f8313])).
fof(f8313,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f594,f8036])).
fof(f8282,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8281,f81])).
fof(f8281,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8280,f8131])).
fof(f8280,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8279,f131])).
fof(f8279,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(condensation,[],[f8278])).
fof(f8278,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f2892,f8036])).
fof(f8277,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8276,f81])).
fof(f8276,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8275,f8131])).
fof(f8275,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8274,f131])).
fof(f8274,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(condensation,[],[f8273])).
fof(f8273,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7580,f8036])).
fof(f7580,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f7579,f85])).
fof(f7579,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))))))))) ) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f7578,f84])).
fof(f7578,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))))))))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f7577,f213])).
fof(f7577,plain,(
  ( ! [X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))))))))) ) | ($spl17 | $spl33)),
  inference(forward_demodulation,[],[f7576,f85])).
fof(f7576,plain,(
  ( ! [X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))))))))) ) | ($spl17 | $spl33)),
  inference(subsumption_resolution,[],[f7517,f109])).
fof(f7517,plain,(
  ( ! [X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5)))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f110,f5813])).
fof(f8260,plain,(
  $false | ($spl0 | $spl17 | $spl28 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8259,f81])).
fof(f8259,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl28 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8258,f8131])).
fof(f8258,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl28 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8257,f131])).
fof(f8257,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl28 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7688,f8036])).
fof(f7688,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl28 | $spl130)),
  inference(backward_demodulation,[],[f7600,f5689])).
fof(f5689,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl28),
  inference(forward_demodulation,[],[f5688,f85])).
fof(f5688,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | $spl28),
  inference(subsumption_resolution,[],[f5687,f84])).
fof(f5687,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | $spl28),
  inference(forward_demodulation,[],[f5686,f1234])).
fof(f5686,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | $spl28),
  inference(forward_demodulation,[],[f5685,f85])).
fof(f5685,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | $spl28),
  inference(subsumption_resolution,[],[f5665,f252])).
fof(f5665,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl28),
  inference(superposition,[],[f702,f647])).
fof(f8256,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8255,f81])).
fof(f8255,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8254,f8131])).
fof(f8254,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8253,f131])).
fof(f8253,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8252,f8036])).
fof(f8252,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8251,f81])).
fof(f8251,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8250,f8131])).
fof(f8250,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8249,f131])).
fof(f8249,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f1504,f8036])).
fof(f8239,plain,(
  $false | ($spl0 | $spl17 | $spl28 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8238,f81])).
fof(f8238,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl28 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8237,f8131])).
fof(f8237,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl28 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8236,f131])).
fof(f8236,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl28 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7717,f8036])).
fof(f7717,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl28 | $spl33 | $spl130)),
  inference(backward_demodulation,[],[f7600,f7585])).
fof(f8235,plain,(
  $false | ($spl0 | $spl15 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8219,f8131])).
fof(f8219,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl15 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f387])).
fof(f8139,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8138,f81])).
fof(f8138,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8137,f8131])).
fof(f8137,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8136,f131])).
fof(f8136,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7675,f8036])).
fof(f7675,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl130),
  inference(backward_demodulation,[],[f7600,f3581])).
fof(f8234,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8188,f8139])).
fof(f8188,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f7118])).
fof(f8233,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8169,f8139])).
fof(f8169,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f6199])).
fof(f8232,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8163,f8139])).
fof(f8163,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f5796])).
fof(f8231,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8162,f8139])).
fof(f8162,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f5794])).
fof(f8230,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8161,f8139])).
fof(f8161,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f5792])).
fof(f8229,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8160,f8139])).
fof(f8160,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl17 | $spl33 | $spl39 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8139,f5790])).
fof(f8135,plain,(
  $false | ($spl0 | $spl15 | $spl49)),
  inference(subsumption_resolution,[],[f5089,f8131])).
fof(f8134,plain,(
  $false | ($spl0 | $spl11 | $spl49)),
  inference(subsumption_resolution,[],[f4976,f8131])).
fof(f8133,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8132,f81])).
fof(f8132,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl17 | $spl33 | $spl49 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8131,f8042])).
fof(f8042,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(global_subsumption,[],[f8041,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1515,f698,f1366,f1363,f1248,f1234,f710,f426,f256,f88,f2526,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2796,f2801,f2812,f2882,f2841,f2896,f2898,f2860,f2861,f4075,f759,f2214,f1179,f2263,f2580,f1373,f846,f813,f591,f85,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f2238,f2236,f2223,f2221,f2259,f142,f269,f652,f659,f1255,f3883,f4168,f653,f657,f658,f1253,f3884,f2816,f2825,f2828,f2846,f2850,f2856,f2817,f2821,f2824,f2827,f2845,f2857,f4166,f4409,f4337,f4256,f4250,f110,f705,f4410,f701,f4398,f4382,f2724,f2807,f4243,f590,f4255,f3581,f2892,f4309,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f2744,f2743,f2563,f1504,f2530,f2731,f2734,f2735,f2761,f2765,f4241,f4253,f4254,f4338,f4313,f3351,f656,f2849,f3354,f2577,f592,f594,f1262,f1524,f2820,f6828,f6841,f6846,f6825,f6966,f7155,f7165,f7167,f7168,f7143,f7170,f7171,f7208,f7209,f7210,f7222,f7237,f7364,f7365,f7366,f7426,f7442])).
fof(f8041,plain,(
  ( ! [X45] : (~p(s(bool,'F_0')) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(duplicate_literal_removal,[],[f8040])).
fof(f8040,plain,(
  ( ! [X45] : (~p(s(bool,'F_0')) | ~p(s(bool,'F_0')) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8039,f131])).
fof(f8039,plain,(
  ( ! [X45] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'F_0')) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8038,f8036])).
fof(f8038,plain,(
  ( ! [X45] : (~p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X45))))))),s('type.realax.real',X45))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8037,f131])).
fof(f8037,plain,(
  ( ! [X45] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X45))))))),s('type.realax.real',X45))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X45) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f7542,f8036])).
fof(f8130,plain,(
  $false | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f2674,f131])).
fof(f2674,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl49),
  inference(backward_demodulation,[],[f2657,f2639])).
fof(f8117,plain,(
  $false | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8089,f252])).
fof(f8089,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(backward_demodulation,[],[f8072,f7600])).
fof(f8072,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8071,f8042])).
fof(f8071,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8070,f131])).
fof(f8070,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(condensation,[],[f8069])).
fof(f8069,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f5846,f8036])).
fof(f8059,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl138),
  inference(cnf_transformation,[],[f8059_D])).
fof(f8059_D,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f8067,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8066,f7660])).
fof(f7660,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl130),
  inference(backward_demodulation,[],[f7600,f710])).
fof(f8066,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8065,f6048])).
fof(f8065,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8064,f8042])).
fof(f8064,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8063,f131])).
fof(f8063,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8062,f8036])).
fof(f8062,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8061,f8042])).
fof(f8061,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8060,f131])).
fof(f8060,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6799,f8036])).
fof(f8056,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8055,f7660])).
fof(f8055,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8054,f6048])).
fof(f8054,plain,(
  ( ! [X4] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8053,f8042])).
fof(f8053,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8052,f131])).
fof(f8052,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8051,f8036])).
fof(f8051,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(subsumption_resolution,[],[f8050,f8042])).
fof(f8050,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl0 | $spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f8049,f131])).
fof(f8049,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | ($spl17 | $spl33 | $spl112 | $spl130)),
  inference(forward_demodulation,[],[f6925,f8036])).
fof(f8035,plain,(
  $false | ($spl6 | $spl49 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f2674,f7773])).
fof(f7773,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7748,f123])).
fof(f7748,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,X1))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7722,f156])).
fof(f7722,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl102 | $spl130)),
  inference(forward_demodulation,[],[f7661,f5186])).
fof(f5186,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl102),
  inference(superposition,[],[f4437,f85])).
fof(f8034,plain,(
  $false | ($spl6 | $spl11 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f4976,f7773])).
fof(f8033,plain,(
  $false | ($spl6 | $spl15 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f5089,f7773])).
fof(f8030,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8029,f7901])).
fof(f7901,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1)) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7777,f81])).
fof(f7777,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) != s('type.realax.real',X1)) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f109])).
fof(f8029,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7881,f81])).
fof(f7881,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7675])).
fof(f8028,plain,(
  $false | ($spl6 | $spl17 | $spl28 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7880,f81])).
fof(f7880,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl28 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7717])).
fof(f8027,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8026,f7901])).
fof(f8026,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7879,f81])).
fof(f7879,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7674])).
fof(f7674,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl130),
  inference(backward_demodulation,[],[f7600,f3351])).
fof(f8025,plain,(
  $false | ($spl6 | $spl102 | $spl128 | $spl130)),
  inference(subsumption_resolution,[],[f8024,f7901])).
fof(f8024,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl128 | $spl130)),
  inference(subsumption_resolution,[],[f7883,f81])).
fof(f7883,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | ($spl6 | $spl102 | $spl128 | $spl130)),
  inference(condensation,[],[f7877])).
fof(f7877,plain,(
  ( ! [X59,X58] : (~p(s(bool,'T_0')) | s('type.realax.real',X58) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X59) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl128 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7599])).
fof(f8023,plain,(
  $false | ($spl6 | $spl102 | $spl130 | $spl132)),
  inference(subsumption_resolution,[],[f8022,f7901])).
fof(f8022,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130 | $spl132)),
  inference(subsumption_resolution,[],[f7884,f81])).
fof(f7884,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | ($spl6 | $spl102 | $spl130 | $spl132)),
  inference(condensation,[],[f7875])).
fof(f7875,plain,(
  ( ! [X61,X60] : (~p(s(bool,'T_0')) | s('type.realax.real',X60) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130 | $spl132)),
  inference(backward_demodulation,[],[f7773,f7602])).
fof(f8021,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8020,f7901])).
fof(f8020,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7885,f81])).
fof(f7885,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | ($spl6 | $spl102 | $spl130)),
  inference(condensation,[],[f7874])).
fof(f7874,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f3354])).
fof(f8015,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8014,f81])).
fof(f8014,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f8013,f7773])).
fof(f8013,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7867,f81])).
fof(f7867,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f1504])).
fof(f8012,plain,(
  $false | ($spl6 | $spl82 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8011,f7901])).
fof(f8011,plain,(
  ( ! [X40] : (s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl82 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7865,f81])).
fof(f7865,plain,(
  ( ! [X40] : (~p(s(bool,'T_0')) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl82 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f2903])).
fof(f8010,plain,(
  $false | ($spl6 | $spl30 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f8009,f7901])).
fof(f8009,plain,(
  ( ! [X2] : (s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl30 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7864,f81])).
fof(f7864,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl30 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f1516])).
fof(f8008,plain,(
  $false | ($spl6 | $spl28 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7859,f81])).
fof(f7859,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl28 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7688])).
fof(f7977,plain,(
  $false | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7976,f81])).
fof(f7976,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7975,f7773])).
fof(f7975,plain,(
  ( ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7841,f81])).
fof(f7841,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7580])).
fof(f7948,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7947,f81])).
fof(f7947,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7946,f7773])).
fof(f7946,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7822,f81])).
fof(f7822,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f2892])).
fof(f7935,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7934,f81])).
fof(f7934,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7933,f7773])).
fof(f7933,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X21))),s('type.realax.real',X21))),s('type.realax.real',X22))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7932,f7901])).
fof(f7932,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X21))),s('type.realax.real',X21))),s('type.realax.real',X22)))) | s('type.realax.real',X21) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7814,f81])).
fof(f7814,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X21))),s('type.realax.real',X21))),s('type.realax.real',X22)))) | s('type.realax.real',X21) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7657])).
fof(f7657,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X21))),s('type.realax.real',X21))),s('type.realax.real',X22)))) | s('type.realax.real',X21) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl130),
  inference(backward_demodulation,[],[f7600,f659])).
fof(f7931,plain,(
  $false | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7930,f81])).
fof(f7930,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7929,f7773])).
fof(f7929,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X32))))))),s('type.realax.real',X32))),s('type.realax.real',X33))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7928,f7901])).
fof(f7928,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X32))))))),s('type.realax.real',X32))),s('type.realax.real',X33)))) | s('type.realax.real',X32) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7811,f81])).
fof(f7811,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X32))))))),s('type.realax.real',X32))),s('type.realax.real',X33)))) | s('type.realax.real',X32) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7534])).
fof(f7534,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X32))))))),s('type.realax.real',X32))),s('type.realax.real',X33)))) | s('type.realax.real',X32) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f115,f5813])).
fof(f7927,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7926,f81])).
fof(f7926,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7925,f7773])).
fof(f7925,plain,(
  ( ! [X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X20))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7924,f7901])).
fof(f7924,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X20)))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7806,f81])).
fof(f7806,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X20)))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7656])).
fof(f7656,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X20)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))),s('type.realax.real',X19)))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl130),
  inference(backward_demodulation,[],[f7600,f658])).
fof(f7923,plain,(
  $false | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7922,f81])).
fof(f7922,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7921,f7773])).
fof(f7921,plain,(
  ( ! [X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X31))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7920,f7901])).
fof(f7920,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X31)))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7801,f81])).
fof(f7801,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X31)))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f7533])).
fof(f7533,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X31)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X30))))))),s('type.realax.real',X30)))))) | s('type.realax.real',X30) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f115,f5813])).
fof(f7919,plain,(
  $false | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7918,f81])).
fof(f7918,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7917,f7773])).
fof(f7917,plain,(
  ( ! [X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7916,f7901])).
fof(f7916,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36)))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7796,f81])).
fof(f7796,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X36)))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f5808])).
fof(f7913,plain,(
  $false | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7912,f81])).
fof(f7912,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7911,f7773])).
fof(f7911,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7910,f7901])).
fof(f7910,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7790,f81])).
fof(f7790,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl6 | $spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f5807])).
fof(f5807,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(backward_demodulation,[],[f5783,f2824])).
fof(f7907,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7906,f81])).
fof(f7906,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7905,f7773])).
fof(f7905,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7783,f81])).
fof(f7783,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f594])).
fof(f7904,plain,(
  $false | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7903,f81])).
fof(f7903,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl102 | $spl130)),
  inference(forward_demodulation,[],[f7902,f7773])).
fof(f7902,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7778,f81])).
fof(f7778,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | ($spl6 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7773,f115])).
fof(f7772,plain,(
  $false | ($spl15 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7747,f387])).
fof(f7747,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl102 | $spl130)),
  inference(backward_demodulation,[],[f7722,f142])).
fof(f7771,plain,(
  $false | ($spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7746,f81])).
fof(f7746,plain,(
  ~p(s(bool,'T_0')) | ($spl102 | $spl130)),
  inference(backward_demodulation,[],[f7722,f84])).
fof(f7736,plain,(
  $false | ($spl17 | $spl33 | $spl102 | $spl115 | $spl130)),
  inference(subsumption_resolution,[],[f7731,f7728])).
fof(f7728,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl102 | $spl130)),
  inference(backward_demodulation,[],[f7690,f4437])).
fof(f7731,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl17 | $spl33 | $spl115 | $spl130)),
  inference(backward_demodulation,[],[f7690,f6365])).
fof(f7723,plain,(
  $false | ($spl3 | $spl102 | $spl130)),
  inference(subsumption_resolution,[],[f7722,f133])).
fof(f7600,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl130),
  inference(cnf_transformation,[],[f7600_D])).
fof(f7600_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7647,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl134),
  inference(cnf_transformation,[],[f7647_D])).
fof(f7647_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7648,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl123),
  inference(subsumption_resolution,[],[f7623,f81])).
fof(f7623,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl123),
  inference(superposition,[],[f7592,f1280])).
fof(f7645,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl123),
  inference(subsumption_resolution,[],[f7624,f81])).
fof(f7624,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl123),
  inference(superposition,[],[f7592,f1280])).
fof(f7640,plain,(
  $false | ($spl123 | $spl127 | $spl131)),
  inference(subsumption_resolution,[],[f7639,f81])).
fof(f7639,plain,(
  ~p(s(bool,'T_0')) | ($spl123 | $spl127 | $spl131)),
  inference(backward_demodulation,[],[f7611,f7596])).
fof(f7611,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl123 | $spl131)),
  inference(unit_resulting_resolution,[],[f7601,f7592,f465])).
fof(f7638,plain,(
  $false | ($spl123 | $spl127 | $spl131)),
  inference(subsumption_resolution,[],[f7612,f7596])).
fof(f7612,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl123 | $spl131)),
  inference(unit_resulting_resolution,[],[f7601,f7592,f108])).
fof(f7637,plain,(
  $false | ($spl123 | $spl127 | $spl131)),
  inference(subsumption_resolution,[],[f7613,f7596])).
fof(f7613,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl123 | $spl131)),
  inference(unit_resulting_resolution,[],[f7601,f7592,f108])).
fof(f7599,plain,(
  ( ! [X59,X58] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X59))),s('type.realax.real',X59))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X58))))))),s('type.realax.real',X58)))))) | s('type.realax.real',X58) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X59) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f7599_D])).
fof(f7599_D,plain,(
  ( ! [X59,X58] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X59))),s('type.realax.real',X59))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X58))))))),s('type.realax.real',X58)))))) | s('type.realax.real',X58) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X59) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7592,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl123),
  inference(cnf_transformation,[],[f7592_D])).
fof(f7592_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f7552,plain,(
  ( ! [X64,X65] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X65))),s('type.realax.real',X65))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X64))))))),s('type.realax.real',X64)))))) | s('type.realax.real',X65) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X64) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f659,f5813])).
fof(f7602,plain,(
  ( ! [X61,X60] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X60))))))),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X61))),s('type.realax.real',X61)))))) | s('type.realax.real',X60) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl132),
  inference(cnf_transformation,[],[f7602_D])).
fof(f7602_D,plain,(
  ( ! [X61,X60] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X60))))))),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X61))),s('type.realax.real',X61)))))) | s('type.realax.real',X60) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7596,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl127),
  inference(cnf_transformation,[],[f7596_D])).
fof(f7596_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f7551,plain,(
  ( ! [X62,X63] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X62))))))),s('type.realax.real',X62))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X63))),s('type.realax.real',X63)))))) | s('type.realax.real',X63) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X62) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f658,f5813])).
fof(f7601,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl131),
  inference(cnf_transformation,[],[f7601_D])).
fof(f7601_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f7550,plain,(
  ( ! [X61,X60] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X60))))))),s('type.realax.real',X60))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X61))),s('type.realax.real',X61)))))) | s('type.realax.real',X61) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X60) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f657,f5813])).
fof(f7549,plain,(
  ( ! [X59,X58] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X59))),s('type.realax.real',X59))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X58))))))),s('type.realax.real',X58)))))) | s('type.realax.real',X59) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X58) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f656,f5813])).
fof(f7595,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53)))))) | s('type.realax.real',X52) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X53) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52)))) ) | $spl124),
  inference(cnf_transformation,[],[f7595_D])).
fof(f7595_D,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53)))))) | s('type.realax.real',X52) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X53) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f7597,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl126),
  inference(cnf_transformation,[],[f7597_D])).
fof(f7597_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7546,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53)))))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X53))),s('type.realax.real',X53))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X52))))))),s('type.realax.real',X52))) | s('type.realax.real',X53) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X52) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f653,f5813])).
fof(f7591,plain,(
  ( ! [X47,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46)))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46)))) ) | $spl120),
  inference(cnf_transformation,[],[f7591_D])).
fof(f7591_D,plain,(
  ( ! [X47,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46)))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7593,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl122),
  inference(cnf_transformation,[],[f7593_D])).
fof(f7593_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7543,plain,(
  ( ! [X47,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46)))))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))),s('type.realax.real',X47))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f652,f5813])).
fof(f7565,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl116),
  inference(cnf_transformation,[],[f7565_D])).
fof(f7565_D,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f7567,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl118),
  inference(cnf_transformation,[],[f7567_D])).
fof(f7567_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f7506,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f5813,f5813])).
fof(f3996,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl92),
  inference(cnf_transformation,[],[f3996_D])).
fof(f3996_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7371,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33 | $spl74)),
  inference(superposition,[],[f5942,f88])).
fof(f5942,plain,(
  ( ! [X23] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23))))) ) | ($spl17 | $spl33 | $spl74)),
  inference(backward_demodulation,[],[f5904,f2897])).
fof(f5904,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl17 | $spl33)),
  inference(superposition,[],[f590,f5783])).
fof(f3993,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl90),
  inference(cnf_transformation,[],[f3993_D])).
fof(f3993_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7240,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl17 | $spl33 | $spl76)),
  inference(superposition,[],[f5822,f88])).
fof(f5822,plain,(
  ( ! [X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))))) ) | ($spl17 | $spl33 | $spl76)),
  inference(backward_demodulation,[],[f5783,f2899])).
fof(f6363,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl112),
  inference(cnf_transformation,[],[f6363_D])).
fof(f6363_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f6365,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl115),
  inference(cnf_transformation,[],[f6365_D])).
fof(f6365_D,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f6295,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl102),
  inference(equality_factoring,[],[f5313])).
fof(f5313,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl102),
  inference(subsumption_resolution,[],[f5305,f81])).
fof(f5305,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl102),
  inference(superposition,[],[f5215,f121])).
fof(f5215,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl102),
  inference(subsumption_resolution,[],[f5195,f81])).
fof(f5195,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl102),
  inference(superposition,[],[f115,f4437])).
fof(f4885,plain,(
  $false | ($spl8 | $spl15 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f387,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269,f142,f4880])).
fof(f4880,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f2263,f4664])).
fof(f4876,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f2260,f4664])).
fof(f2260,plain,(
  ( ! [X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2198,f123])).
fof(f2198,plain,(
  ( ! [X5] : (p(s(bool,X5)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f108,f269])).
fof(f4869,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f1939,f4664])).
fof(f1939,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f1922,f84])).
fof(f1922,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f4867,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f1936,f4664])).
fof(f1936,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f1921,f84])).
fof(f1921,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f4840,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f1217,f4664])).
fof(f1217,plain,(
  ( ! [X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1165,f123])).
fof(f1165,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f108,f269])).
fof(f4838,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f1214,f4664])).
fof(f1214,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1164,f123])).
fof(f1164,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f108,f269])).
fof(f4836,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f636,f4664])).
fof(f636,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f625,f84])).
fof(f625,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f4834,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f633,f4664])).
fof(f633,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f624,f84])).
fof(f624,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f4832,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4831,f4664])).
fof(f4831,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4830,f316])).
fof(f4830,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4829,f4664])).
fof(f4829,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f652,f316])).
fof(f4826,plain,(
  ( ! [X4,X3] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4825,f4664])).
fof(f4825,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))),s('type.realax.real',X4))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f1255,f316])).
fof(f4821,plain,(
  ( ! [X10,X9] : (s('type.realax.real',X10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4820,f4664])).
fof(f4820,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4819,f316])).
fof(f4819,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9)))))) | s('type.realax.real',X10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4818,f4664])).
fof(f4818,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9)))))) | s('type.realax.real',X10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f653,f316])).
fof(f4814,plain,(
  ( ! [X4,X3] : (s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4813,f4664])).
fof(f4813,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X4) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f1253,f316])).
fof(f4810,plain,(
  ( ! [X17,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4809,f4664])).
fof(f4809,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4808,f316])).
fof(f4808,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4807,f4664])).
fof(f4807,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2816,f316])).
fof(f4804,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4803,f4704])).
fof(f4803,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X0)))) ) | ($spl8 | $spl28)),
  inference(condensation,[],[f4802])).
fof(f4802,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4801,f4704])).
fof(f4801,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4800,f4664])).
fof(f4800,plain,(
  ( ! [X41,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2828,f316])).
fof(f4798,plain,(
  ( ! [X12,X13] : (s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4797,f4664])).
fof(f4797,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4796,f316])).
fof(f4796,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12)))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4795,f4664])).
fof(f4795,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12)))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2846,f316])).
fof(f4791,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl8 | $spl28)),
  inference(condensation,[],[f4790])).
fof(f4790,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X32) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4789,f4675])).
fof(f4789,plain,(
  ( ! [X31,X32] : (s('type.realax.real',X32) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4788,f4675])).
fof(f4788,plain,(
  ( ! [X31,X32] : (s('type.realax.real',X32) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4787,f4664])).
fof(f4787,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X32) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X31))))) | s('type.realax.real',X31) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2856,f316])).
fof(f4785,plain,(
  ( ! [X19,X20] : (s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4784,f4664])).
fof(f4784,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4783,f316])).
fof(f4783,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4782,f4664])).
fof(f4782,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2817,f316])).
fof(f4777,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4776,f4704])).
fof(f4776,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) | ($spl8 | $spl28)),
  inference(condensation,[],[f4775])).
fof(f4775,plain,(
  ( ! [X39,X38] : (s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4774,f4704])).
fof(f4774,plain,(
  ( ! [X39,X38] : (s('type.realax.real',X39) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4773,f4664])).
fof(f4773,plain,(
  ( ! [X39,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X39) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X38) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2827,f316])).
fof(f4771,plain,(
  ( ! [X10,X11] : (s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4770,f4664])).
fof(f4770,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4769,f316])).
fof(f4769,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4768,f4664])).
fof(f4768,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real',X11) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X10))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2845,f316])).
fof(f4766,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | ($spl8 | $spl28)),
  inference(condensation,[],[f4765])).
fof(f4765,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X34))) | s('type.realax.real',X34) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4764,f4675])).
fof(f4764,plain,(
  ( ! [X33,X34] : (s('type.realax.real',X34) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4763,f4675])).
fof(f4763,plain,(
  ( ! [X33,X34] : (s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4762,f4664])).
fof(f4762,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X33))))),s('type.realax.real',X33))),s('type.realax.real',X34))) | s('type.realax.real',X33) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(forward_demodulation,[],[f2857,f316])).
fof(f4694,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) ) | ($spl8 | $spl28)),
  inference(backward_demodulation,[],[f4675,f107])).
fof(f4721,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4701,f4704])).
fof(f4701,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl8 | $spl28)),
  inference(condensation,[],[f4691])).
fof(f4691,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(backward_demodulation,[],[f4675,f2839])).
fof(f4720,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X44))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X44) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4686,f4704])).
fof(f4686,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X44))) | s('type.realax.real',X44) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X43))))),s('type.realax.real',X43))) | s('type.realax.real',X43) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(backward_demodulation,[],[f4675,f2863])).
fof(f4704,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl8 | $spl28)),
  inference(condensation,[],[f4676])).
fof(f4676,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(backward_demodulation,[],[f4675,f647])).
fof(f4675,plain,(
  ( ! [X8,X9] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4674,f4664])).
fof(f4674,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f2744,f4664])).
fof(f4673,plain,(
  ( ! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X7))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f2743,f4664])).
fof(f4664,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f4663,f84])).
fof(f4663,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4662,f236])).
fof(f4662,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f3015,f316])).
fof(f4650,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4649,f84])).
fof(f4649,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) ) | $spl8),
  inference(forward_demodulation,[],[f4648,f236])).
fof(f4648,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) ) | $spl8),
  inference(forward_demodulation,[],[f592,f316])).
fof(f3016,plain,(
  ( ! [X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f3004,f84])).
fof(f3004,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) | $spl28),
  inference(superposition,[],[f702,f86])).
fof(f3015,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl28),
  inference(forward_demodulation,[],[f3001,f85])).
fof(f3001,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) ) | $spl28),
  inference(superposition,[],[f702,f590])).
fof(f3012,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl28),
  inference(forward_demodulation,[],[f3011,f85])).
fof(f3011,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))) ) | $spl28),
  inference(forward_demodulation,[],[f2996,f2724])).
fof(f2996,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))) ) | $spl28),
  inference(superposition,[],[f702,f107])).
fof(f4884,plain,(
  $false | ($spl8 | $spl17 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f455,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269,f142,f4880])).
fof(f4883,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(subsumption_resolution,[],[f458,f4664])).
fof(f4882,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl102)),
  inference(global_subsumption,[],[f4437,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269,f142,f4880])).
fof(f4881,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269,f142,f4880])).
fof(f4879,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269,f142])).
fof(f4878,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876,f269])).
fof(f4877,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259,f4876])).
fof(f4875,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221,f2259])).
fof(f4874,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223,f2221])).
fof(f4873,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236,f2223])).
fof(f4872,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238,f2236])).
fof(f4871,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869,f2238])).
fof(f4870,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867,f4869])).
fof(f4868,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865,f4867])).
fof(f4866,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455,f4865])).
fof(f4864,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4863,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4863,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1919,f133])).
fof(f1919,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f1280,f142])).
fof(f4862,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4861,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4861,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1918,f133])).
fof(f1918,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f1280,f142])).
fof(f4860,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4859,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4859,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1912,f133])).
fof(f1912,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f142,f1280])).
fof(f4858,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4857,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4857,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1911,f133])).
fof(f1911,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f142,f1280])).
fof(f4856,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4855,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4855,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1427,f133])).
fof(f1427,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f142,f1280])).
fof(f4854,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468,f1455])).
fof(f4853,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4852,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468])).
fof(f4852,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1403,f133])).
fof(f1403,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f142,f1280])).
fof(f4851,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477,f1468])).
fof(f4850,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4849,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477])).
fof(f4849,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1389,f133])).
fof(f1389,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f1280,f142])).
fof(f4848,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485,f1477])).
fof(f4847,plain,(
  $false | ($spl3 | $spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f4846,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485])).
fof(f4846,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl3),
  inference(subsumption_resolution,[],[f1372,f133])).
fof(f1372,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f1280,f142])).
fof(f4845,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843,f1485])).
fof(f4844,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279,f4843])).
fof(f4842,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840,f1279])).
fof(f4841,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838,f4840])).
fof(f4839,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836,f4838])).
fof(f4837,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834,f4836])).
fof(f4835,plain,(
  $false | ($spl8 | $spl18 | $spl28)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4005,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832,f4834])).
fof(f4833,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826,f4832])).
fof(f2904,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl84),
  inference(cnf_transformation,[],[f2904_D])).
fof(f2904_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4439,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl33 | $spl84)),
  inference(subsumption_resolution,[],[f4424,f81])).
fof(f4424,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl33 | $spl84)),
  inference(superposition,[],[f4260,f121])).
fof(f4419,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f84,f4260,f111])).
fof(f4345,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4295,f88])).
fof(f4346,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4295,f89])).
fof(f4294,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4263,f88])).
fof(f4295,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4263,f89])).
fof(f4260,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4218,f594])).
fof(f4263,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f4218,f109])).
fof(f4268,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f81,f4218,f114])).
fof(f4218,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl33 | $spl84)),
  inference(forward_demodulation,[],[f2904,f4087])).
fof(f4827,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821,f4826])).
fof(f4822,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814,f4821])).
fof(f4815,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810,f4814])).
fof(f4811,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804,f4810])).
fof(f4805,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798,f4804])).
fof(f4799,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791,f4798])).
fof(f4792,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785,f4791])).
fof(f4786,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777,f4785])).
fof(f4778,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771,f4777])).
fof(f4772,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766,f4771])).
fof(f4767,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694,f4766])).
fof(f4734,plain,(
  $false | ($spl8 | $spl28 | $spl86)),
  inference(subsumption_resolution,[],[f4733,f4664])).
fof(f4733,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl86)),
  inference(forward_demodulation,[],[f2984,f316])).
fof(f4724,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721,f4694])).
fof(f4723,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl70 | $spl84)),
  inference(global_subsumption,[],[f4692,f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721])).
fof(f4692,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28 | $spl70)),
  inference(backward_demodulation,[],[f4675,f2885])).
fof(f4722,plain,(
  $false | ($spl8 | $spl18 | $spl28 | $spl33 | $spl84)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f271,f252,f251,f236,f129,f89,f458,f579,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1262,f594,f592,f1515,f1504,f705,f701,f698,f590,f110,f608,f583,f269,f142,f1528,f1359,f1225,f783,f709,f707,f623,f1749,f1527,f1490,f1366,f1363,f1358,f1255,f1253,f1248,f1234,f1224,f782,f710,f708,f706,f659,f658,f657,f656,f653,f652,f426,f256,f88,f644,f643,f702,f2463,f2563,f2526,f2530,f2577,f2579,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2778,f2777,f2796,f2801,f2807,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3012,f3015,f3016,f3354,f3442,f3678,f3883,f3884,f4000,f4001,f4005,f4006,f4008,f4009,f4057,f4022,f4023,f4024,f4025,f4026,f4027,f4028,f4029,f4058,f4059,f4037,f4038,f4039,f4040,f4060,f4061,f4062,f4063,f4045,f4046,f4047,f4048,f4049,f4050,f4051,f4052,f4053,f1284,f2788,f2787,f2786,f2785,f2784,f2783,f2782,f2521,f2460,f1702,f4075,f4071,f4087,f4089,f4090,f4091,f4138,f4104,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4139,f4140,f4119,f4120,f4121,f4122,f4141,f4142,f4143,f4144,f4127,f4128,f4129,f4130,f4131,f4132,f4133,f4134,f4135,f4166,f4168,f4218,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4268,f4263,f4260,f4295,f4294,f4309,f4313,f4337,f4338,f4346,f4345,f4398,f4400,f4382,f4409,f4410,f4419,f4433,f4438,f4439,f759,f2214,f1179,f2904,f1518,f2857,f2856,f2850,f2849,f2846,f2845,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f3581,f3351,f2263,f316,f2580,f1373,f846,f813,f591,f85,f2779,f107,f2839,f2875,f2863,f2862,f2834,f2833,f2794,f647,f4650,f4664,f4673,f4675,f4704,f4716,f4720,f4721])).
fof(f4548,plain,(
  ( ! [X17] : (s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f4548_D])).
fof(f4548_D,plain,(
  ( ! [X17] : (s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4620,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl108),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4689,plain,(
  ( ! [X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28 | $spl66)),
  inference(backward_demodulation,[],[f4675,f2876])).
fof(f4717,plain,(
  ( ! [X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl110),
  inference(cnf_transformation,[],[f4717_D])).
fof(f4717_D,plain,(
  ( ! [X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4716,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f4683,f4704])).
fof(f4683,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X49))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl28)),
  inference(backward_demodulation,[],[f4675,f2833])).
fof(f4642,plain,(
  $false | ($spl8 | $spl84)),
  inference(subsumption_resolution,[],[f4641,f84])).
fof(f4641,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl84)),
  inference(forward_demodulation,[],[f4640,f236])).
fof(f4640,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl84)),
  inference(forward_demodulation,[],[f2904,f316])).
fof(f4634,plain,(
  $false | ($spl12 | $spl16 | $spl28 | $spl84)),
  inference(subsumption_resolution,[],[f4526,f4622])).
fof(f4622,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl12 | $spl16 | $spl28)),
  inference(subsumption_resolution,[],[f4516,f4569])).
fof(f4569,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4568,f84])).
fof(f4568,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4533,f213])).
fof(f4533,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl12 | $spl16)),
  inference(condensation,[],[f4475])).
fof(f4475,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f110])).
fof(f4447,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f385,f4442])).
fof(f4442,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl16),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f456,f480,f501,f502,f515,f517,f505,f507,f508,f509,f512,f513,f271,f252,f251,f236,f129,f89,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f269,f142,f1366,f1363,f1255,f1253,f1248,f1234,f710,f659,f658,f657,f656,f653,f652,f426,f256,f88,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2794,f2796,f2801,f2807,f2875,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3354,f3883,f3884,f4075,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2863,f2862,f2857,f2856,f2850,f2849,f2846,f2845,f2839,f2834,f2833,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f2779,f647,f3581,f3351,f2263])).
fof(f513,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl16),
  inference(backward_demodulation,[],[f480,f456])).
fof(f512,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl16),
  inference(backward_demodulation,[],[f480,f426])).
fof(f509,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl16),
  inference(backward_demodulation,[],[f480,f256])).
fof(f508,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl16),
  inference(backward_demodulation,[],[f480,f252])).
fof(f507,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))) | $spl16),
  inference(backward_demodulation,[],[f480,f251])).
fof(f505,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))) | $spl16),
  inference(backward_demodulation,[],[f480,f129])).
fof(f517,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) ) | $spl16),
  inference(forward_demodulation,[],[f516,f480])).
fof(f516,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl16),
  inference(forward_demodulation,[],[f504,f480])).
fof(f504,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl16),
  inference(backward_demodulation,[],[f480,f110])).
fof(f515,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',sK0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))) ) | $spl16),
  inference(forward_demodulation,[],[f503,f480])).
fof(f503,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl16),
  inference(backward_demodulation,[],[f480,f89])).
fof(f502,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) ) | $spl16),
  inference(backward_demodulation,[],[f480,f88])).
fof(f501,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl16),
  inference(backward_demodulation,[],[f480,f85])).
fof(f480,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f456,f89])).
fof(f4516,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl12 | $spl16 | $spl28)),
  inference(backward_demodulation,[],[f4447,f3015])).
fof(f4526,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl12 | $spl16 | $spl84)),
  inference(backward_demodulation,[],[f4447,f2904])).
fof(f4589,plain,(
  ( ! [X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16 | $spl66)),
  inference(backward_demodulation,[],[f4575,f2876])).
fof(f4575,plain,(
  ( ! [X8,X9] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4574,f4569])).
fof(f4574,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4573,f4447])).
fof(f4573,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4499,f4569])).
fof(f4499,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real',X9)))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f2744])).
fof(f4549,plain,(
  ( ! [X18] : (s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl106),
  inference(cnf_transformation,[],[f4549_D])).
fof(f4549_D,plain,(
  ( ! [X18] : (s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4618,plain,(
  ( ! [X50,X51] : (s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | ($spl12 | $spl16)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f456,f480,f501,f502,f515,f517,f505,f507,f508,f509,f512,f513,f271,f252,f251,f236,f129,f89,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f269,f142,f1366,f1363,f1255,f1253,f1248,f1234,f710,f659,f658,f657,f656,f653,f652,f426,f256,f88,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2794,f2796,f2801,f2807,f2875,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3354,f3883,f3884,f4075,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2863,f2862,f2857,f2856,f2850,f2849,f2846,f2845,f2839,f2834,f2833,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f2779,f647,f3581,f3351,f2263,f4442,f4447,f4538,f4547,f4552,f4553,f4558,f4564,f4569,f4575,f4603,f4615,f4616,f4617])).
fof(f4617,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X51))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4584,f4603])).
fof(f4584,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X51))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X51) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X50) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4575,f2834])).
fof(f4615,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X49))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4583,f4603])).
fof(f4583,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X49),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X49))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4575,f2833])).
fof(f4603,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl12 | $spl16)),
  inference(condensation,[],[f4576])).
fof(f4576,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4575,f647])).
fof(f4558,plain,(
  ( ! [X12,X13] : (s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4557,f84])).
fof(f4557,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4556,f213])).
fof(f4556,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4555,f4447])).
fof(f4555,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X13)))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4554,f84])).
fof(f4554,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X13)))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4461,f213])).
fof(f4461,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X13)))) | s('type.realax.real',X13) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))),s('type.realax.real',X12))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f2846])).
fof(f4552,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4551,f84])).
fof(f4551,plain,(
  ( ! [X41,X40] : (p(s(bool,'F_0')) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4460,f213])).
fof(f4460,plain,(
  ( ! [X41,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X41))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f2828])).
fof(f4538,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4537,f84])).
fof(f4537,plain,(
  ( ! [X8,X7] : (p(s(bool,'F_0')) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4536,f213])).
fof(f4536,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4535,f4447])).
fof(f4535,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4534,f84])).
fof(f4534,plain,(
  ( ! [X8,X7] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4448,f213])).
fof(f4448,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))),s('type.realax.real',X7))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f652])).
fof(f4616,plain,(
  ( ! [X48,X49] : (s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X49) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | ($spl12 | $spl16)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f456,f480,f501,f502,f515,f517,f505,f507,f508,f509,f512,f513,f271,f252,f251,f236,f129,f89,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f269,f142,f1366,f1363,f1255,f1253,f1248,f1234,f710,f659,f658,f657,f656,f653,f652,f426,f256,f88,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2794,f2796,f2801,f2807,f2875,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3354,f3883,f3884,f4075,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2863,f2862,f2857,f2856,f2850,f2849,f2846,f2845,f2839,f2834,f2833,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f2779,f647,f3581,f3351,f2263,f4442,f4447,f4538,f4547,f4552,f4553,f4558,f4564,f4569,f4575,f4603,f4615])).
fof(f4572,plain,(
  $false | ($spl12 | $spl16 | $spl86)),
  inference(subsumption_resolution,[],[f4487,f4569])).
fof(f4487,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl12 | $spl16 | $spl86)),
  inference(backward_demodulation,[],[f4447,f2984])).
fof(f4564,plain,(
  ( ! [X19,X20] : (s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4563,f84])).
fof(f4563,plain,(
  ( ! [X19,X20] : (p(s(bool,'F_0')) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4562,f213])).
fof(f4562,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4561,f4447])).
fof(f4561,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4560,f84])).
fof(f4560,plain,(
  ( ! [X19,X20] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4464,f213])).
fof(f4464,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))),s('type.realax.real',X19))),s('type.realax.real',X20)))) | s('type.realax.real',X20) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X19) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f2817])).
fof(f4553,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X41) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | ($spl12 | $spl16)),
  inference(global_subsumption,[],[f80,f79,f81,f86,f114,f111,f116,f121,f123,f128,f154,f157,f158,f108,f109,f115,f83,f84,f209,f214,f213,f232,f270,f302,f366,f454,f459,f465,f470,f471,f472,f456,f480,f501,f502,f515,f517,f505,f507,f508,f509,f512,f513,f271,f252,f251,f236,f129,f89,f621,f622,f646,f693,f695,f831,f888,f889,f1115,f1116,f947,f948,f1280,f1500,f1502,f1552,f1561,f1524,f1373,f1262,f846,f813,f594,f592,f591,f85,f1515,f1504,f705,f701,f698,f590,f110,f107,f269,f142,f1366,f1363,f1255,f1253,f1248,f1234,f710,f659,f658,f657,f656,f653,f652,f426,f256,f88,f2563,f2526,f2530,f2577,f2579,f2580,f2612,f2613,f2603,f2634,f2619,f2620,f2621,f2622,f2623,f2645,f2646,f2633,f2724,f2761,f2731,f2734,f2735,f2765,f2743,f2744,f2794,f2796,f2801,f2807,f2875,f2812,f2882,f2892,f2841,f2896,f2898,f2860,f2861,f3354,f3883,f3884,f4075,f4166,f4168,f4241,f4243,f4250,f4253,f4254,f4255,f4256,f4309,f4313,f4337,f4338,f4398,f4382,f4409,f4410,f759,f2214,f1179,f2863,f2862,f2857,f2856,f2850,f2849,f2846,f2845,f2839,f2834,f2833,f2828,f2827,f2825,f2824,f2821,f2820,f2817,f2816,f2779,f647,f3581,f3351,f2263,f4442,f4447,f4538,f4547,f4552])).
fof(f4547,plain,(
  ( ! [X17,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4546,f84])).
fof(f4546,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4545,f213])).
fof(f4545,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4544,f4447])).
fof(f4544,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f4543,f84])).
fof(f4543,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f4458,f213])).
fof(f4458,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X18)))) | s('type.realax.real',X18) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X17) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f4447,f2816])).
fof(f385,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl12),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f387,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl15),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1179,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) )),
  inference(equality_factoring,[],[f269])).
fof(f2214,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) )),
  inference(equality_factoring,[],[f269])).
fof(f316,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl8),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f318,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl11),
  inference(cnf_transformation,[],[f318_D])).
fof(f318_D,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f759,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_factoring,[],[f271])).
fof(f4435,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl100),
  inference(cnf_transformation,[],[f4435_D])).
fof(f4435_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4437,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl102),
  inference(cnf_transformation,[],[f4437_D])).
fof(f4437_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4438,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(subsumption_resolution,[],[f4422,f81])).
fof(f4422,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(superposition,[],[f4260,f1280])).
fof(f4433,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(subsumption_resolution,[],[f4421,f81])).
fof(f4421,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl33 | $spl84)),
  inference(superposition,[],[f4260,f1280])).
fof(f4213,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl98),
  inference(cnf_transformation,[],[f4213_D])).
fof(f4213_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4214,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl85)),
  inference(subsumption_resolution,[],[f4198,f81])).
fof(f4198,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl85)),
  inference(superposition,[],[f4137,f1280])).
fof(f4209,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl85)),
  inference(subsumption_resolution,[],[f4197,f81])).
fof(f4197,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl33 | $spl85)),
  inference(superposition,[],[f4137,f1280])).
fof(f1518,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl33),
  inference(cnf_transformation,[],[f1518_D])).
fof(f1518_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4075,plain,(
  ( ! [X47] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))))))),s('type.realax.real',X47))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2832,f107])).
fof(f2832,plain,(
  ( ! [X47] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X47))))),s('type.realax.real',X47))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X47) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f594,f647])).
fof(f4074,plain,(
  $false | ($spl18 | $spl59 | $spl68)),
  inference(subsumption_resolution,[],[f4073,f84])).
fof(f4073,plain,(
  p(s(bool,'F_0')) | ($spl18 | $spl59 | $spl68)),
  inference(forward_demodulation,[],[f4072,f2884])).
fof(f4072,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl59)),
  inference(subsumption_resolution,[],[f4071,f2867])).
fof(f2867,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl59),
  inference(cnf_transformation,[],[f2867_D])).
fof(f2867_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f4070,plain,(
  $false | ($spl18 | $spl59 | $spl68)),
  inference(subsumption_resolution,[],[f4069,f84])).
fof(f4069,plain,(
  p(s(bool,'F_0')) | ($spl18 | $spl59 | $spl68)),
  inference(forward_demodulation,[],[f4068,f2884])).
fof(f4068,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl59)),
  inference(subsumption_resolution,[],[f4067,f2867])).
fof(f4067,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(forward_demodulation,[],[f4005,f1284])).
fof(f4066,plain,(
  $false | ($spl18 | $spl33 | $spl59)),
  inference(subsumption_resolution,[],[f2867,f4008])).
fof(f2903,plain,(
  ( ! [X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f2903_D])).
fof(f2903_D,plain,(
  ( ! [X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3994,plain,(
  ( ! [X49] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X49))))),s('type.realax.real',X49))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real',X49) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl28 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3964,f590])).
fof(f3964,plain,(
  ( ! [X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X49))))))),s('type.realax.real',X49)))))) | s('type.realax.real',X49) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl28 | $spl58 | $spl85)),
  inference(superposition,[],[f702,f3144])).
fof(f3144,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl58 | $spl85)),
  inference(backward_demodulation,[],[f3142,f2946])).
fof(f2946,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl58),
  inference(backward_demodulation,[],[f2868,f2839])).
fof(f3142,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3141,f2908])).
fof(f2908,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl58)),
  inference(backward_demodulation,[],[f2868,f1284])).
fof(f3141,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl58 | $spl85)),
  inference(subsumption_resolution,[],[f3140,f84])).
fof(f3140,plain,(
  p(s(bool,'F_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3139,f213])).
fof(f3139,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3138,f2908])).
fof(f3138,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl58 | $spl85)),
  inference(subsumption_resolution,[],[f3103,f2957])).
fof(f2957,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl58 | $spl85)),
  inference(backward_demodulation,[],[f2868,f2905])).
fof(f3103,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl58)),
  inference(superposition,[],[f592,f2908])).
fof(f3991,plain,(
  ( ! [X48] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))),s('type.realax.real',X48))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl26 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3963,f590])).
fof(f3963,plain,(
  ( ! [X48] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))))),s('type.realax.real',X48)))))) | s('type.realax.real',X48) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl26 | $spl58 | $spl85)),
  inference(superposition,[],[f700,f3144])).
fof(f3981,plain,(
  ( ! [X23] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23))))) ) | $spl88),
  inference(cnf_transformation,[],[f3981_D])).
fof(f3981_D,plain,(
  ( ! [X23] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3980,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X22))))),s('type.realax.real',X22))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) | s('type.realax.real',X22) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl58 | $spl85)),
  inference(forward_demodulation,[],[f3948,f590])).
fof(f3948,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X22))))))),s('type.realax.real',X22)))))) | s('type.realax.real',X22) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl18 | $spl58 | $spl85)),
  inference(superposition,[],[f110,f3144])).
fof(f2984,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl86),
  inference(cnf_transformation,[],[f2984_D])).
fof(f2984_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2970,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(superposition,[],[f700,f88])).
fof(f2905,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl85),
  inference(cnf_transformation,[],[f2905_D])).
fof(f2905_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f2861,plain,(
  ( ! [X40] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))))),s('type.realax.real',X40))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X40) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f594,f647])).
fof(f2900,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X39))))),s('type.realax.real',X39))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f2900_D])).
fof(f2900_D,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X39))))),s('type.realax.real',X39))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2902,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl80),
  inference(cnf_transformation,[],[f2902_D])).
fof(f2902_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2860,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X39))))),s('type.realax.real',X39))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X39) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f591,f647])).
fof(f2899,plain,(
  ( ! [X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))))) ) | $spl76),
  inference(cnf_transformation,[],[f2899_D])).
fof(f2899_D,plain,(
  ( ! [X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1516,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl30),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2898,plain,(
  ( ! [X24,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X24))),s('type.realax.real',X24))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25)))) | s('type.realax.real',X24) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2852,f590])).
fof(f2852,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X24))))),s('type.realax.real',X24)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25)))) | s('type.realax.real',X24) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f110,f647])).
fof(f2897,plain,(
  ( ! [X23] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X23))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23))))) ) | $spl74),
  inference(cnf_transformation,[],[f2897_D])).
fof(f2897_D,plain,(
  ( ! [X23] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X23))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2896,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X22))),s('type.realax.real',X22))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X23))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) | s('type.realax.real',X22) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2895,f590])).
fof(f2895,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X23))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X22))))),s('type.realax.real',X22)))))) | s('type.realax.real',X22) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2851,f590])).
fof(f2851,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X23)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X22))))),s('type.realax.real',X22)))))) | s('type.realax.real',X22) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f110,f647])).
fof(f2871,plain,(
  ( ! [X4] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl62),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  ( ! [X4] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2894,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl72),
  inference(cnf_transformation,[],[f2894_D])).
fof(f2894_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2841,plain,(
  ( ! [X4] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f88,f647])).
fof(f2884,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl68),
  inference(cnf_transformation,[],[f2884_D])).
fof(f2884_D,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2885,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl70),
  inference(cnf_transformation,[],[f2885_D])).
fof(f2885_D,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2882,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))))),s('type.realax.real',X46))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2831,f107])).
fof(f2831,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))),s('type.realax.real',X46))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X46) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f591,f647])).
fof(f2868,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl58),
  inference(cnf_transformation,[],[f2868_D])).
fof(f2868_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2876,plain,(
  ( ! [X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))) | s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl66),
  inference(cnf_transformation,[],[f2876_D])).
fof(f2876_D,plain,(
  ( ! [X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))) | s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2812,plain,(
  ( ! [X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))),s('type.realax.real',X11))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f88,f647])).
fof(f2873,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl64),
  inference(cnf_transformation,[],[f2873_D])).
fof(f2873_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2801,plain,(
  ( ! [X4] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f647,f647])).
fof(f2869,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f2869_D])).
fof(f2869_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2796,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X2))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f647,f647])).
fof(f2640,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool))))) ) | $spl50),
  inference(cnf_transformation,[],[f2640_D])).
fof(f2640_D,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2646,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X18,X19,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X18,X19,bool,bool))) )),
  inference(forward_demodulation,[],[f2625,f236])).
fof(f2625,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,f(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X18,X19,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X18,X19,bool,bool))) )),
  inference(superposition,[],[f622,f271])).
fof(f2645,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X16,X17,bool,bool))) )),
  inference(forward_demodulation,[],[f2624,f142])).
fof(f2624,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X16,X17,bool,bool))) )),
  inference(superposition,[],[f622,f269])).
fof(f2643,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f2643_D])).
fof(f2643_D,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2623,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X13,X14,bool,bool)))) | s(bool,X15) = s(bool,sK1(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f622,f121])).
fof(f2644,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl56),
  inference(cnf_transformation,[],[f2644_D])).
fof(f2644_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2622,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f622,f121])).
fof(f2621,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK1(X7,X8,bool,bool))) )),
  inference(superposition,[],[f622,f121])).
fof(f2620,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool)))) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f622,f121])).
fof(f2619,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool))) )),
  inference(superposition,[],[f622,f86])).
fof(f2636,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) | $spl45),
  inference(cnf_transformation,[],[f2636_D])).
fof(f2636_D,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f2634,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f2618])).
fof(f2618,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f622,f621])).
fof(f2564,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl40),
  inference(cnf_transformation,[],[f2564_D])).
fof(f2564_D,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2566,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl42),
  inference(cnf_transformation,[],[f2566_D])).
fof(f2566_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2526,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f591,f88])).
fof(f2494,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl36),
  inference(cnf_transformation,[],[f2494_D])).
fof(f2494_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2496,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl39),
  inference(cnf_transformation,[],[f2496_D])).
fof(f2496_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f2463,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl18),
  inference(superposition,[],[f89,f1284])).
fof(f1905,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1865,f1821])).
fof(f1821,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(subsumption_resolution,[],[f1820,f81])).
fof(f1820,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(forward_demodulation,[],[f1778,f256])).
fof(f1778,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f1766])).
fof(f1766,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl24),
  inference(superposition,[],[f699,f88])).
fof(f1865,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f608])).
fof(f1904,plain,(
  $false | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f1863,f1823])).
fof(f1823,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl24),
  inference(backward_demodulation,[],[f1821,f256])).
fof(f1863,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl15 | $spl24)),
  inference(backward_demodulation,[],[f1821,f387])).
fof(f1903,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1858,f1821])).
fof(f1858,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f1528])).
fof(f1902,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1857,f1821])).
fof(f1857,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f1359])).
fof(f1901,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1856,f1821])).
fof(f1856,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f1225])).
fof(f1900,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1855,f1821])).
fof(f1855,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f783])).
fof(f1899,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1854,f1821])).
fof(f1854,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f709])).
fof(f1898,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1853,f1821])).
fof(f1853,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f707])).
fof(f1897,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1852,f1821])).
fof(f1852,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1821,f623])).
fof(f1886,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1824,f1821])).
fof(f1824,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl24),
  inference(backward_demodulation,[],[f1821,f426])).
fof(f1819,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1818,f252])).
fof(f1818,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1817,f644])).
fof(f1817,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1816,f81])).
fof(f1816,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1815,f256])).
fof(f1815,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1765,f710])).
fof(f1765,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl24)),
  inference(superposition,[],[f699,f644])).
fof(f1814,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1813,f252])).
fof(f1813,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f1812,f643])).
fof(f1812,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1811,f81])).
fof(f1811,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1810,f256])).
fof(f1810,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1764,f710])).
fof(f1764,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl24)),
  inference(superposition,[],[f699,f643])).
fof(f1805,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1804,f81])).
fof(f1804,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f1803,f256])).
fof(f1803,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl24),
  inference(forward_demodulation,[],[f1753,f710])).
fof(f1753,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl24),
  inference(unit_resulting_resolution,[],[f252,f699])).
fof(f1802,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1801,f81])).
fof(f1801,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1800,f256])).
fof(f1800,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1754,f643])).
fof(f1754,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f579,f699])).
fof(f1799,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1798,f81])).
fof(f1798,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1797,f256])).
fof(f1797,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1755,f644])).
fof(f1755,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f608,f699])).
fof(f1796,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1795,f81])).
fof(f1795,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1794,f256])).
fof(f1794,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1756,f706])).
fof(f1756,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f623,f699])).
fof(f1793,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1792,f81])).
fof(f1792,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1791,f256])).
fof(f1791,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1757,f708])).
fof(f1757,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f707,f699])).
fof(f1790,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1789,f81])).
fof(f1789,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1788,f256])).
fof(f1788,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1758,f782])).
fof(f1758,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f709,f699])).
fof(f1787,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1786,f81])).
fof(f1786,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1785,f256])).
fof(f1785,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1759,f1224])).
fof(f1759,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f783,f699])).
fof(f1784,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1783,f81])).
fof(f1783,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1782,f256])).
fof(f1782,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1760,f1358])).
fof(f1760,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f1225,f699])).
fof(f1781,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f1780,f81])).
fof(f1780,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1779,f256])).
fof(f1779,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl18 | $spl24)),
  inference(forward_demodulation,[],[f1761,f1527])).
fof(f1761,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))))))))))))))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f1359,f699])).
fof(f1606,plain,(
  ( ! [X35] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))),s('type.realax.real',X35))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X35) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f594,f662])).
fof(f1639,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1638,f579])).
fof(f1638,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1585,f252])).
fof(f1585,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl18 | $spl20)),
  inference(superposition,[],[f643,f662])).
fof(f1637,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1612,f252])).
fof(f1612,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl20),
  inference(condensation,[],[f1611])).
fof(f1611,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f1584])).
fof(f1584,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f88,f662])).
fof(f1636,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1635,f579])).
fof(f1635,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1583,f252])).
fof(f1583,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ($spl18 | $spl20)),
  inference(superposition,[],[f662,f643])).
fof(f1634,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1614,f252])).
fof(f1614,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl20),
  inference(condensation,[],[f1613])).
fof(f1613,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f1582])).
fof(f1582,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f662,f88])).
fof(f1632,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1631,f252])).
fof(f1631,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl20),
  inference(backward_demodulation,[],[f1569,f710])).
fof(f1569,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl20),
  inference(unit_resulting_resolution,[],[f252,f662])).
fof(f1630,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1629,f252])).
fof(f1629,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1570,f643])).
fof(f1570,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f579,f662])).
fof(f1628,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1627,f252])).
fof(f1627,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1571,f644])).
fof(f1571,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f608,f662])).
fof(f1626,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1625,f252])).
fof(f1625,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1572,f706])).
fof(f1572,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f623,f662])).
fof(f1624,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1623,f252])).
fof(f1623,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1573,f708])).
fof(f1573,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f707,f662])).
fof(f1622,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1621,f252])).
fof(f1621,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1574,f782])).
fof(f1574,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f709,f662])).
fof(f1620,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1619,f252])).
fof(f1619,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1575,f1224])).
fof(f1575,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f783,f662])).
fof(f1618,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1617,f252])).
fof(f1617,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1576,f1358])).
fof(f1576,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f1225,f662])).
fof(f1616,plain,(
  $false | ($spl18 | $spl20)),
  inference(subsumption_resolution,[],[f1615,f252])).
fof(f1615,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl18 | $spl20)),
  inference(backward_demodulation,[],[f1577,f1527])).
fof(f1577,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))))))))))))))))) | ($spl18 | $spl20)),
  inference(unit_resulting_resolution,[],[f1359,f662])).
fof(f1502,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f594,f88])).
fof(f702,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | $spl28),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f695,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f110,f88])).
fof(f700,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | $spl26),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  ( ! [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f693,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f110,f88])).
fof(f664,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl22),
  inference(cnf_transformation,[],[f664_D])).
fof(f664_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f646,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f88,f88])).
fof(f519,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f514,f505])).
fof(f514,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f480,f477])).
fof(f477,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f456,f387])).
fof(f518,plain,(
  $false | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f511,f505])).
fof(f511,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))) | ($spl11 | $spl16)),
  inference(backward_demodulation,[],[f480,f318])).
fof(f456,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl16),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f458,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl18),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f472,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f443,f84])).
fof(f443,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f471,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(global_subsumption,[],[f80,f79,f81,f86,f88,f89,f107,f114,f111,f116,f121,f123,f129,f128,f154,f157,f158,f108,f109,f115,f85,f83,f110,f84,f142,f209,f214,f213,f232,f236,f256,f252,f251,f269,f270,f271,f302,f366,f426,f454,f459,f465,f470])).
fof(f459,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f432,f84])).
fof(f432,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))),
  inference(superposition,[],[f108,f142])).
fof(f454,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f431,f123])).
fof(f431,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f269])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) )),
  inference(equality_factoring,[],[f269])).
fof(f302,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X0) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_factoring,[],[f271])).
fof(f172,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f171,f81])).
fof(f171,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f170,f155])).
fof(f170,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f164,f81])).
fof(f164,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl4),
  inference(backward_demodulation,[],[f155,f115])).
fof(f167,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f160,f81])).
fof(f160,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f155,f84])).
fof(f156,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f150,f123])).
fof(f150,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f111,f86])).
fof(f154,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f147,f123])).
fof(f147,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f111,f86])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f133,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f128,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f86])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 326
% Passive clauses: 2184
% Generated clauses: 28880
% Final active clauses: 103
% Final passive clauses: 282
% Input formulas: 23
% Initial clauses: 38
% 
% Duplicate literals: 1333
% Trivial inequalities: 95
% Fw subsumption resolutions: 2024
% Fw demodulations: 7500
% Bw demodulations: 1024
% Condensations: 1266
% Global subsumptions: 75
% 
% Simple tautologies: 165
% Equational tautologies: 196
% Forward subsumptions: 13694
% Fw demodulations to eq. taut.: 31
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 1464
% Factoring: 3
% Forward superposition: 6139
% Backward superposition: 5805
% Self superposition: 162
% Equality factoring: 78
% Equality resolution: 13
% 
% Split clauses: 136
% Split components: 275
% SAT solver clauses: 2725
% SAT solver unit clauses: 89
% SAT solver binary clauses: 558
% 
% Sat splits: 136
% Sat splitting refutations: 192
% 
% Memory used [KB]: 10874
% Time elapsed: 0.965 s
% ------------------------------
% ------------------------------
% Success in time 2.282 s
