Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (83975)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1085
% Passive clauses: 7185
% Generated clauses: 31161
% Final active clauses: 787
% Final passive clauses: 3831
% Input formulas: 43
% Initial clauses: 177
% 
% Duplicate literals: 658
% Trivial inequalities: 180
% Fw subsumption resolutions: 3008
% Fw demodulations: 2027
% Bw demodulations: 44
% 
% Simple tautologies: 271
% Equational tautologies: 97
% Forward subsumptions: 16061
% 
% Binary resolution: 12791
% Factoring: 1214
% Forward superposition: 4142
% Backward superposition: 5830
% Self superposition: 40
% Equality factoring: 28
% Equality resolution: 18
% 
% Split clauses: 1561
% Split components: 3137
% SAT solver clauses: 2357
% SAT solver unit clauses: 23
% SAT solver binary clauses: 163
% SAT solver learnt clauses: 6
% SAT solver learnt literals: 8
% 
% Sat splits: 1561
% Sat splitting refutations: 40
% 
% TWLsolver clauses: 2344
% TWLsolver calls for satisfiability: 600
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (83987)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 521
% Passive clauses: 3166
% Generated clauses: 19001
% Final active clauses: 274
% Final passive clauses: 1662
% Input formulas: 43
% Initial clauses: 177
% 
% Duplicate literals: 278
% Trivial inequalities: 81
% Fw subsumption resolutions: 1405
% Fw demodulations: 1331
% Bw demodulations: 1021
% Condensations: 359
% Global subsumptions: 14
% 
% Simple tautologies: 266
% Equational tautologies: 26
% Forward subsumptions: 8083
% Bw demodulations to eq. taut.: 8
% 
% Binary resolution: 3946
% Factoring: 143
% Forward superposition: 1950
% Backward superposition: 1996
% Self superposition: 45
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 360
% Split components: 726
% SAT solver clauses: 3302
% SAT solver unit clauses: 177
% SAT solver binary clauses: 102
% SAT solver learnt clauses: 34
% SAT solver learnt literals: 21
% 
% Sat splits: 360
% Sat splitting refutations: 54
% 
% TWLsolver clauses: 3240
% TWLsolver calls for satisfiability: 1727
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f21599,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f230,f417,f229,f419,f415,f228,f416,f418,f414,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f3365,f3397,f3329,f3399,f3333,f3401,f3336,f3403,f3339,f3405,f3342,f3407,f3345,f3409,f3348,f3411,f3351,f3413,f3354,f3415,f3357,f3417,f3360,f3419,f3363,f3421,f3366,f3534,f3535,f3536,f3537,f3538,f3539,f3540,f3541,f3542,f3543,f3544,f3545,f3546,f3547,f3548,f3550,f3552,f3553,f3554,f3555,f3556,f3557,f3558,f3559,f3560,f3561,f3562,f3563,f3564,f3565,f3566,f3567,f3568,f3569,f3570,f3571,f3572,f3573,f3574,f3575,f3576,f3577,f3578,f3579,f3580,f3581,f3582,f3583,f3584,f3585,f3586,f3587,f3588,f3589,f3590,f3591,f3592,f3593,f3594,f3595,f3596,f3597,f3598,f3599,f3600,f3601,f3603,f3607,f3610,f3613,f3616,f3619,f3608,f3621,f3611,f3623,f3614,f3625,f3617,f6865,f6868,f6866,f6908,f6911,f6909,f6935,f6938,f6936,f6968,f6971,f6969,f7020,f7023,f7021,f7071,f7074,f7072,f7098,f7101,f7099,f7129,f7132,f7130,f7760,f7763,f7762,f7761,f7798,f7801,f7799,f7837,f7840,f7838,f7879,f7882,f7880,f7926,f7929,f7927,f7970,f7973,f7971,f8019,f8022,f8020,f8065,f8068,f8066,f8101,f8104,f8102,f8137,f8222,f8138,f8223,f8139,f8224,f8140,f8225,f8141,f8226,f8142,f8227,f8143,f8228,f8144,f8229,f8145,f8230,f8221,f10842,f10843,f10844,f10845,f10846,f10847,f10848,f10849,f10850,f10851,f10852,f10853,f10854,f10855,f10856,f10857,f10858,f10859,f10860,f10861,f10862,f10863,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f10873,f10874,f10875,f10876,f10877,f10878,f10879,f10880,f10881,f10882,f10883,f10884,f10885,f10886,f10887,f10888,f10889,f10890,f10891,f10892,f10893,f10894,f10895,f10896,f10897,f10898,f10899,f10900,f10901,f10902,f10903,f10904,f13279,f13282,f13280,f13331,f13334,f13332,f13384,f13387,f13385,f13442,f13445,f3330,f13443,f13475,f13476,f13483,f13480,f13477,f13478,f13486,f13479,f13488,f13481,f13485,f13499,f13500,f13512,f13510,f16516,f16521,f16525,f16544,f16551,f16558,f16561,f16568,f16570,f16572,f16574,f16576,f16577,f16578,f16579,f16580,f16581,f16582,f16583,f16584,f16585,f16586,f16587,f16588,f16591,f16592,f16593,f16594,f16595,f16596,f16597,f16598,f16599,f16600,f16601,f16602,f16603,f16604,f16605,f21583,f21592,f21597])).
fof(f21597,plain,(
  $false | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(subsumption_resolution,[],[f21596,f420])).
fof(f420,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f234,f231])).
fof(f231,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.le_max',aHL_TRUTH)).
fof(f234,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.le_max','thm.bool.NOT_CLAUSES|split|1')).
fof(f21596,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(forward_demodulation,[],[f21595,f17610])).
fof(f17610,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f462,f16606,f710])).
fof(f710,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f703,f420])).
fof(f703,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f282,f239])).
fof(f239,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.le_max',aHL_BOOL_CASES)).
fof(f282,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.extreal.le_max','thm.bool.IMP_ANTISYM_AX')).
fof(f16606,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl1 | $spl5)),
  inference(global_subsumption,[],[f11608,f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f907,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1042,f1052,f1061,f1079,f1081,f1100,f1151,f1156,f1158,f1169,f1179,f1188,f1227,f1253,f1535,f1532,f1587,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1815,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2015,f2022,f2045,f2051,f2089,f2099,f2128,f2138,f2166,f2176,f2202,f2212,f2233,f2255,f2387,f2411,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2306,f2285,f2213,f2198,f2177,f2163,f2161,f2141,f2139,f2125,f2102,f2100,f2086,f2084,f2054,f2052,f2025,f2023,f2017,f1816,f1812,f1747,f1746,f1639,f1228,f1159,f1157,f1101,f1082,f1080,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8194,f8195,f8198,f8199,f8200,f8201,f419,f8561,f8560,f8559,f8558,f8556,f8555,f8554,f8552,f8551,f8550,f8548,f8546,f8545,f8544,f8543,f8542,f8541,f8540,f8536,f8535,f8532,f8530,f8528,f8525,f8521,f8520,f8516,f8515,f8514,f8509,f8492,f8490,f8489,f8487,f8485,f8484,f8481,f8480,f8479,f8478,f8477,f8474,f8472,f8467,f8466,f8462,f8461,f8456,f8431,f8430,f8429,f8428,f8423,f8422,f8417,f8416,f8415,f8410,f8405,f8400,f8399,f8398,f8397,f8396,f8394,f8392,f8389,f8388,f8387,f8386,f8381,f8380,f8374,f8373,f8372,f8367,f8366,f8365,f8360,f8354,f8349,f8346,f8345,f8343,f8337,f8336,f8334,f8604,f8955,f8960,f8973,f8981,f9002,f9000,f9010,f9015,f9020,f9585,f9719,f9735,f9752,f10289,f10288,f10287,f10286,f10390,f10391,f10397,f10398,f10399,f10400,f10401,f10402,f10403,f10404,f10405,f10406,f10407,f10408,f10409,f10410,f10323,f10325,f10327,f10329,f10419,f10420,f10423,f10342,f10344,f10345,f10347,f10432,f10433,f10436,f10437,f10438,f10439,f10440,f10445,f10446,f10462,f10467,f10474,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514])).
fof(f13514,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK15(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK15(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f13513])).
fof(f13513,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK15(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK15(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f13506,f239])).
fof(f13506,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK15(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK15(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1747,f651])).
fof(f13439,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)))),X74,X75) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)))),X74,X75) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f877,f845])).
fof(f13440,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13423,f231])).
fof(f13423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f877,f651])).
fof(f13381,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)))),X74,X75) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X73),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)))),X74,X75) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f876,f845])).
fof(f13382,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13365,f231])).
fof(f13365,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f876,f651])).
fof(f13328,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))),s(bool,X73)),X74,X75) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))),s(bool,X73)),X74,X75) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f819,f845])).
fof(f13329,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13312,f231])).
fof(f13312,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f819,f651])).
fof(f13276,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))),s(bool,X73)),X74,X75) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))),s(bool,X73)),X74,X75) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f818,f845])).
fof(f13277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13260,f231])).
fof(f13260,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X2,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f818,f651])).
fof(f10474,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f440,f430,f10288,f337])).
fof(f10467,plain,(
  sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) | $spl5),
  inference(unit_resulting_resolution,[],[f10287,f333])).
fof(f10462,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'T_0','F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f442,f430,f10286,f335])).
fof(f10446,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X115)),X116,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10367,f231])).
fof(f10367,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,X116)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X115)),X116,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f1228,f8334])).
fof(f10445,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,X112)) | p(s(bool,X113)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X112),s(bool,X114)),X113,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10366,f231])).
fof(f10366,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X112)) | p(s(bool,X113)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X112),s(bool,X114)),X113,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f1227,f8334])).
fof(f10440,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP7(X104,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X103)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10363,f231])).
fof(f10363,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X103)) | p(s(bool,X104)) | ~sP7(X104,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X103)))) ) | $spl5),
  inference(superposition,[],[f1009,f8334])).
fof(f10439,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X100),s(bool,X102)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10362,f231])).
fof(f10362,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'T_0')) | ~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X100),s(bool,X102)))) ) | $spl5),
  inference(superposition,[],[f1008,f8334])).
fof(f10438,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X99),s(bool,X97)),X98)) ) | $spl5),
  inference(subsumption_resolution,[],[f10361,f231])).
fof(f10361,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X99),s(bool,X97)),X98)) ) | $spl5),
  inference(superposition,[],[f908,f8334])).
fof(f10437,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X94),s(bool,X96)),X95)) ) | $spl5),
  inference(subsumption_resolution,[],[f10360,f231])).
fof(f10360,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X94)) | p(s(bool,X95)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X94),s(bool,X96)),X95)) ) | $spl5),
  inference(superposition,[],[f907,f8334])).
fof(f10436,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X91)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X93)) ) | $spl5),
  inference(subsumption_resolution,[],[f10359,f231])).
fof(f10359,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X91)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X91)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X93)) ) | $spl5),
  inference(superposition,[],[f875,f8334])).
fof(f10433,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,X85)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X86),s(bool,X85)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X87)) ) | $spl5),
  inference(subsumption_resolution,[],[f10357,f231])).
fof(f10357,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X85)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X86),s(bool,X85)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X87)) ) | $spl5),
  inference(superposition,[],[f873,f8334])).
fof(f10432,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,X82)),X84,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10356,f231])).
fof(f10356,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X82)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,X82)),X84,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f872,f8334])).
fof(f10347,plain,(
  ( ! [X67] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,'T_0')))) ) | $spl5),
  inference(superposition,[],[f844,f8334])).
fof(f10345,plain,(
  ( ! [X65] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,'T_0')))) ) | $spl5),
  inference(superposition,[],[f843,f8334])).
fof(f10344,plain,(
  ( ! [X64] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f842,f8334])).
fof(f10342,plain,(
  ( ! [X62] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X62),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f841,f8334])).
fof(f10423,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X58),s(bool,X59)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X60)) ) | $spl5),
  inference(subsumption_resolution,[],[f10340,f231])).
fof(f10340,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X58),s(bool,X59)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X60)) ) | $spl5),
  inference(superposition,[],[f817,f8334])).
fof(f10420,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X54)) ) | $spl5),
  inference(subsumption_resolution,[],[f10338,f231])).
fof(f10338,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X54)) ) | $spl5),
  inference(superposition,[],[f815,f8334])).
fof(f10419,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10337,f231])).
fof(f10337,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f814,f8334])).
fof(f10329,plain,(
  ( ! [X37] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X37)))) ) | $spl5),
  inference(superposition,[],[f790,f8334])).
fof(f10327,plain,(
  ( ! [X35] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X35)))) ) | $spl5),
  inference(superposition,[],[f789,f8334])).
fof(f10325,plain,(
  ( ! [X33] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X33)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f788,f8334])).
fof(f10323,plain,(
  ( ! [X31] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f787,f8334])).
fof(f10410,plain,(
  ( ! [X26,X27] : (~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X26,X27) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X26,X27)) ) | $spl5),
  inference(subsumption_resolution,[],[f10319,f231])).
fof(f10319,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X26,X27) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X26,X27)) ) | $spl5),
  inference(superposition,[],[f406,f8334])).
fof(f10409,plain,(
  ( ! [X24,X25] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X24,X25) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X24,X25)) ) | $spl5),
  inference(subsumption_resolution,[],[f10318,f231])).
fof(f10318,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X24,X25) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X24,X25)) ) | $spl5),
  inference(superposition,[],[f405,f8334])).
fof(f10408,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,X23)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X23,X22)) ) | $spl5),
  inference(subsumption_resolution,[],[f10317,f231])).
fof(f10317,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | p(s(bool,X23)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X23,X22)) ) | $spl5),
  inference(superposition,[],[f402,f8334])).
fof(f10407,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X20,X21)) ) | $spl5),
  inference(subsumption_resolution,[],[f10316,f231])).
fof(f10316,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X20,X21)) ) | $spl5),
  inference(superposition,[],[f401,f8334])).
fof(f10406,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X19,X18)) ) | $spl5),
  inference(subsumption_resolution,[],[f10315,f231])).
fof(f10315,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X19,X18)) ) | $spl5),
  inference(superposition,[],[f400,f8334])).
fof(f10405,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X17,X16)) ) | $spl5),
  inference(subsumption_resolution,[],[f10314,f231])).
fof(f10314,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X17,X16)) ) | $spl5),
  inference(superposition,[],[f385,f8334])).
fof(f10404,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X14,X15)) ) | $spl5),
  inference(subsumption_resolution,[],[f10313,f231])).
fof(f10313,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X14,X15)) ) | $spl5),
  inference(superposition,[],[f384,f8334])).
fof(f10403,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X13,X12)) ) | $spl5),
  inference(subsumption_resolution,[],[f10312,f231])).
fof(f10312,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X13,X12)) ) | $spl5),
  inference(superposition,[],[f383,f8334])).
fof(f10402,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | ~sP7(X11,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10311,f231])).
fof(f10311,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP7(X11,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f370,f8334])).
fof(f10401,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X9)) ) | $spl5),
  inference(subsumption_resolution,[],[f10310,f231])).
fof(f10310,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X9)) ) | $spl5),
  inference(superposition,[],[f334,f8334])).
fof(f10400,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10309,f231])).
fof(f10309,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f310,f8334])).
fof(f10399,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X6)) ) | $spl5),
  inference(subsumption_resolution,[],[f10308,f231])).
fof(f10308,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X6)) ) | $spl5),
  inference(superposition,[],[f307,f8334])).
fof(f10398,plain,(
  ( ! [X5] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X5)) ) | $spl5),
  inference(subsumption_resolution,[],[f10307,f231])).
fof(f10307,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),X5)) ) | $spl5),
  inference(superposition,[],[f292,f8334])).
fof(f10397,plain,(
  ( ! [X4] : (sP0(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(subsumption_resolution,[],[f10306,f231])).
fof(f10306,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) ) | $spl5),
  inference(superposition,[],[f291,f8334])).
fof(f10391,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) | $spl5),
  inference(forward_demodulation,[],[f10295,f276])).
fof(f10295,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(superposition,[],[f280,f8334])).
fof(f10390,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10294,f231])).
fof(f10294,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0))))) ) | $spl5),
  inference(superposition,[],[f324,f8334])).
fof(f10286,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f8334,f353])).
fof(f10287,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f8334,f351])).
fof(f10288,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f8334,f353])).
fof(f10289,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f8334,f351])).
fof(f9752,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f8388,f331])).
fof(f9735,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f8387,f331])).
fof(f9719,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f8381,f331])).
fof(f9585,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f8337,f1695])).
fof(f9020,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f8386,f331])).
fof(f9015,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f8380,f331])).
fof(f9010,plain,(
  ~sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f8374,f331])).
fof(f9000,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f442,f430,f8373,f335])).
fof(f9002,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f442,f8366,f8373,f335])).
fof(f8981,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f440,f430,f8372,f337])).
fof(f8973,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0','T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f8367,f332])).
fof(f8960,plain,(
  sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f438,f436,f8365,f336])).
fof(f8955,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f8360,f333])).
fof(f8604,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(subsumption_resolution,[],[f8565,f231])).
fof(f8565,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f419,f651])).
fof(f8334,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f1273])).
fof(f8336,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f281])).
fof(f8337,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f249])).
fof(f8343,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f305])).
fof(f8345,plain,(
  ~sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f305])).
fof(f8346,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f305])).
fof(f8349,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f334])).
fof(f8354,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f438,f428,f434,f419,f336])).
fof(f8360,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f338])).
fof(f8365,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f340])).
fof(f8366,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f341])).
fof(f8367,plain,(
  ~sP3('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f342])).
fof(f8372,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f344])).
fof(f8373,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f345])).
fof(f8374,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f346])).
fof(f8380,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f346])).
fof(f8381,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f419,f420,f419,f346])).
fof(f8386,plain,(
  ~sP2('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f346])).
fof(f8387,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f420,f419,f346])).
fof(f8388,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f419,f346])).
fof(f8389,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f419,f346])).
fof(f8392,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f462,f419,f350])).
fof(f8394,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0','F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f462,f419,f350])).
fof(f8396,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f276,f419,f353])).
fof(f8397,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f277,f419,f353])).
fof(f8398,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f276,f419,f353])).
fof(f8399,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f277,f419,f353])).
fof(f8400,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f355])).
fof(f8405,plain,(
  ( ! [X0] : (~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f356])).
fof(f8410,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f361])).
fof(f8415,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f363])).
fof(f8416,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f364])).
fof(f8417,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f365])).
fof(f8422,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f365])).
fof(f8423,plain,(
  ( ! [X0] : (sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f366])).
fof(f8428,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f366])).
fof(f8429,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0','T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f370])).
fof(f8430,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'T_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f371])).
fof(f8431,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0',X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f372])).
fof(f8456,plain,(
  ~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f420,f419,f378])).
fof(f8461,plain,(
  ~sP8('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f420,f419,f378])).
fof(f8462,plain,(
  ~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f419,f378])).
fof(f8466,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f381])).
fof(f8467,plain,(
  sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f382])).
fof(f8472,plain,(
  sP9('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f382])).
fof(f8474,plain,(
  sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f419,f382])).
fof(f8477,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f420,f419,f382])).
fof(f8478,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f419,f420,f419,f382])).
fof(f8479,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f419,f382])).
fof(f8480,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f419,f382])).
fof(f8481,plain,(
  ~sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f231,f419,f385])).
fof(f8484,plain,(
  ~sP9('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f420,f419,f385])).
fof(f8485,plain,(
  ~sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f419,f385])).
fof(f8487,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'F_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f388])).
fof(f8489,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f388])).
fof(f8490,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f388])).
fof(f8492,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0',X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f389])).
fof(f8509,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f395])).
fof(f8514,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f396])).
fof(f8515,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f398])).
fof(f8516,plain,(
  sP11('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f231,f419,f399])).
fof(f8520,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f231,f419,f399])).
fof(f8521,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f419,f231,f419,f399])).
fof(f8525,plain,(
  ~sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f402])).
fof(f8528,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,'F_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f446,f419,f407])).
fof(f8530,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0',X0)) ) | $spl5),
  inference(unit_resulting_resolution,[],[f448,f419,f407])).
fof(f8532,plain,(
  ~sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f421])).
fof(f8535,plain,(
  ~sP0('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f419,f421])).
fof(f8536,plain,(
  ~sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f419,f421])).
fof(f8540,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f812])).
fof(f8541,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f813])).
fof(f8542,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f816])).
fof(f8543,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f870])).
fof(f8544,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f871])).
fof(f8545,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f874])).
fof(f8546,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f907])).
fof(f8548,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f908])).
fof(f8550,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f1008])).
fof(f8551,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f1009])).
fof(f8552,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f420,f419,f1100])).
fof(f8554,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f231,f419,f1100])).
fof(f8555,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f419,f231,f419,f1100])).
fof(f8556,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f420,f231,f419,f1101])).
fof(f8558,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f420,f419,f1101])).
fof(f8559,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f419,f419,f1101])).
fof(f8560,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f1227])).
fof(f8561,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0')) ) | $spl5),
  inference(unit_resulting_resolution,[],[f231,f231,f419,f1228])).
fof(f8201,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X66),s('type.extreal.extreal',X68)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X68),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X66),s('type.extreal.extreal',X67))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X66),s('type.extreal.extreal',X67)))) )),
  inference(superposition,[],[f1273,f1695])).
fof(f8200,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X65),s('type.extreal.extreal',X63)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X63),s('type.extreal.extreal',X64))),s('type.extreal.extreal',X65))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X63),s('type.extreal.extreal',X64)))) )),
  inference(superposition,[],[f1273,f1695])).
fof(f8199,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X59)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X59),s('type.extreal.extreal',X60))),s('type.extreal.extreal',X62)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X62)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X59),s('type.extreal.extreal',X60)))) )),
  inference(superposition,[],[f324,f1695])).
fof(f8198,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X55),s('type.extreal.extreal',X57)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X57),s('type.extreal.extreal',X58)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X55),s('type.extreal.extreal',X56))),s('type.extreal.extreal',X58)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X55),s('type.extreal.extreal',X56)))) )),
  inference(superposition,[],[f324,f1695])).
fof(f8195,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X46)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X47))),s('type.extreal.extreal',X48)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X47)))) )),
  inference(superposition,[],[f281,f1695])).
fof(f8194,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X45)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X45),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X44)))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X44)))) )),
  inference(superposition,[],[f281,f1695])).
fof(f8133,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.extreal.extreal',X18) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',X17))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.bool.COND_3'(s(bool,X15),s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1695,f845])).
fof(f8130,plain,(
  ( ! [X6,X8,X7] : (s('type.extreal.extreal',X8) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) )),
  inference(superposition,[],[f1695,f1695])).
fof(f8125,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.extreal.extreal',X17) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X17),s('type.extreal.extreal',X18))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.bool.COND_3'(s(bool,X15),s('type.extreal.extreal',X16),s('type.extreal.extreal',X17))),s('type.extreal.extreal',X18))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1695,f845])).
fof(f8122,plain,(
  ( ! [X6,X8,X7] : (s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) )),
  inference(superposition,[],[f1695,f1695])).
fof(f8098,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X73)),X74,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1228,f845])).
fof(f8097,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X67)),X68,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) )),
  inference(superposition,[],[f1228,f277])).
fof(f8096,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,X62)),X63,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)))) )),
  inference(superposition,[],[f1228,f276])).
fof(f8108,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57)),X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f8095,f231])).
fof(f8095,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57)),X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f1228,f859])).
fof(f8107,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52)),X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f8094,f231])).
fof(f8094,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52)),X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f1228,f801])).
fof(f8099,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8082,f231])).
fof(f8082,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1228,f651])).
fof(f8081,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f1228])).
fof(f8062,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X73),s(bool,X75)),X74,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1227,f845])).
fof(f8061,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X69)),X68,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) )),
  inference(superposition,[],[f1227,f277])).
fof(f8060,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X64)),X63,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)))) )),
  inference(superposition,[],[f1227,f276])).
fof(f8072,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59)),X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f8059,f231])).
fof(f8059,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59)),X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f1227,f859])).
fof(f8071,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54)),X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f8058,f231])).
fof(f8058,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54)),X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f1227,f801])).
fof(f8063,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8046,f231])).
fof(f8046,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1227,f651])).
fof(f8045,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f1227])).
fof(f8016,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)))),X74,X73) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1101,f845])).
fof(f8017,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X0)),X2,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8000,f231])).
fof(f8000,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X0)),X2,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1101,f651])).
fof(f7999,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f420,f1101])).
fof(f7967,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))),s(bool,X75)),X74,X73) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1100,f845])).
fof(f7968,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X3)),X2,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7951,f231])).
fof(f7951,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X3)),X2,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1100,f651])).
fof(f7950,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f420,f1100])).
fof(f7921,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP7(X74,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X73))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1009,f845])).
fof(f7920,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP7(X68,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X67)))) )),
  inference(superposition,[],[f1009,f277])).
fof(f7919,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP7(X63,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,X62)))) )),
  inference(superposition,[],[f1009,f276])).
fof(f7933,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7918,f231])).
fof(f7918,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f1009,f859])).
fof(f7932,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f7917,f231])).
fof(f7917,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f1009,f801])).
fof(f7924,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7905,f231])).
fof(f7905,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1009,f651])).
fof(f7901,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f1009])).
fof(f7874,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP7(X74,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X73),s(bool,X75))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1008,f845])).
fof(f7873,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP7(X68,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X69)))) )),
  inference(superposition,[],[f1008,f277])).
fof(f7872,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP7(X63,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X64)))) )),
  inference(superposition,[],[f1008,f276])).
fof(f7886,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7871,f231])).
fof(f7871,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59))) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f1008,f859])).
fof(f7885,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f7870,f231])).
fof(f7870,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f1008,f801])).
fof(f7877,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7858,f231])).
fof(f7858,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1008,f651])).
fof(f7854,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f1008])).
fof(f7830,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5('const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X73)),X74) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f908,f845])).
fof(f7829,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X69),s(bool,X67)),X68)) )),
  inference(superposition,[],[f908,f277])).
fof(f7828,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,X62)),X63)) )),
  inference(superposition,[],[f908,f276])).
fof(f7844,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57)),X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7827,f231])).
fof(f7827,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X57)),X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f908,f859])).
fof(f7843,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52)),X53) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f7826,f231])).
fof(f7826,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52)),X53) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f908,f801])).
fof(f7835,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7814,f231])).
fof(f7814,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f908,f651])).
fof(f7809,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f908])).
fof(f7791,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP5('const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X73),s(bool,X75)),X74) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f907,f845])).
fof(f7790,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X69)),X68)) )),
  inference(superposition,[],[f907,f277])).
fof(f7789,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X60),s(bool,X61)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X64)),X63)) )),
  inference(superposition,[],[f907,f276])).
fof(f7805,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59)),X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7788,f231])).
fof(f7788,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP5('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X59)),X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f907,f859])).
fof(f7804,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54)),X53) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f7787,f231])).
fof(f7787,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5('const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X54)),X53) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f907,f801])).
fof(f7796,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7775,f231])).
fof(f7775,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f907,f651])).
fof(f7770,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f907])).
fof(f7764,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X13),s('type.extreal.extreal',X12))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13)))) )),
  inference(subsumption_resolution,[],[f7745,f420])).
fof(f7745,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X13),s('type.extreal.extreal',X12))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13)))) )),
  inference(superposition,[],[f1273,f239])).
fof(f7736,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',X14)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.bool.COND_3'(s(bool,X12),s('type.extreal.extreal',X13),s('type.extreal.extreal',X14))),s('type.extreal.extreal',X15))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1273,f845])).
fof(f7729,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X14),s('type.extreal.extreal',X15)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal','const.bool.COND_3'(s(bool,X12),s('type.extreal.extreal',X13),s('type.extreal.extreal',X14))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1273,f845])).
fof(f7711,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,X212)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X213),s(bool,X212)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X210),s(bool,X211)),X214) | s(bool,'F_0') = s(bool,X211)) )),
  inference(subsumption_resolution,[],[f7628,f231])).
fof(f7628,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,'T_0')) | ~p(s(bool,X212)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X213),s(bool,X212)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X210),s(bool,X211)),X214) | s(bool,'F_0') = s(bool,X211)) )),
  inference(superposition,[],[f875,f859])).
fof(f7708,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,X202)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X203),s(bool,X202)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X200),s(bool,X201)),X204) | s(bool,'F_0') = s(bool,X201)) )),
  inference(subsumption_resolution,[],[f7626,f231])).
fof(f7626,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X202)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X203),s(bool,X202)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X200),s(bool,X201)),X204) | s(bool,'F_0') = s(bool,X201)) )),
  inference(superposition,[],[f873,f859])).
fof(f7707,plain,(
  ( ! [X198,X196,X199,X197,X195] : (~p(s(bool,X197)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X198),s(bool,X197)),X199,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X195),s(bool,X196))) | s(bool,'F_0') = s(bool,X196)) )),
  inference(subsumption_resolution,[],[f7625,f231])).
fof(f7625,plain,(
  ( ! [X198,X196,X199,X197,X195] : (~p(s(bool,'T_0')) | ~p(s(bool,X197)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X198),s(bool,X197)),X199,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X195),s(bool,X196))) | s(bool,'F_0') = s(bool,X196)) )),
  inference(superposition,[],[f872,f859])).
fof(f7618,plain,(
  ( ! [X169,X170,X168] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X168),s(bool,X169)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X170),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X169)) )),
  inference(superposition,[],[f844,f859])).
fof(f7616,plain,(
  ( ! [X163,X164,X162] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X162),s(bool,X163)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X164),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X163)) )),
  inference(superposition,[],[f843,f859])).
fof(f7615,plain,(
  ( ! [X161,X159,X160] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X161),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X159),s(bool,X160))) | s(bool,'F_0') = s(bool,X160)) )),
  inference(superposition,[],[f842,f859])).
fof(f7613,plain,(
  ( ! [X154,X155,X153] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X155),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X153),s(bool,X154))) | s(bool,'F_0') = s(bool,X154)) )),
  inference(superposition,[],[f841,f859])).
fof(f7698,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,X147)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X147),s(bool,X148)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X146)),X149) | s(bool,'F_0') = s(bool,X146)) )),
  inference(subsumption_resolution,[],[f7611,f231])).
fof(f7611,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,X147)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X147),s(bool,X148)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X146)),X149) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f817,f859])).
fof(f7695,plain,(
  ( ! [X138,X136,X135,X139,X137] : (~p(s(bool,X137)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X135),s(bool,X136)),X139) | s(bool,'F_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f7609,f231])).
fof(f7609,plain,(
  ( ! [X138,X136,X135,X139,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,X137)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X137),s(bool,X138)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X135),s(bool,X136)),X139) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f815,f859])).
fof(f7694,plain,(
  ( ! [X134,X132,X130,X133,X131] : (~p(s(bool,X132)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X132),s(bool,X133)),X134,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X130),s(bool,X131))) | s(bool,'F_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f7608,f231])).
fof(f7608,plain,(
  ( ! [X134,X132,X130,X133,X131] : (~p(s(bool,'T_0')) | ~p(s(bool,X132)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X132),s(bool,X133)),X134,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X130),s(bool,X131))) | s(bool,'F_0') = s(bool,X131)) )),
  inference(superposition,[],[f814,f859])).
fof(f7600,plain,(
  ( ! [X103,X102,X104] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X102),s(bool,X103)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X104))) | s(bool,'F_0') = s(bool,X103)) )),
  inference(superposition,[],[f790,f859])).
fof(f7598,plain,(
  ( ! [X97,X98,X96] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X96),s(bool,X97)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X98))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f789,f859])).
fof(f7596,plain,(
  ( ! [X92,X90,X91] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X92)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X90),s(bool,X91))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f788,f859])).
fof(f7594,plain,(
  ( ! [X85,X86,X84] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X86)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X84),s(bool,X85))) | s(bool,'F_0') = s(bool,X85)) )),
  inference(superposition,[],[f787,f859])).
fof(f7685,plain,(
  ( ! [X70,X68,X69,X67] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68)),X69,X70) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68)),X69,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f7590,f231])).
fof(f7590,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68)),X69,X70) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68)),X69,X70) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f406,f859])).
fof(f7684,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64)),X65,X66) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64)),X65,X66) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f7589,f231])).
fof(f7589,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64)),X65,X66) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X63),s(bool,X64)),X65,X66) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f405,f859])).
fof(f7683,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X60)),X62,X61) | s(bool,'F_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f7588,f231])).
fof(f7588,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X60)),X62,X61) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f402,f859])).
fof(f7682,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X57,X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f7587,f231])).
fof(f7587,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X57,X58) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f401,f859])).
fof(f7681,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X50,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f7585,f231])).
fof(f7585,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X50,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f385,f859])).
fof(f7680,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X43),s(bool,X44)),X45,X46) | s(bool,'F_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f7584,f231])).
fof(f7584,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X43),s(bool,X44)),X45,X46) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f384,f859])).
fof(f7679,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)),X42,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7583,f231])).
fof(f7583,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)),X42,X41) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f383,f859])).
fof(f7678,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP7(X38,X37,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f7582,f231])).
fof(f7582,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP7(X38,X37,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(superposition,[],[f370,f859])).
fof(f7677,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X32)),X34) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f7581,f231])).
fof(f7581,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X32)),X34) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f334,f859])).
fof(f7674,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f7572,f231])).
fof(f7572,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f253,f859])).
fof(f7563,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,X40))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f859,f845])).
fof(f7562,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X37),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36)))) )),
  inference(superposition,[],[f859,f277])).
fof(f7561,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f859,f276])).
fof(f7535,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f859,f651])).
fof(f7482,plain,(
  ( ! [X344,X346,X345,X347,X348] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X347),s('type.extreal.extreal',X346)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.bool.COND_3'(s(bool,X344),s('type.extreal.extreal',X345),s('type.extreal.extreal',X346))),s('type.extreal.extreal',X348)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X347),s('type.extreal.extreal',X348)))) | s(bool,'T_0') = s(bool,X344)) )),
  inference(superposition,[],[f324,f845])).
fof(f7481,plain,(
  ( ! [X341,X343,X339,X340,X342] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X341),s('type.extreal.extreal',X342)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X342),s('type.extreal.extreal',X343)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.bool.COND_3'(s(bool,X339),s('type.extreal.extreal',X340),s('type.extreal.extreal',X341))),s('type.extreal.extreal',X343)))) | s(bool,'T_0') = s(bool,X339)) )),
  inference(superposition,[],[f324,f845])).
fof(f7478,plain,(
  ( ! [X327,X329,X328,X330] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X330),s('type.extreal.extreal',X329)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.bool.COND_3'(s(bool,X327),s('type.extreal.extreal',X328),s('type.extreal.extreal',X329))),s('type.extreal.extreal',X330)))) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f281,f845])).
fof(f7477,plain,(
  ( ! [X325,X323,X324,X326] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X325),s('type.extreal.extreal',X326)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X326),s('type.extreal.extreal','const.bool.COND_3'(s(bool,X323),s('type.extreal.extreal',X324),s('type.extreal.extreal',X325)))))) | s(bool,'T_0') = s(bool,X323)) )),
  inference(superposition,[],[f281,f845])).
fof(f7463,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (~p(s(bool,X271)) | ~p(s(bool,X272)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X273),s(bool,X272)),'const.bool.COND_3'(s(bool,X269),s(bool,X270),s(bool,X271)),X274) | s(bool,'T_0') = s(bool,X269)) )),
  inference(superposition,[],[f875,f845])).
fof(f7462,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (~p(s(bool,X265)) | p(s(bool,X266)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X267),s(bool,'const.bool.COND_3'(s(bool,X263),s(bool,X264),s(bool,X265)))),X268,X266) | s(bool,'T_0') = s(bool,X263)) )),
  inference(superposition,[],[f874,f845])).
fof(f7461,plain,(
  ( ! [X261,X257,X259,X260,X262,X258] : (~p(s(bool,X259)) | ~p(s(bool,X260)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X261),s(bool,X260)),'const.bool.COND_3'(s(bool,X257),s(bool,X258),s(bool,X259)),X262) | s(bool,'T_0') = s(bool,X257)) )),
  inference(superposition,[],[f873,f845])).
fof(f7460,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (~p(s(bool,X253)) | ~p(s(bool,X254)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X255),s(bool,X254)),X256,'const.bool.COND_3'(s(bool,X251),s(bool,X252),s(bool,X253))) | s(bool,'T_0') = s(bool,X251)) )),
  inference(superposition,[],[f872,f845])).
fof(f7459,plain,(
  ( ! [X249,X246,X250,X248,X247] : (~p(s(bool,X248)) | p(s(bool,X249)) | sP1(X249,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X250),s(bool,'const.bool.COND_3'(s(bool,X246),s(bool,X247),s(bool,X248))))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f871,f845])).
fof(f7458,plain,(
  ( ! [X243,X241,X244,X242,X245] : (~p(s(bool,X243)) | p(s(bool,X244)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X245),s(bool,'const.bool.COND_3'(s(bool,X241),s(bool,X242),s(bool,X243)))),X244) | s(bool,'T_0') = s(bool,X241)) )),
  inference(superposition,[],[f870,f845])).
fof(f7457,plain,(
  ( ! [X239,X237,X240,X238,X236] : (~p(s(bool,X238)) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X239),s(bool,'const.bool.COND_3'(s(bool,X236),s(bool,X237),s(bool,X238)))),X240) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f869,f845])).
fof(f7456,plain,(
  ( ! [X231,X235,X233,X234,X232] : (~p(s(bool,X233)) | sP0(X234,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X235),s(bool,'const.bool.COND_3'(s(bool,X231),s(bool,X232),s(bool,X233))))) | s(bool,'T_0') = s(bool,X231)) )),
  inference(superposition,[],[f868,f845])).
fof(f7454,plain,(
  ( ! [X225,X223,X226,X224] : (sP5('T_0','const.bool.COND_3'(s(bool,X223),s(bool,X224),s(bool,X225)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X226),s(bool,X225))) | s(bool,'T_0') = s(bool,X223)) )),
  inference(superposition,[],[f844,f845])).
fof(f7452,plain,(
  ( ! [X218,X216,X215,X217] : (~sP5('F_0','const.bool.COND_3'(s(bool,X215),s(bool,X216),s(bool,X217)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X218),s(bool,X217))) | s(bool,'T_0') = s(bool,X215)) )),
  inference(superposition,[],[f843,f845])).
fof(f7451,plain,(
  ( ! [X212,X213,X211,X214] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X214),s(bool,X213)),'const.bool.COND_3'(s(bool,X211),s(bool,X212),s(bool,X213))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f842,f845])).
fof(f7449,plain,(
  ( ! [X206,X204,X205,X203] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X206),s(bool,X205)),'const.bool.COND_3'(s(bool,X203),s(bool,X204),s(bool,X205))) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f841,f845])).
fof(f7447,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (~p(s(bool,X195)) | ~p(s(bool,X196)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X196),s(bool,X197)),'const.bool.COND_3'(s(bool,X193),s(bool,X194),s(bool,X195)),X198) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f817,f845])).
fof(f7446,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (~p(s(bool,X189)) | p(s(bool,X190)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X187),s(bool,X188),s(bool,X189))),s(bool,X191)),X192,X190) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f816,f845])).
fof(f7445,plain,(
  ( ! [X185,X182,X186,X184,X183,X181] : (~p(s(bool,X183)) | ~p(s(bool,X184)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X184),s(bool,X185)),'const.bool.COND_3'(s(bool,X181),s(bool,X182),s(bool,X183)),X186) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f815,f845])).
fof(f7444,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (~p(s(bool,X177)) | ~p(s(bool,X178)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X178),s(bool,X179)),X180,'const.bool.COND_3'(s(bool,X175),s(bool,X176),s(bool,X177))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f814,f845])).
fof(f7443,plain,(
  ( ! [X173,X171,X174,X172,X170] : (~p(s(bool,X172)) | p(s(bool,X173)) | sP1(X173,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X170),s(bool,X171),s(bool,X172))),s(bool,X174))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f813,f845])).
fof(f7442,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X167)) | p(s(bool,X168)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X165),s(bool,X166),s(bool,X167))),s(bool,X169)),X168) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f812,f845])).
fof(f7441,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X162)) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X160),s(bool,X161),s(bool,X162))),s(bool,X163)),X164) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f811,f845])).
fof(f7440,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X157)) | sP0(X158,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X155),s(bool,X156),s(bool,X157))),s(bool,X159))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f810,f845])).
fof(f7438,plain,(
  ( ! [X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X149),s(bool,X150))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X147),s(bool,X148),s(bool,X149))) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f801,f845])).
fof(f7436,plain,(
  ( ! [X142,X140,X141,X139] : (sP5('T_0','const.bool.COND_3'(s(bool,X139),s(bool,X140),s(bool,X141)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f790,f845])).
fof(f7434,plain,(
  ( ! [X134,X132,X133,X131] : (~sP5('F_0','const.bool.COND_3'(s(bool,X131),s(bool,X132),s(bool,X133)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X133),s(bool,X134))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f789,f845])).
fof(f7432,plain,(
  ( ! [X125,X123,X126,X124] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X125),s(bool,X126)),'const.bool.COND_3'(s(bool,X123),s(bool,X124),s(bool,X125))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f788,f845])).
fof(f7430,plain,(
  ( ! [X118,X116,X117,X115] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X117),s(bool,X118)),'const.bool.COND_3'(s(bool,X115),s(bool,X116),s(bool,X117))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f787,f845])).
fof(f7429,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,X113)) | s(bool,X114) = s(bool,'const.bool.COND_3'(s(bool,X111),s(bool,X112),s(bool,X113))) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f710,f845])).
fof(f7424,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,X92)) | ~sP8('const.bool.COND_3'(s(bool,X90),s(bool,X91),s(bool,X92)),X93,X94) | sP9('const.bool.COND_3'(s(bool,X90),s(bool,X91),s(bool,X92)),X93,X94) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f406,f845])).
fof(f7423,plain,(
  ( ! [X88,X87,X85,X89,X86] : (~p(s(bool,X87)) | ~sP10('const.bool.COND_3'(s(bool,X85),s(bool,X86),s(bool,X87)),X88,X89) | sP11('const.bool.COND_3'(s(bool,X85),s(bool,X86),s(bool,X87)),X88,X89) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f405,f845])).
fof(f7422,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82)),X84,X83) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f402,f845])).
fof(f7421,plain,(
  ( ! [X78,X76,X79,X77,X75] : (~p(s(bool,X77)) | ~p(s(bool,X78)) | sP11('const.bool.COND_3'(s(bool,X75),s(bool,X76),s(bool,X77)),X78,X79) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f401,f845])).
fof(f7420,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | sP11('const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72)),X74,X73) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f400,f845])).
fof(f7419,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('const.bool.COND_3'(s(bool,X65),s(bool,X66),s(bool,X67)),X69,X68) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f385,f845])).
fof(f7418,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | sP9('const.bool.COND_3'(s(bool,X60),s(bool,X61),s(bool,X62)),X63,X64) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f384,f845])).
fof(f7417,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | sP9('const.bool.COND_3'(s(bool,X55),s(bool,X56),s(bool,X57)),X59,X58) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f383,f845])).
fof(f7416,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'const.bool.COND_3'(s(bool,X50),s(bool,X51),s(bool,X52))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f370,f845])).
fof(f7415,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,'const.bool.COND_3'(s(bool,X45),s(bool,X46),s(bool,X47)),X49) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f334,f845])).
fof(f7414,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | sP1(X44,'const.bool.COND_3'(s(bool,X41),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f310,f845])).
fof(f7413,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | sP1('const.bool.COND_3'(s(bool,X37),s(bool,X38),s(bool,X39)),X40) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f307,f845])).
fof(f7412,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | sP0('const.bool.COND_3'(s(bool,X33),s(bool,X34),s(bool,X35)),X36) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f292,f845])).
fof(f7411,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X31)) | sP0(X32,'const.bool.COND_3'(s(bool,X29),s(bool,X30),s(bool,X31))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f291,f845])).
fof(f7410,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X26),s(bool,X27))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f285,f845])).
fof(f7408,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) = s(bool,'const.bool.COND_3'(s(bool,X17),s(bool,X18),s(bool,X19))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f282,f845])).
fof(f7407,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X14),s(bool,X15),s(bool,X16))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f254,f845])).
fof(f7406,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X11),s(bool,X12),s(bool,X13))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f253,f845])).
fof(f7405,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X10)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f250,f845])).
fof(f7404,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.COND_3'(s(bool,X5),s(bool,X6),s(bool,X7)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f249,f845])).
fof(f7399,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X18),s(bool,X19),s(bool,X20)))) )),
  inference(superposition,[],[f845,f239])).
fof(f7397,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X10),s(bool,X11),s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f845,f651])).
fof(f7395,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X2) | s(bool,X5) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f845,f651])).
fof(f7353,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(X66,X67) = s(X66,'const.bool.COND_3'(s(bool,X65),s(X66,X68),s(X66,X67))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f845,f845])).
fof(f7352,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(X60,X61) = s(X60,'const.bool.COND_3'(s(bool,X59),s(X60,X62),s(X60,X61))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X59)))) )),
  inference(superposition,[],[f845,f277])).
fof(f7351,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(X55,X56) = s(X55,'const.bool.COND_3'(s(bool,X53),s(X55,X57),s(X55,X56))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X53),s(bool,X54)))) )),
  inference(superposition,[],[f845,f276])).
fof(f7324,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,X201)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,X201)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X199),s(bool,X200)),X203) | s(bool,'F_0') = s(bool,X199)) )),
  inference(subsumption_resolution,[],[f7238,f231])).
fof(f7238,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X201)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,X201)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X199),s(bool,X200)),X203) | s(bool,'F_0') = s(bool,X199)) )),
  inference(superposition,[],[f875,f801])).
fof(f7323,plain,(
  ( ! [X198,X196,X194,X197] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X197),s(bool,X194)),X198,X196) | p(s(bool,X196)) | s(bool,'F_0') = s(bool,X194)) )),
  inference(forward_demodulation,[],[f7322,f276])).
fof(f7322,plain,(
  ( ! [X198,X196,X194,X197,X195] : (p(s(bool,X196)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X197),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X194),s(bool,X195)))),X198,X196) | s(bool,'F_0') = s(bool,X194)) )),
  inference(subsumption_resolution,[],[f7237,f231])).
fof(f7237,plain,(
  ( ! [X198,X196,X194,X197,X195] : (~p(s(bool,'T_0')) | p(s(bool,X196)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X197),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X194),s(bool,X195)))),X198,X196) | s(bool,'F_0') = s(bool,X194)) )),
  inference(superposition,[],[f874,f801])).
fof(f7321,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,X191)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X192),s(bool,X191)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X189),s(bool,X190)),X193) | s(bool,'F_0') = s(bool,X189)) )),
  inference(subsumption_resolution,[],[f7236,f231])).
fof(f7236,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,'T_0')) | ~p(s(bool,X191)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X192),s(bool,X191)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X189),s(bool,X190)),X193) | s(bool,'F_0') = s(bool,X189)) )),
  inference(superposition,[],[f873,f801])).
fof(f7320,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,X186)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X187),s(bool,X186)),X188,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X184),s(bool,X185))) | s(bool,'F_0') = s(bool,X184)) )),
  inference(subsumption_resolution,[],[f7235,f231])).
fof(f7235,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,'T_0')) | ~p(s(bool,X186)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X187),s(bool,X186)),X188,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X184),s(bool,X185))) | s(bool,'F_0') = s(bool,X184)) )),
  inference(superposition,[],[f872,f801])).
fof(f7319,plain,(
  ( ! [X182,X180,X183] : (sP1(X182,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X183),s(bool,X180))) | p(s(bool,X182)) | s(bool,'F_0') = s(bool,X180)) )),
  inference(forward_demodulation,[],[f7318,f276])).
fof(f7318,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X182)) | sP1(X182,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X183),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X180),s(bool,X181))))) | s(bool,'F_0') = s(bool,X180)) )),
  inference(subsumption_resolution,[],[f7234,f231])).
fof(f7234,plain,(
  ( ! [X182,X180,X183,X181] : (~p(s(bool,'T_0')) | p(s(bool,X182)) | sP1(X182,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X183),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X180),s(bool,X181))))) | s(bool,'F_0') = s(bool,X180)) )),
  inference(superposition,[],[f871,f801])).
fof(f7317,plain,(
  ( ! [X178,X176,X179] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,X176)),X178) | p(s(bool,X178)) | s(bool,'F_0') = s(bool,X176)) )),
  inference(forward_demodulation,[],[f7316,f276])).
fof(f7316,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X178)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X176),s(bool,X177)))),X178) | s(bool,'F_0') = s(bool,X176)) )),
  inference(subsumption_resolution,[],[f7233,f231])).
fof(f7233,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'T_0')) | p(s(bool,X178)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X179),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X176),s(bool,X177)))),X178) | s(bool,'F_0') = s(bool,X176)) )),
  inference(superposition,[],[f870,f801])).
fof(f7315,plain,(
  ( ! [X175,X174,X172] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X174),s(bool,X172)),X175) | s(bool,'F_0') = s(bool,X172)) )),
  inference(forward_demodulation,[],[f7314,f276])).
fof(f7314,plain,(
  ( ! [X175,X173,X174,X172] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X174),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X172),s(bool,X173)))),X175) | s(bool,'F_0') = s(bool,X172)) )),
  inference(subsumption_resolution,[],[f7232,f231])).
fof(f7232,plain,(
  ( ! [X175,X173,X174,X172] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X174),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X172),s(bool,X173)))),X175) | s(bool,'F_0') = s(bool,X172)) )),
  inference(superposition,[],[f869,f801])).
fof(f7313,plain,(
  ( ! [X171,X170,X168] : (sP0(X170,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X171),s(bool,X168))) | s(bool,'F_0') = s(bool,X168)) )),
  inference(forward_demodulation,[],[f7312,f276])).
fof(f7312,plain,(
  ( ! [X171,X169,X170,X168] : (sP0(X170,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X171),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X168),s(bool,X169))))) | s(bool,'F_0') = s(bool,X168)) )),
  inference(subsumption_resolution,[],[f7231,f231])).
fof(f7231,plain,(
  ( ! [X171,X169,X170,X168] : (~p(s(bool,'T_0')) | sP0(X170,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X171),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X168),s(bool,X169))))) | s(bool,'F_0') = s(bool,X168)) )),
  inference(superposition,[],[f868,f801])).
fof(f7229,plain,(
  ( ! [X163,X164,X162] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X162),s(bool,X163)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X164),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X162)) )),
  inference(superposition,[],[f844,f801])).
fof(f7227,plain,(
  ( ! [X158,X156,X157] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X156),s(bool,X157)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X158),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f843,f801])).
fof(f7226,plain,(
  ( ! [X154,X155,X153] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X155),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X153),s(bool,X154))) | s(bool,'F_0') = s(bool,X153)) )),
  inference(superposition,[],[f842,f801])).
fof(f7224,plain,(
  ( ! [X149,X147,X148] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X149),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X147),s(bool,X148))) | s(bool,'F_0') = s(bool,X147)) )),
  inference(superposition,[],[f841,f801])).
fof(f7311,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,X141)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X139),s(bool,X140)),X143) | s(bool,'F_0') = s(bool,X139)) )),
  inference(subsumption_resolution,[],[f7222,f231])).
fof(f7222,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | ~p(s(bool,X141)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X141),s(bool,X142)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X139),s(bool,X140)),X143) | s(bool,'F_0') = s(bool,X139)) )),
  inference(superposition,[],[f817,f801])).
fof(f7308,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X131)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X131),s(bool,X132)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X129),s(bool,X130)),X133) | s(bool,'F_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f7220,f231])).
fof(f7220,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X131)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X131),s(bool,X132)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X129),s(bool,X130)),X133) | s(bool,'F_0') = s(bool,X129)) )),
  inference(superposition,[],[f815,f801])).
fof(f7307,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X126)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X126),s(bool,X127)),X128,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X124),s(bool,X125))) | s(bool,'F_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f7219,f231])).
fof(f7219,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | ~p(s(bool,X126)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X126),s(bool,X127)),X128,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X124),s(bool,X125))) | s(bool,'F_0') = s(bool,X124)) )),
  inference(superposition,[],[f814,f801])).
fof(f7213,plain,(
  ( ! [X103,X102,X104] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X102),s(bool,X103)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X104))) | s(bool,'F_0') = s(bool,X102)) )),
  inference(superposition,[],[f790,f801])).
fof(f7211,plain,(
  ( ! [X97,X98,X96] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X96),s(bool,X97)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X98))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f789,f801])).
fof(f7209,plain,(
  ( ! [X92,X90,X91] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X92)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X90),s(bool,X91))) | s(bool,'F_0') = s(bool,X90)) )),
  inference(superposition,[],[f788,f801])).
fof(f7207,plain,(
  ( ! [X85,X86,X84] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X86)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85))) | s(bool,'F_0') = s(bool,X84)) )),
  inference(superposition,[],[f787,f801])).
fof(f7298,plain,(
  ( ! [X70,X68,X69,X67] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X68)),X69,X70) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X68)),X69,X70) | s(bool,'F_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f7203,f231])).
fof(f7203,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X68)),X69,X70) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X68)),X69,X70) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f406,f801])).
fof(f7297,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)),X65,X66) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)),X65,X66) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f7202,f231])).
fof(f7202,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)),X65,X66) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X63),s(bool,X64)),X65,X66) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f405,f801])).
fof(f7296,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X59),s(bool,X60)),X62,X61) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f7201,f231])).
fof(f7201,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X59),s(bool,X60)),X62,X61) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f402,f801])).
fof(f7295,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56)),X57,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f7200,f231])).
fof(f7200,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56)),X57,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f401,f801])).
fof(f7294,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X53)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X54,X53) | s(bool,'F_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f7199,f231])).
fof(f7199,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X54,X53) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f400,f801])).
fof(f7293,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X50,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f7198,f231])).
fof(f7198,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X50,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f385,f801])).
fof(f7292,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f7197,f231])).
fof(f7197,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f384,f801])).
fof(f7291,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f7196,f231])).
fof(f7196,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f383,f801])).
fof(f7290,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP7(X38,X37,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f7195,f231])).
fof(f7195,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP7(X38,X37,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f370,f801])).
fof(f7289,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34) | s(bool,'F_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f7194,f231])).
fof(f7194,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f334,f801])).
fof(f7288,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7193,f231])).
fof(f7193,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f310,f801])).
fof(f7287,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f7192,f231])).
fof(f7192,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f307,f801])).
fof(f7286,plain,(
  ( ! [X24,X23,X22] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f7191,f231])).
fof(f7191,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f292,f801])).
fof(f7285,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f7190,f231])).
fof(f7190,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(superposition,[],[f291,f801])).
fof(f7284,plain,(
  ( ! [X18,X16] : (s(bool,X16) = s(bool,X18) | ~p(s(bool,X18)) | s(bool,'F_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f7283,f276])).
fof(f7283,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f7189,f231])).
fof(f7189,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X16),s(bool,X17))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f285,f801])).
fof(f7282,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7185,f231])).
fof(f7185,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f253,f801])).
fof(f7162,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X33),s(bool,X34))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f801,f277])).
fof(f7161,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X31))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X30)))) )),
  inference(superposition,[],[f801,f276])).
fof(f7142,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f801,f651])).
fof(f7126,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X59)) )),
  inference(superposition,[],[f875,f277])).
fof(f7125,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X52)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),X54)) )),
  inference(superposition,[],[f875,f276])).
fof(f7127,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7113,f231])).
fof(f7113,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f875,f651])).
fof(f7112,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f875])).
fof(f7096,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7082,f231])).
fof(f7082,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X0)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f874,f651])).
fof(f7081,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f874])).
fof(f7068,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X59)) )),
  inference(superposition,[],[f873,f277])).
fof(f7067,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X52)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),X54)) )),
  inference(superposition,[],[f873,f276])).
fof(f7069,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7055,f231])).
fof(f7055,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f873,f651])).
fof(f7054,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f873])).
fof(f7017,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X57)),X59,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)))) )),
  inference(superposition,[],[f872,f277])).
fof(f7016,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X52)),X54,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)))) )),
  inference(superposition,[],[f872,f276])).
fof(f7018,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7004,f231])).
fof(f7004,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f872,f651])).
fof(f7003,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f231,f872])).
fof(f6965,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X59)) )),
  inference(superposition,[],[f817,f277])).
fof(f6964,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),X54)) )),
  inference(superposition,[],[f817,f276])).
fof(f6966,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6952,f231])).
fof(f6952,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f817,f651])).
fof(f6951,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f817])).
fof(f6933,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6919,f231])).
fof(f6919,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2)),X3,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f816,f651])).
fof(f6918,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f816])).
fof(f6905,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X59)) )),
  inference(superposition,[],[f815,f277])).
fof(f6904,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)),X54)) )),
  inference(superposition,[],[f815,f276])).
fof(f6906,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6892,f231])).
fof(f6892,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X0,X3) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f815,f651])).
fof(f6891,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f815])).
fof(f6862,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X57),s(bool,X58)),X59,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)))) )),
  inference(superposition,[],[f814,f277])).
fof(f6861,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),X54,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X50),s(bool,X51)))) )),
  inference(superposition,[],[f814,f276])).
fof(f6863,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6849,f231])).
fof(f6849,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,X0) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f814,f651])).
fof(f6848,plain,(
  ( ! [X0,X1] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f231,f814])).
fof(f253,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.extreal.le_max','thm.bool.EQ_CLAUSES|split|3')).
fof(f254,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f159])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12] : ! [X9,X10] : s(X12,X10) = s(X12,'const.bool.COND_3'(s(bool,'F_0'),s(X12,X9),s(X12,X10)))),
  file('thm.extreal.le_max','thm.bool.COND_CLAUSES|split|1')).
fof(f801,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f276,f239])).
fof(f841,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f420,f277,f353])).
fof(f842,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f231,f277,f351])).
fof(f843,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f420,f277,f353])).
fof(f844,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f277,f351])).
fof(f859,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f277,f239])).
fof(f868,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | sP0(X21,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X19),s(bool,X20)))) )),
  inference(superposition,[],[f291,f277])).
fof(f869,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X23)) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X22),s(bool,X23)),X24)) )),
  inference(superposition,[],[f292,f277])).
fof(f870,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26)),X27)) )),
  inference(superposition,[],[f307,f277])).
fof(f871,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f310,f277])).
fof(f872,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,X32)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X32)),X34,X33)) )),
  inference(superposition,[],[f383,f277])).
fof(f873,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,X36)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36)),X37,X38)) )),
  inference(superposition,[],[f384,f277])).
fof(f874,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f400,f277])).
fof(f875,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X44)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f401,f277])).
fof(f876,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50)) )),
  inference(superposition,[],[f405,f277])).
fof(f877,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f406,f277])).
fof(f908,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X16)) | ~p(s(bool,X15)) | p(s(bool,X17)) | ~sP5(X16,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)),X17)) )),
  inference(superposition,[],[f334,f277])).
fof(f1009,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X16)) | ~p(s(bool,X15)) | p(s(bool,X17)) | ~sP7(X17,X16,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)))) )),
  inference(superposition,[],[f370,f277])).
fof(f1080,plain,(
  ( ! [X4,X5] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X5)),X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f787,f239])).
fof(f1082,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) )),
  inference(superposition,[],[f787,f277])).
fof(f1101,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)),X17,X16)) )),
  inference(superposition,[],[f385,f277])).
fof(f1157,plain,(
  ( ! [X4,X5] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X5)),X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f788,f239])).
fof(f1159,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) )),
  inference(superposition,[],[f788,f277])).
fof(f1228,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X16)) | ~p(s(bool,X15)) | p(s(bool,X17)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)),X17,X16)) )),
  inference(superposition,[],[f402,f277])).
fof(f1639,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f787,f428,f434,f420,f337])).
fof(f1746,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,bool))) )),
  inference(superposition,[],[f323,f239])).
fof(f1747,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK15(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK15(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f323,f239])).
fof(f1812,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X2)),X2)) )),
  inference(unit_resulting_resolution,[],[f277,f1253,f351])).
fof(f1816,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f277,f1253,f351])).
fof(f2017,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f436,f430,f789,f336])).
fof(f2023,plain,(
  ( ! [X4,X5] : (~sP5('F_0',X4,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f789,f239])).
fof(f2025,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f789,f277])).
fof(f2052,plain,(
  ( ! [X4,X5] : (sP5('T_0',X4,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f790,f239])).
fof(f2054,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f790,f277])).
fof(f2084,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f434,f428,f841,f337])).
fof(f2086,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f442,f430,f841,f335])).
fof(f2100,plain,(
  ( ! [X4,X5] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'F_0')),X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f841,f239])).
fof(f2102,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) )),
  inference(superposition,[],[f841,f277])).
fof(f2125,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f842,f333])).
fof(f2139,plain,(
  ( ! [X4,X5] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'F_0')),X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f842,f239])).
fof(f2141,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X10)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) )),
  inference(superposition,[],[f842,f277])).
fof(f2161,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f440,f426,f843,f337])).
fof(f2163,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f436,f430,f843,f336])).
fof(f2177,plain,(
  ( ! [X4,X5] : (~sP5('F_0',X4,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f843,f239])).
fof(f2198,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) )),
  inference(unit_resulting_resolution,[],[f844,f333])).
fof(f2213,plain,(
  ( ! [X4,X5] : (sP5('T_0',X4,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f844,f239])).
fof(f2285,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f231,f868])).
fof(f2306,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f231,f869])).
fof(f2433,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f870])).
fof(f2460,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f871])).
fof(f2506,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26)))) )),
  inference(superposition,[],[f651,f277])).
fof(f2531,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X18),s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f651,f277])).
fof(f2581,plain,(
  ( ! [X118,X119,X117] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X119),s(bool,X118)),X117) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f841,f651])).
fof(f2583,plain,(
  ( ! [X125,X123,X124] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X125),s(bool,X124)),X123) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f842,f651])).
fof(f2584,plain,(
  ( ! [X127,X128,X126] : (~sP5('F_0',X126,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X128),s(bool,X127))) | s(bool,'T_0') = s(bool,X127) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f843,f651])).
fof(f2586,plain,(
  ( ! [X134,X132,X133] : (sP5('T_0',X132,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X134),s(bool,X133))) | s(bool,'T_0') = s(bool,X133) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f844,f651])).
fof(f2789,plain,(
  ( ! [X173,X171,X172] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X172),s(bool,X171)),'T_0') | s(bool,X171) = s(bool,X173) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f841,f651])).
fof(f2791,plain,(
  ( ! [X177,X178,X179] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X178),s(bool,X177)),'T_0') | s(bool,X177) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f842,f651])).
fof(f2792,plain,(
  ( ! [X182,X180,X181] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X181),s(bool,X180))) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f843,f651])).
fof(f2794,plain,(
  ( ! [X187,X188,X186] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X187),s(bool,X186))) | s(bool,X186) = s(bool,X188) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f844,f651])).
fof(f2849,plain,(
  ( ! [X118,X119,X117] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X118),s(bool,'T_0')),X117) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f841,f651])).
fof(f2851,plain,(
  ( ! [X125,X123,X124] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X124),s(bool,'T_0')),X123) | s(bool,X123) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f842,f651])).
fof(f2852,plain,(
  ( ! [X127,X128,X126] : (~sP5('F_0',X126,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X127),s(bool,'T_0'))) | s(bool,X126) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f843,f651])).
fof(f2854,plain,(
  ( ! [X134,X132,X133] : (sP5('T_0',X132,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X133),s(bool,'T_0'))) | s(bool,X132) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f844,f651])).
fof(f3387,plain,(
  ( ! [X140,X138,X141,X139] : (sP0(X139,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X140),s(bool,X138))) | s(bool,X138) = s(bool,X141) | s(bool,'T_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f2856,f231])).
fof(f2856,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,'T_0')) | sP0(X139,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X140),s(bool,X138))) | s(bool,X138) = s(bool,X141) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f868,f651])).
fof(f3388,plain,(
  ( ! [X144,X142,X145,X143] : (sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X143),s(bool,X142)),X144) | s(bool,X142) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(subsumption_resolution,[],[f2857,f231])).
fof(f2857,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'F_0'),s(bool,X143),s(bool,X142)),X144) | s(bool,X142) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f869,f651])).
fof(f2905,plain,(
  ( ! [X277,X275,X276,X274] : (s(bool,'T_0') != s(bool,f(s(X274,sK15(X275,X276,bool,X274)))) | s(fun(X274,bool),X275) = s(fun(X274,bool),X276) | s(bool,X277) = s(bool,g(s(X274,sK15(X275,X276,bool,X274)))) | s(bool,'T_0') = s(bool,X277)) )),
  inference(superposition,[],[f323,f651])).
fof(f2638,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X281) != s(bool,f(s(X278,sK15(X279,X280,bool,X278)))) | s(fun(X278,bool),X279) = s(fun(X278,bool),X280) | s(bool,'T_0') = s(bool,X281) | s(bool,'T_0') = s(bool,g(s(X278,sK15(X279,X280,bool,X278))))) )),
  inference(superposition,[],[f323,f651])).
fof(f280,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7,X8] : s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))),s('type.extreal.extreal',X8),s('type.extreal.extreal',X7)))),
  file('thm.extreal.le_max','thm.extreal.extreal_max_def')).
fof(f281,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7,X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.le_max','thm.extreal.le_total')).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7,X8,X6] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X6))))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_max','thm.extreal.le_trans')).
fof(f1273,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(subsumption_resolution,[],[f1260,f420])).
fof(f1260,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f281,f239])).
fof(f1695,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(forward_demodulation,[],[f1678,f277])).
fof(f1678,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'F_0'),s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f280,f239])).
fof(f2631,plain,(
  ( ! [X255,X253,X254] : (s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X253),s('type.extreal.extreal',X254))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X255),s('type.extreal.extreal',X254),s('type.extreal.extreal',X253))) | s(bool,'T_0') = s(bool,X255) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X253),s('type.extreal.extreal',X254)))) )),
  inference(superposition,[],[f280,f651])).
fof(f3390,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X247),s('type.extreal.extreal',X248)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X246),s('type.extreal.extreal',X248)))) | s(bool,X249) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X246),s('type.extreal.extreal',X247))) | s(bool,'T_0') = s(bool,X249)) )),
  inference(subsumption_resolution,[],[f2897,f231])).
fof(f2897,plain,(
  ( ! [X249,X246,X248,X247] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X247),s('type.extreal.extreal',X248)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X246),s('type.extreal.extreal',X248)))) | s(bool,X249) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X246),s('type.extreal.extreal',X247))) | s(bool,'T_0') = s(bool,X249)) )),
  inference(superposition,[],[f324,f651])).
fof(f3391,plain,(
  ( ! [X251,X252,X250] : (s('type.extreal.extreal',X251) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X250),s('type.extreal.extreal',X251))) | s(bool,X252) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X250),s('type.extreal.extreal',X251))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(forward_demodulation,[],[f2898,f276])).
fof(f2898,plain,(
  ( ! [X251,X252,X250] : (s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X250),s('type.extreal.extreal',X251))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',X251),s('type.extreal.extreal',X250))) | s(bool,X252) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X250),s('type.extreal.extreal',X251))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f280,f651])).
fof(f3605,plain,(
  ( ! [X146,X149,X147,X148] : (p(s(bool,X147)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X148),s(bool,X146)),X147) | s(bool,X146) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f3028,f231])).
fof(f3028,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,'T_0')) | p(s(bool,X147)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X148),s(bool,X146)),X147) | s(bool,X146) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f870,f651])).
fof(f3604,plain,(
  ( ! [X152,X151,X153,X150] : (p(s(bool,X151)) | sP1(X151,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X152),s(bool,X150))) | s(bool,X150) = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f3029,f231])).
fof(f3029,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X151)) | sP1(X151,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X152),s(bool,X150))) | s(bool,X150) = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f871,f651])).
fof(f2908,plain,(
  ( ! [X285,X287,X284,X286] : (s(X286,f(s(bool,'T_0'))) != s(X286,g(s(bool,'T_0'))) | s(fun(bool,X286),X284) = s(fun(bool,X286),X285) | s(bool,X287) = s(bool,sK15(X284,X285,X286,bool)) | s(bool,'T_0') = s(bool,X287)) )),
  inference(superposition,[],[f323,f651])).
fof(f3386,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X111)) | sP1(X111,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X110),s(bool,X112))) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f2847,f231])).
fof(f2847,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,X111)) | sP1(X111,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X110),s(bool,X112))) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f813,f651])).
fof(f3385,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X107)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X106),s(bool,X108)),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f2846,f231])).
fof(f2846,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,X107)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X106),s(bool,X108)),X107) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f812,f651])).
fof(f3384,plain,(
  ( ! [X103,X105,X102,X104] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X102),s(bool,X103)),X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f2845,f231])).
fof(f2845,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X102),s(bool,X103)),X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f811,f651])).
fof(f3383,plain,(
  ( ! [X101,X99,X100,X98] : (sP0(X99,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X98),s(bool,X100))) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f2844,f231])).
fof(f2844,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | sP0(X99,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X98),s(bool,X100))) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f810,f651])).
fof(f2842,plain,(
  ( ! [X94,X92,X93] : (sP5('T_0',X92,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X93))) | s(bool,X92) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f790,f651])).
fof(f2840,plain,(
  ( ! [X88,X87,X86] : (~sP5('F_0',X86,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X87))) | s(bool,X86) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f789,f651])).
fof(f2838,plain,(
  ( ! [X80,X81,X82] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X81)),X80) | s(bool,X80) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f788,f651])).
fof(f2836,plain,(
  ( ! [X76,X74,X75] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X75)),X74) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f787,f651])).
fof(f3382,plain,(
  ( ! [X70,X68,X69,X67] : (~sP8(X67,X68,X69) | sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2834,f231])).
fof(f2834,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~sP8(X67,X68,X69) | sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f406,f651])).
fof(f3381,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X63,X64,X65) | sP11(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2833,f231])).
fof(f2833,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~sP10(X63,X64,X65) | sP11(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f405,f651])).
fof(f3380,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP11(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2832,f231])).
fof(f2832,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP11(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f402,f651])).
fof(f3379,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP11(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2831,f231])).
fof(f2831,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP11(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f401,f651])).
fof(f3378,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | sP11(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2830,f231])).
fof(f2830,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | sP11(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f400,f651])).
fof(f3377,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP9(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2829,f231])).
fof(f2829,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP9(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f385,f651])).
fof(f3376,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | sP9(X43,X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2828,f231])).
fof(f2828,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | sP9(X43,X44,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f384,f651])).
fof(f3375,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | sP9(X39,X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2827,f231])).
fof(f2827,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP9(X39,X41,X40) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f383,f651])).
fof(f3374,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7(X37,X36,X35) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2826,f231])).
fof(f2826,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7(X37,X36,X35) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f370,f651])).
fof(f3373,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2825,f231])).
fof(f2825,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f334,f651])).
fof(f3372,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2824,f231])).
fof(f2824,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f310,f651])).
fof(f3371,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2823,f231])).
fof(f2823,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f307,f651])).
fof(f3370,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2822,f231])).
fof(f2822,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f292,f651])).
fof(f3369,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2821,f231])).
fof(f2821,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f291,f651])).
fof(f3368,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2820,f231])).
fof(f2820,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f285,f651])).
fof(f3367,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2816,f231])).
fof(f2816,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f253,f651])).
fof(f2782,plain,(
  ( ! [X146,X147,X148] : (sP5('T_0','T_0','const.bool.COND_3'(s(bool,X146),s(bool,X146),s(bool,X147))) | s(bool,X146) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f790,f651])).
fof(f2780,plain,(
  ( ! [X142,X140,X141] : (~sP5('F_0','T_0','const.bool.COND_3'(s(bool,X140),s(bool,X140),s(bool,X141))) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f789,f651])).
fof(f2778,plain,(
  ( ! [X134,X136,X135] : (sP5('T_0','const.bool.COND_3'(s(bool,X134),s(bool,X134),s(bool,X135)),'T_0') | s(bool,X134) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f788,f651])).
fof(f2776,plain,(
  ( ! [X130,X128,X129] : (~sP5('F_0','const.bool.COND_3'(s(bool,X128),s(bool,X128),s(bool,X129)),'T_0') | s(bool,X128) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f787,f651])).
fof(f2745,plain,(
  ( ! [X26,X24,X27,X25] : (sP5('T_0',X25,'const.bool.COND_3'(s(bool,X24),s(bool,X25),s(bool,X26))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f790,f651])).
fof(f2744,plain,(
  ( ! [X23,X21,X22,X20] : (~sP5('F_0',X21,'const.bool.COND_3'(s(bool,X20),s(bool,X21),s(bool,X22))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f789,f651])).
fof(f2743,plain,(
  ( ! [X19,X17,X18,X16] : (sP5('T_0','const.bool.COND_3'(s(bool,X16),s(bool,X17),s(bool,X18)),X17) | s(bool,X16) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f788,f651])).
fof(f2742,plain,(
  ( ! [X14,X12,X15,X13] : (~sP5('F_0','const.bool.COND_3'(s(bool,X12),s(bool,X13),s(bool,X14)),X13) | s(bool,X12) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f787,f651])).
fof(f2741,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X5),s(X6,X7),s(X6,X8))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f276,f651])).
fof(f2740,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f249,f651])).
fof(f2641,plain,(
  ( ! [X288,X290,X289,X291] : (s(X290,f(s(bool,X291))) != s(X290,g(s(bool,X291))) | s(fun(bool,X290),X288) = s(fun(bool,X290),X289) | s(bool,'T_0') = s(bool,X291) | s(bool,'T_0') = s(bool,sK15(X288,X289,X290,bool))) )),
  inference(superposition,[],[f323,f651])).
fof(f2574,plain,(
  ( ! [X94,X92,X93] : (sP5('T_0',X92,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X93),s(bool,X94))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f790,f651])).
fof(f2572,plain,(
  ( ! [X88,X87,X86] : (~sP5('F_0',X86,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X87),s(bool,X88))) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f789,f651])).
fof(f2570,plain,(
  ( ! [X80,X81,X82] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X81),s(bool,X82)),X80) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f788,f651])).
fof(f2568,plain,(
  ( ! [X76,X74,X75] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X75),s(bool,X76)),X74) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f787,f651])).
fof(f2530,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X17) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f651,f276])).
fof(f2519,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f651,f651])).
fof(f2505,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)))) )),
  inference(superposition,[],[f651,f276])).
fof(f2411,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f813])).
fof(f2387,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f812])).
fof(f2255,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f231,f811])).
fof(f2233,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f231,f810])).
fof(f2212,plain,(
  ( ! [X2,X3] : (sP5('T_0','F_0','const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f844,f239])).
fof(f2202,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f844,f239])).
fof(f2176,plain,(
  ( ! [X2,X3] : (~sP5('F_0','F_0','const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f843,f239])).
fof(f2166,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f843,f239])).
fof(f2138,plain,(
  ( ! [X2,X3] : (sP5('T_0','const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X2)),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f842,f239])).
fof(f2128,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f842,f239])).
fof(f2099,plain,(
  ( ! [X2,X3] : (~sP5('F_0','const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X2)),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f841,f239])).
fof(f2089,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f841,f239])).
fof(f2051,plain,(
  ( ! [X2,X3] : (sP5('T_0','F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f790,f239])).
fof(f2045,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f790,f333])).
fof(f2022,plain,(
  ( ! [X2,X3] : (~sP5('F_0','F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f789,f239])).
fof(f2015,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f440,f426,f789,f337])).
fof(f1758,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f285])).
fof(f1759,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f285])).
fof(f1760,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f291])).
fof(f1761,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1253,f292])).
fof(f1762,plain,(
  ( ! [X0] : (~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f306])).
fof(f1764,plain,(
  ( ! [X0] : (~sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f306])).
fof(f1765,plain,(
  ( ! [X0,X1] : (~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f306])).
fof(f1766,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f307])).
fof(f1767,plain,(
  ( ! [X0] : (sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f310])).
fof(f1770,plain,(
  ( ! [X0] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f334])).
fof(f1771,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f334])).
fof(f1773,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f438,f426,f432,f1253,f335])).
fof(f1780,plain,(
  ( ! [X0] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f338])).
fof(f1786,plain,(
  ( ! [X0] : (~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f338])).
fof(f1787,plain,(
  ( ! [X0,X1] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f338])).
fof(f1792,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f339])).
fof(f1793,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f342])).
fof(f1799,plain,(
  ( ! [X0] : (~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f342])).
fof(f1800,plain,(
  ( ! [X0,X1] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f342])).
fof(f1805,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f343])).
fof(f1806,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f347])).
fof(f1807,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f348])).
fof(f1808,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f349])).
fof(f1811,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f276,f1253,f351])).
fof(f1815,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)))) )),
  inference(unit_resulting_resolution,[],[f276,f1253,f351])).
fof(f1817,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f462,f1253,f352])).
fof(f1819,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f462,f1253,f352])).
fof(f1821,plain,(
  ( ! [X0,X1] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f355])).
fof(f1822,plain,(
  ( ! [X0,X1] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f356])).
fof(f1833,plain,(
  ( ! [X0] : (~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f361])).
fof(f1839,plain,(
  ( ! [X0] : (~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f361])).
fof(f1840,plain,(
  ( ! [X0,X1] : (~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f361])).
fof(f1845,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f362])).
fof(f1846,plain,(
  ( ! [X0] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f231,f1253,f367])).
fof(f1851,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f231,f1253,f367])).
fof(f1853,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f1253,f367])).
fof(f1856,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f231,f1253,f367])).
fof(f1857,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1253,f231,f1253,f367])).
fof(f1858,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f1253,f367])).
fof(f1859,plain,(
  ( ! [X2,X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f1253,f367])).
fof(f1860,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f370])).
fof(f1864,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f370])).
fof(f1865,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f370])).
fof(f1868,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f371])).
fof(f1873,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f372])).
fof(f1882,plain,(
  ( ! [X0] : (~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f420,f1253,f378])).
fof(f1887,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f379])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f380])).
fof(f1889,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f383])).
fof(f1893,plain,(
  ( ! [X2,X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f383])).
fof(f1894,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f384])).
fof(f1898,plain,(
  ( ! [X2,X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f384])).
fof(f1899,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f420,f1253,f385])).
fof(f1900,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f420,f1253,f389])).
fof(f1921,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f420,f231,f1253,f395])).
fof(f1927,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f395])).
fof(f1928,plain,(
  ( ! [X0,X1] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f395])).
fof(f1933,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f397])).
fof(f1934,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f420,f420,f1253,f399])).
fof(f1937,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f231,f1253,f401])).
fof(f1941,plain,(
  ( ! [X2,X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1253,f1253,f401])).
fof(f1942,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f402])).
fof(f1946,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f1253,f402])).
fof(f1947,plain,(
  ( ! [X0,X1] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f1253,f420,f1253,f402])).
fof(f1948,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f456,f1253,f405])).
fof(f1950,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f458,f1253,f405])).
fof(f1952,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f450,f1253,f406])).
fof(f1954,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f452,f1253,f406])).
fof(f1587,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f438,f428,f434,f420,f336])).
fof(f1532,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f442,f787,f430,f231,f335])).
fof(f1535,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f442,f426,f432,f231,f335])).
fof(f1253,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(factoring,[],[f281])).
fof(f1227,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X12)) | ~p(s(bool,X10)) | p(s(bool,X13)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),X13,X12)) )),
  inference(superposition,[],[f402,f276])).
fof(f1188,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1179,f390])).
fof(f1179,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f231,f231,f420,f395])).
fof(f1169,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f460,f231,f420,f394])).
fof(f1158,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f788,f276])).
fof(f1156,plain,(
  ( ! [X2,X3] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3)),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f788,f239])).
fof(f1151,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f788,f333])).
fof(f1100,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X10)) | p(s(bool,X12)) | p(s(bool,X13)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),X13,X12)) )),
  inference(superposition,[],[f385,f276])).
fof(f1081,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X8)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f787,f276])).
fof(f1079,plain,(
  ( ! [X2,X3] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X3)),'F_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f787,f239])).
fof(f1061,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1052,f373])).
fof(f1052,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f231,f420,f420,f378])).
fof(f1042,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f454,f420,f420,f377])).
fof(f1008,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X12)) | ~p(s(bool,X10)) | p(s(bool,X13)) | ~sP7(X13,X12,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f370,f276])).
fof(f983,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f974,f354])).
fof(f974,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f361])).
fof(f964,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f444,f231,f231,f357])).
fof(f960,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f951,f331])).
fof(f951,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f420,f420,f420,f346])).
fof(f945,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f936,f332])).
fof(f936,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f342])).
fof(f920,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f231,f420,f231,f338])).
fof(f907,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X12)) | ~p(s(bool,X10)) | p(s(bool,X13)) | ~sP5(X12,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),X13)) )),
  inference(superposition,[],[f334,f276])).
fof(f845,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f277,f239])).
fof(f819,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f406,f276])).
fof(f818,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50)) )),
  inference(superposition,[],[f405,f276])).
fof(f817,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X43)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f401,f276])).
fof(f816,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | p(s(bool,X41)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f400,f276])).
fof(f815,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,X35)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36)),X37,X38)) )),
  inference(superposition,[],[f384,f276])).
fof(f814,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,X31)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33)) )),
  inference(superposition,[],[f383,f276])).
fof(f813,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | p(s(bool,X30)) | sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f310,f276])).
fof(f812,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X27)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27)) )),
  inference(superposition,[],[f307,f276])).
fof(f811,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24)) )),
  inference(superposition,[],[f292,f276])).
fof(f810,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20)))) )),
  inference(superposition,[],[f291,f276])).
fof(f787,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f420,f276,f353])).
fof(f788,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) )),
  inference(unit_resulting_resolution,[],[f231,f276,f351])).
fof(f789,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f420,f276,f353])).
fof(f790,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) )),
  inference(unit_resulting_resolution,[],[f231,f276,f351])).
fof(f764,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f231,f462,f352])).
fof(f767,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f231,f462,f352])).
fof(f741,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f420,f462,f350])).
fof(f743,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f420,f462,f350])).
fof(f651,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f239,f239])).
fof(f589,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f401])).
fof(f575,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f231,f400])).
fof(f571,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f420,f389])).
fof(f565,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f420,f388])).
fof(f557,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f231,f384])).
fof(f547,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f231,f383])).
fof(f545,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f420,f372])).
fof(f544,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f371])).
fof(f529,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f420,f420,f366])).
fof(f521,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f420,f365])).
fof(f518,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f420,f356])).
fof(f515,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f231,f420,f355])).
fof(f509,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f420,f420,f421])).
fof(f498,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f420,f231,f310])).
fof(f490,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f420,f231,f307])).
fof(f487,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f231,f231,f306])).
fof(f486,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f420,f420,f305])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f420,f398])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f397])).
fof(f456,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f420,f396])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f420,f381])).
fof(f452,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f380])).
fof(f450,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f379])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f420,f364])).
fof(f446,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f420,f363])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f362])).
fof(f442,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f231,f349])).
fof(f440,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f231,f348])).
fof(f438,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f347])).
fof(f436,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f420,f345])).
fof(f434,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f420,f344])).
fof(f432,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f343])).
fof(f430,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f420,f341])).
fof(f428,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f420,f340])).
fof(f426,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f339])).
fof(f424,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f231,f292])).
fof(f422,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f231,f291])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f133,f150,f149])).
fof(f149,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) => p(s(bool,X15)))) <=> ((p(s(bool,X17)) | p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X15))) & (~p(s(bool,X16)) | p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.extreal.le_max','thm.sat.sat_27_prove_ltb')).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f225])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f222])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f105,f147,f146])).
fof(f146,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f147,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) | p(s(bool,X15)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.extreal.le_max','thm.sat.sat_25_prove_ltb')).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f214])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f214])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f103,f144,f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) & p(s(bool,X15)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X16)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.extreal.le_max','thm.sat.sat_23_prove_ltb')).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f211])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f143])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f207])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f101,f141,f140,f139,f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f139,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f140,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f141,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> s(bool,X15) = s(bool,X16)) <=> ((p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,X15))) & (p(s(bool,X17)) | ~p(s(bool,X15)) | ~p(s(bool,X16))) & (p(s(bool,X16)) | ~p(s(bool,X15)) | ~p(s(bool,X17))) & (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)))))),
  file('thm.extreal.le_max','thm.sat.sat_2|49|_prove_ltb')).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f204])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f138])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f203])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f203])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f139])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f198])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f198])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f323,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.extreal.le_max',aHL_EXT)).
fof(f305,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f183])).
fof(f183,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f306,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f184])).
fof(f307,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f310,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f421,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f290,f420])).
fof(f290,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f291,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f292,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f285,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f123])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : ! [X9,X10] : s(X12,X9) = s(X12,'const.bool.COND_3'(s(bool,'T_0'),s(X12,X9),s(X12,X10)))),
  file('thm.extreal.le_max','thm.bool.COND_CLAUSES|split|0')).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.extreal.le_max','thm.bool.EQ_CLAUSES|split|0')).
fof(f250,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f11608,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl1),
  inference(superposition,[],[f415,f1695])).
fof(f462,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f420,f249])).
fof(f21595,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl2 | $spl79)),
  inference(subsumption_resolution,[],[f21570,f231])).
fof(f21570,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl2 | $spl79)),
  inference(superposition,[],[f17259,f18404])).
fof(f18404,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))) | $spl79),
  inference(unit_resulting_resolution,[],[f231,f18259,f285])).
fof(f18259,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))) | $spl79),
  inference(unit_resulting_resolution,[],[f8220,f1273])).
fof(f8220,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl79),
  inference(cnf_transformation,[],[f8220_D])).
fof(f8220_D,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f17259,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f17163,f231])).
fof(f17163,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl2),
  inference(superposition,[],[f324,f16623])).
fof(f16623,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f416,f285])).
fof(f21592,plain,(
  $false | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(subsumption_resolution,[],[f21591,f420])).
fof(f21591,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(forward_demodulation,[],[f21590,f17610])).
fof(f21590,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl2 | $spl79)),
  inference(forward_demodulation,[],[f21589,f18258])).
fof(f18258,plain,(
  s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl79),
  inference(unit_resulting_resolution,[],[f8220,f1695])).
fof(f21589,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl2 | $spl79)),
  inference(subsumption_resolution,[],[f21588,f231])).
fof(f21588,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl2 | $spl79)),
  inference(forward_demodulation,[],[f21566,f18404])).
fof(f21566,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl2 | $spl79)),
  inference(superposition,[],[f17259,f18258])).
fof(f21583,plain,(
  $false | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(subsumption_resolution,[],[f21582,f420])).
fof(f21582,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl5 | $spl79)),
  inference(forward_demodulation,[],[f21581,f17610])).
fof(f21581,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl2 | $spl79)),
  inference(forward_demodulation,[],[f21580,f20644])).
fof(f20644,plain,(
  s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))) | $spl79),
  inference(forward_demodulation,[],[f20528,f276])).
fof(f20528,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl79),
  inference(superposition,[],[f280,f18404])).
fof(f21580,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))))) | ($spl2 | $spl79)),
  inference(subsumption_resolution,[],[f21579,f231])).
fof(f21579,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))))) | ($spl2 | $spl79)),
  inference(forward_demodulation,[],[f21564,f18404])).
fof(f21564,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13)))))) | ($spl2 | $spl79)),
  inference(superposition,[],[f17259,f20644])).
fof(f16605,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666,f11664,f11581,f415,f416])).
fof(f11581,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f462,f415,f710])).
fof(f11664,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f11359,f277])).
fof(f11359,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f842,f415,f334])).
fof(f11666,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f11358,f276])).
fof(f11358,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),s(bool,X0))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f788,f415,f334])).
fof(f11349,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f282])).
fof(f11347,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f282])).
fof(f11346,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f254])).
fof(f12121,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f11351,f307])).
fof(f12132,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f11353,f310])).
fof(f12144,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f11573,f292])).
fof(f12155,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f11577,f291])).
fof(f11342,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f1273])).
fof(f11345,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f249])).
fof(f11351,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f305])).
fof(f11353,plain,(
  ~sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f305])).
fof(f11354,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f305])).
fof(f11357,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f334])).
fof(f11364,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f438,f428,f434,f415,f336])).
fof(f11376,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f338])).
fof(f11381,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f340])).
fof(f11382,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f341])).
fof(f11383,plain,(
  ~sP3('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f342])).
fof(f11388,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f344])).
fof(f11389,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f345])).
fof(f11390,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f346])).
fof(f11398,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f346])).
fof(f11399,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f415,f420,f415,f346])).
fof(f11406,plain,(
  ~sP2('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f346])).
fof(f11407,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f420,f415,f346])).
fof(f11408,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f415,f346])).
fof(f11409,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f415,f346])).
fof(f11414,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f462,f415,f350])).
fof(f11416,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f462,f415,f350])).
fof(f11418,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f276,f415,f353])).
fof(f11419,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f277,f415,f353])).
fof(f11420,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f276,f415,f353])).
fof(f11421,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f277,f415,f353])).
fof(f11422,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f355])).
fof(f11427,plain,(
  ( ! [X0] : (~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f356])).
fof(f11432,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f361])).
fof(f11437,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f363])).
fof(f11438,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f364])).
fof(f11439,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f365])).
fof(f11445,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f365])).
fof(f11446,plain,(
  ( ! [X0] : (sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f366])).
fof(f11452,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f366])).
fof(f11453,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f370])).
fof(f11454,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f371])).
fof(f11455,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f372])).
fof(f11488,plain,(
  ~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f415,f378])).
fof(f11494,plain,(
  ~sP8('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f415,f378])).
fof(f11495,plain,(
  ~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f415,f378])).
fof(f11500,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f381])).
fof(f11501,plain,(
  sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f382])).
fof(f11507,plain,(
  sP9('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f382])).
fof(f11509,plain,(
  sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f415,f382])).
fof(f11513,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f415,f382])).
fof(f11514,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f415,f420,f415,f382])).
fof(f11515,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f415,f382])).
fof(f11516,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f415,f382])).
fof(f11517,plain,(
  ~sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f415,f385])).
fof(f11521,plain,(
  ~sP9('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f415,f385])).
fof(f11522,plain,(
  ~sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f415,f385])).
fof(f11525,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f388])).
fof(f11527,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f388])).
fof(f11528,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f388])).
fof(f11532,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f389])).
fof(f11549,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f395])).
fof(f11554,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f396])).
fof(f11555,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f398])).
fof(f11556,plain,(
  sP11('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f415,f399])).
fof(f11560,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f415,f399])).
fof(f11561,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f415,f231,f415,f399])).
fof(f11566,plain,(
  ~sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f402])).
fof(f11569,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f446,f415,f407])).
fof(f11571,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f448,f415,f407])).
fof(f11573,plain,(
  ~sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f421])).
fof(f11577,plain,(
  ~sP0('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f415,f421])).
fof(f11578,plain,(
  ~sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f415,f421])).
fof(f11583,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f812])).
fof(f11584,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f813])).
fof(f11585,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f816])).
fof(f11586,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f870])).
fof(f11587,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f871])).
fof(f11588,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f874])).
fof(f11589,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f907])).
fof(f11591,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f908])).
fof(f11593,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f1008])).
fof(f11594,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f1009])).
fof(f11595,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f415,f1100])).
fof(f11597,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f415,f1100])).
fof(f11598,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f231,f415,f1100])).
fof(f11599,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f415,f1101])).
fof(f11601,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f415,f1101])).
fof(f11602,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f415,f415,f1101])).
fof(f11603,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f1227])).
fof(f11604,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f415,f1228])).
fof(f11669,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f11609,f231])).
fof(f11609,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f415,f651])).
fof(f15566,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12405,f651])).
fof(f15554,plain,(
  ( ! [X10,X9] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12405,f651])).
fof(f15543,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12405,f239])).
fof(f15540,plain,(
  ( ! [X0] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0')) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12405,f876])).
fof(f15510,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12404,f651])).
fof(f15507,plain,(
  ( ! [X0] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0')) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12404,f818])).
fof(f15494,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12402,f651])).
fof(f15482,plain,(
  ( ! [X10,X9] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12402,f651])).
fof(f15471,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12402,f239])).
fof(f15468,plain,(
  ( ! [X0] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12402,f877])).
fof(f15455,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12401,f651])).
fof(f15443,plain,(
  ( ! [X10,X9] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12401,f651])).
fof(f15432,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12401,f239])).
fof(f15429,plain,(
  ( ! [X0] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12401,f877])).
fof(f15399,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12399,f651])).
fof(f15396,plain,(
  ( ! [X0] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12399,f819])).
fof(f15366,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12398,f651])).
fof(f15363,plain,(
  ( ! [X0] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(unit_resulting_resolution,[],[f231,f12398,f819])).
fof(f15351,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12397,f651])).
fof(f15339,plain,(
  ( ! [X10,X9] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12397,f651])).
fof(f15328,plain,(
  ( ! [X0,X1] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12397,f239])).
fof(f15299,plain,(
  ( ! [X2,X0,X1] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12396,f651])).
fof(f15282,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12395,f651])).
fof(f15270,plain,(
  ( ! [X10,X9] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12395,f651])).
fof(f15259,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12395,f239])).
fof(f15218,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12394,f651])).
fof(f15197,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12393,f651])).
fof(f15185,plain,(
  ( ! [X14,X15,X16] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X14)),X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12393,f651])).
fof(f15174,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12393,f239])).
fof(f15140,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.COND_3'(s(bool,X0),s(bool,X0),s(bool,X1)),X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12390,f651])).
fof(f14295,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl78),
  inference(unit_resulting_resolution,[],[f231,f420,f14290,f394])).
fof(f14290,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl78),
  inference(unit_resulting_resolution,[],[f231,f420,f12427])).
fof(f14248,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f420,f14209,f377])).
fof(f14209,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f420,f12424])).
fof(f14166,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f14162,f407])).
fof(f14162,plain,(
  ~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(unit_resulting_resolution,[],[f231,f420,f12421])).
fof(f14107,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f231,f420,f12420])).
fof(f13751,plain,(
  ( ! [X2,X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12392,f651])).
fof(f13739,plain,(
  ( ! [X10,X9] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12392,f651])).
fof(f13728,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12392,f239])).
fof(f13707,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12391,f651])).
fof(f13695,plain,(
  ( ! [X10,X9] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X9)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12391,f651])).
fof(f13684,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12391,f239])).
fof(f13246,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl78),
  inference(unit_resulting_resolution,[],[f231,f12426])).
fof(f13235,plain,(
  ( ! [X39,X41,X38,X40] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f845])).
fof(f13232,plain,(
  ( ! [X30,X31,X29] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,X30)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f859])).
fof(f13231,plain,(
  ( ! [X28,X26,X27] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X26),s(bool,X27)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f801])).
fof(f13229,plain,(
  ( ! [X24] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X24),s(bool,'T_0')))) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f8221])).
fof(f13226,plain,(
  ( ! [X21,X20] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.bool.COND_3'(s(bool,X20),s(bool,X21),s(bool,X20))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f239])).
fof(f13224,plain,(
  ( ! [X14,X15,X16] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X15),s(bool,X14))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f651])).
fof(f13223,plain,(
  ( ! [X12,X13] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X12,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f239])).
fof(f13221,plain,(
  ( ! [X6,X8,X7] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X6,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f651])).
fof(f13219,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0'))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f651])).
fof(f13196,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12339,f239])).
fof(f13193,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | $spl78),
  inference(unit_resulting_resolution,[],[f420,f12425])).
fof(f13182,plain,(
  ( ! [X39,X41,X38,X40] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,X40)),'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f845])).
fof(f13181,plain,(
  ( ! [X37,X35,X36] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X37),s(bool,X36)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36)))) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f277])).
fof(f13179,plain,(
  ( ! [X30,X31,X29] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,X30))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f859])).
fof(f13178,plain,(
  ( ! [X28,X26,X27] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,'T_0')),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X26),s(bool,X27))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f801])).
fof(f13176,plain,(
  ( ! [X24] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X24),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f8221])).
fof(f13173,plain,(
  ( ! [X21,X20] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,X20),s(bool,X21),s(bool,X20)),'F_0') | s(bool,'T_0') = s(bool,X20)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f239])).
fof(f13171,plain,(
  ( ! [X14,X15,X16] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X15),s(bool,X14)),'T_0') | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f651])).
fof(f13170,plain,(
  ( ! [X12,X13] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X13),s(bool,'F_0')),X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f239])).
fof(f13168,plain,(
  ( ! [X6,X8,X7] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X8),s(bool,X7)),X6) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f651])).
fof(f13166,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X0) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f651])).
fof(f13143,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl78)),
  inference(superposition,[],[f12337,f239])).
fof(f13140,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl78),
  inference(unit_resulting_resolution,[],[f231,f12423])).
fof(f13085,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X0,'T_0')) ) | $spl78),
  inference(unit_resulting_resolution,[],[f231,f12422])).
fof(f12886,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(unit_resulting_resolution,[],[f420,f12419])).
fof(f12880,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f12418])).
fof(f12490,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f440,f430,f12203,f337])).
fof(f12487,plain,(
  sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(unit_resulting_resolution,[],[f12202,f333])).
fof(f12477,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'T_0','F_0') | $spl78),
  inference(unit_resulting_resolution,[],[f231,f442,f430,f12201,f335])).
fof(f12465,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,X115)) | p(s(bool,X116)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X115)),X116,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12282,f231])).
fof(f12282,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,X116)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X115)),X116,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f1228,f8221])).
fof(f12464,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,X112)) | p(s(bool,X113)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X112),s(bool,X114)),X113,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12281,f231])).
fof(f12281,plain,(
  ( ! [X114,X112,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X112)) | p(s(bool,X113)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X112),s(bool,X114)),X113,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f1227,f8221])).
fof(f12459,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP7(X104,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X103)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12278,f231])).
fof(f12278,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X103)) | p(s(bool,X104)) | ~sP7(X104,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X105),s(bool,X103)))) ) | $spl78),
  inference(superposition,[],[f1009,f8221])).
fof(f12458,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X100),s(bool,X102)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12277,f231])).
fof(f12277,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'T_0')) | ~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X100),s(bool,X102)))) ) | $spl78),
  inference(superposition,[],[f1008,f8221])).
fof(f12457,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X99),s(bool,X97)),X98)) ) | $spl78),
  inference(subsumption_resolution,[],[f12276,f231])).
fof(f12276,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X99),s(bool,X97)),X98)) ) | $spl78),
  inference(superposition,[],[f908,f8221])).
fof(f12456,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X94),s(bool,X96)),X95)) ) | $spl78),
  inference(subsumption_resolution,[],[f12275,f231])).
fof(f12275,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X94)) | p(s(bool,X95)) | ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X94),s(bool,X96)),X95)) ) | $spl78),
  inference(superposition,[],[f907,f8221])).
fof(f12455,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X91)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X93)) ) | $spl78),
  inference(subsumption_resolution,[],[f12274,f231])).
fof(f12274,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X91)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X92),s(bool,X91)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X93)) ) | $spl78),
  inference(superposition,[],[f875,f8221])).
fof(f12452,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,X85)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X86),s(bool,X85)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X87)) ) | $spl78),
  inference(subsumption_resolution,[],[f12272,f231])).
fof(f12272,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X85)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X86),s(bool,X85)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X87)) ) | $spl78),
  inference(superposition,[],[f873,f8221])).
fof(f12451,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,X82)),X84,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12271,f231])).
fof(f12271,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X82)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,X82)),X84,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f872,f8221])).
fof(f12262,plain,(
  ( ! [X67] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,'T_0')))) ) | $spl78),
  inference(superposition,[],[f844,f8221])).
fof(f12260,plain,(
  ( ! [X65] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,'T_0')))) ) | $spl78),
  inference(superposition,[],[f843,f8221])).
fof(f12259,plain,(
  ( ! [X64] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f842,f8221])).
fof(f12257,plain,(
  ( ! [X62] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X62),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f841,f8221])).
fof(f12442,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X58),s(bool,X59)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X60)) ) | $spl78),
  inference(subsumption_resolution,[],[f12255,f231])).
fof(f12255,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X58),s(bool,X59)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X60)) ) | $spl78),
  inference(superposition,[],[f817,f8221])).
fof(f12439,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X54)) ) | $spl78),
  inference(subsumption_resolution,[],[f12253,f231])).
fof(f12253,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X52),s(bool,X53)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X54)) ) | $spl78),
  inference(superposition,[],[f815,f8221])).
fof(f12438,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12252,f231])).
fof(f12252,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)),X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f814,f8221])).
fof(f12244,plain,(
  ( ! [X37] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X37)))) ) | $spl78),
  inference(superposition,[],[f790,f8221])).
fof(f12242,plain,(
  ( ! [X35] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X35)))) ) | $spl78),
  inference(superposition,[],[f789,f8221])).
fof(f12240,plain,(
  ( ! [X33] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X33)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f788,f8221])).
fof(f12238,plain,(
  ( ! [X31] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f787,f8221])).
fof(f12429,plain,(
  ( ! [X26,X27] : (~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X26,X27) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X26,X27)) ) | $spl78),
  inference(subsumption_resolution,[],[f12234,f231])).
fof(f12234,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X26,X27) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X26,X27)) ) | $spl78),
  inference(superposition,[],[f406,f8221])).
fof(f12428,plain,(
  ( ! [X24,X25] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X24,X25) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X24,X25)) ) | $spl78),
  inference(subsumption_resolution,[],[f12233,f231])).
fof(f12233,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X24,X25) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X24,X25)) ) | $spl78),
  inference(superposition,[],[f405,f8221])).
fof(f12427,plain,(
  ( ! [X23,X22] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X23,X22) | p(s(bool,X23)) | ~p(s(bool,X22))) ) | $spl78),
  inference(subsumption_resolution,[],[f12232,f231])).
fof(f12232,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | p(s(bool,X23)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X23,X22)) ) | $spl78),
  inference(superposition,[],[f402,f8221])).
fof(f12426,plain,(
  ( ! [X21,X20] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X20,X21) | ~p(s(bool,X20))) ) | $spl78),
  inference(subsumption_resolution,[],[f12231,f231])).
fof(f12231,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X20,X21)) ) | $spl78),
  inference(superposition,[],[f401,f8221])).
fof(f12425,plain,(
  ( ! [X19,X18] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X19,X18) | p(s(bool,X18))) ) | $spl78),
  inference(subsumption_resolution,[],[f12230,f231])).
fof(f12230,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X19,X18)) ) | $spl78),
  inference(superposition,[],[f400,f8221])).
fof(f12424,plain,(
  ( ! [X17,X16] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X17,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl78),
  inference(subsumption_resolution,[],[f12229,f231])).
fof(f12229,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X17,X16)) ) | $spl78),
  inference(superposition,[],[f385,f8221])).
fof(f12423,plain,(
  ( ! [X14,X15] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X14,X15) | ~p(s(bool,X14))) ) | $spl78),
  inference(subsumption_resolution,[],[f12228,f231])).
fof(f12228,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X14,X15)) ) | $spl78),
  inference(superposition,[],[f384,f8221])).
fof(f12422,plain,(
  ( ! [X12,X13] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X13,X12) | ~p(s(bool,X12))) ) | $spl78),
  inference(subsumption_resolution,[],[f12227,f231])).
fof(f12227,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X13,X12)) ) | $spl78),
  inference(superposition,[],[f383,f8221])).
fof(f12421,plain,(
  ( ! [X10,X11] : (~sP7(X11,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl78),
  inference(subsumption_resolution,[],[f12226,f231])).
fof(f12226,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP7(X11,X10,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f370,f8221])).
fof(f12420,plain,(
  ( ! [X8,X9] : (~sP5(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl78),
  inference(subsumption_resolution,[],[f12225,f231])).
fof(f12225,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X9)) ) | $spl78),
  inference(superposition,[],[f334,f8221])).
fof(f12419,plain,(
  ( ! [X7] : (sP1(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | p(s(bool,X7))) ) | $spl78),
  inference(subsumption_resolution,[],[f12224,f231])).
fof(f12224,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f310,f8221])).
fof(f12418,plain,(
  ( ! [X6] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X6) | p(s(bool,X6))) ) | $spl78),
  inference(subsumption_resolution,[],[f12223,f231])).
fof(f12223,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X6)) ) | $spl78),
  inference(superposition,[],[f307,f8221])).
fof(f12417,plain,(
  ( ! [X5] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X5)) ) | $spl78),
  inference(subsumption_resolution,[],[f12222,f231])).
fof(f12222,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),X5)) ) | $spl78),
  inference(superposition,[],[f292,f8221])).
fof(f12416,plain,(
  ( ! [X4] : (sP0(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(subsumption_resolution,[],[f12221,f231])).
fof(f12221,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl78),
  inference(superposition,[],[f291,f8221])).
fof(f12405,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0')) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11604])).
fof(f12404,plain,(
  ( ! [X0] : (~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0')) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11603])).
fof(f12403,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11602])).
fof(f12402,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11601])).
fof(f12401,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11599])).
fof(f12400,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11598])).
fof(f12399,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11597])).
fof(f12398,plain,(
  ( ! [X0] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11595])).
fof(f12397,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11594])).
fof(f12396,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11593])).
fof(f12395,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11591])).
fof(f12394,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11589])).
fof(f12393,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11588])).
fof(f12392,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11587])).
fof(f12391,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11586])).
fof(f12390,plain,(
  ( ! [X0,X1] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11585])).
fof(f12339,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11421])).
fof(f12337,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | ($spl1 | $spl78)),
  inference(backward_demodulation,[],[f12306,f11419])).
fof(f12306,plain,(
  s('type.extreal.extreal',sK14) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(forward_demodulation,[],[f12210,f276])).
fof(f12210,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))) | $spl78),
  inference(superposition,[],[f280,f8221])).
fof(f12305,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0))))) ) | $spl78),
  inference(subsumption_resolution,[],[f12209,f231])).
fof(f12209,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0))))) ) | $spl78),
  inference(superposition,[],[f324,f8221])).
fof(f12201,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'T_0') | $spl78),
  inference(unit_resulting_resolution,[],[f420,f8221,f353])).
fof(f12202,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)),'T_0') | $spl78),
  inference(unit_resulting_resolution,[],[f231,f8221,f351])).
fof(f12203,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(unit_resulting_resolution,[],[f420,f8221,f353])).
fof(f12204,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(unit_resulting_resolution,[],[f231,f8221,f351])).
fof(f11344,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f281])).
fof(f482,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f254])).
fof(f1078,plain,(
  ( ! [X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(superposition,[],[f787,f482])).
fof(f1155,plain,(
  ( ! [X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(superposition,[],[f788,f482])).
fof(f1333,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f277,f1272,f351])).
fof(f1337,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f277,f1272,f351])).
fof(f1664,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f1651,f482])).
fof(f1651,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),s(bool,X0)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f787,f429,f435,f420,f337])).
fof(f2021,plain,(
  ( ! [X1] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X1)))) ) | $spl1),
  inference(superposition,[],[f789,f482])).
fof(f2033,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2014,f482])).
fof(f2014,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f437,f430,f789,f337])).
fof(f2050,plain,(
  ( ! [X1] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X1)))) ) | $spl1),
  inference(superposition,[],[f790,f482])).
fof(f2098,plain,(
  ( ! [X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(superposition,[],[f841,f482])).
fof(f2103,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2085,f482])).
fof(f2085,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f435,f429,f841,f337])).
fof(f2137,plain,(
  ( ! [X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(superposition,[],[f842,f482])).
fof(f2175,plain,(
  ( ! [X1] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'F_0')))) ) | $spl1),
  inference(superposition,[],[f843,f482])).
fof(f2181,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2160,f482])).
fof(f2160,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f437,f430,f843,f337])).
fof(f2211,plain,(
  ( ! [X1] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'F_0')))) ) | $spl1),
  inference(superposition,[],[f844,f482])).
fof(f2373,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2361,f1287])).
fof(f2361,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f843,f430,f1328,f337])).
fof(f2385,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2380,f1287])).
fof(f2380,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f841,f430,f1329,f337])).
fof(f2747,plain,(
  ( ! [X33,X31,X32] : (~sP5('F_0','const.bool.COND_3'(s(bool,X31),s(bool,'F_0'),s(bool,X32)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl1),
  inference(superposition,[],[f1078,f651])).
fof(f3389,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,X163)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X163,X162) | s(bool,X162) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | $spl2),
  inference(subsumption_resolution,[],[f2864,f231])).
fof(f2864,plain,(
  ( ! [X163,X164,X162] : (~p(s(bool,'T_0')) | p(s(bool,X163)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X163,X162) | s(bool,X162) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | $spl2),
  inference(superposition,[],[f1232,f651])).
fof(f2753,plain,(
  ( ! [X54,X52,X55,X53] : (sP1(X53,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X52))) | p(s(bool,X53)) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | $spl2),
  inference(superposition,[],[f2468,f651])).
fof(f2752,plain,(
  ( ! [X50,X48,X51,X49] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X49),s(bool,X48)),X50) | p(s(bool,X50)) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl2),
  inference(superposition,[],[f2441,f651])).
fof(f2751,plain,(
  ( ! [X47,X45,X46,X44] : (sP1(X45,'const.bool.COND_3'(s(bool,X44),s(bool,X44),s(bool,X46))) | p(s(bool,X45)) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) ) | $spl2),
  inference(superposition,[],[f2419,f651])).
fof(f2750,plain,(
  ( ! [X43,X41,X42,X40] : (sP1('const.bool.COND_3'(s(bool,X40),s(bool,X40),s(bool,X41)),X42) | p(s(bool,X42)) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | $spl2),
  inference(superposition,[],[f2395,f651])).
fof(f2749,plain,(
  ( ! [X39,X37,X38] : (sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,X37),s(bool,X37),s(bool,X38))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl2),
  inference(superposition,[],[f1200,f651])).
fof(f2748,plain,(
  ( ! [X35,X36,X34] : (sP5('T_0','const.bool.COND_3'(s(bool,X34),s(bool,X34),s(bool,X35)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl2),
  inference(superposition,[],[f1154,f651])).
fof(f2746,plain,(
  ( ! [X30,X28,X29] : (~sP5('F_0','const.bool.COND_3'(s(bool,X28),s(bool,X28),s(bool,X29)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl2),
  inference(superposition,[],[f1077,f651])).
fof(f2472,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,'T_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f2468,f239])).
fof(f2468,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0'))) | p(s(bool,X0))) ) | $spl2),
  inference(forward_demodulation,[],[f2467,f466])).
fof(f2467,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP1(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f2461,f231])).
fof(f2461,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))))) ) | $spl2),
  inference(superposition,[],[f871,f466])).
fof(f2445,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,'T_0')),X2) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f2441,f239])).
fof(f2441,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X0) | p(s(bool,X0))) ) | $spl2),
  inference(forward_demodulation,[],[f2440,f466])).
fof(f2440,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))),X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f2434,f231])).
fof(f2434,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))),X0)) ) | $spl2),
  inference(superposition,[],[f870,f466])).
fof(f2419,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1))) | p(s(bool,X0))) ) | $spl2),
  inference(forward_demodulation,[],[f2418,f466])).
fof(f2418,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP1(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))),s(bool,X1)))) ) | $spl2),
  inference(subsumption_resolution,[],[f2412,f231])).
fof(f2412,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))),s(bool,X1)))) ) | $spl2),
  inference(superposition,[],[f813,f466])).
fof(f2395,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X0) | p(s(bool,X0))) ) | $spl2),
  inference(forward_demodulation,[],[f2394,f466])).
fof(f2394,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))),s(bool,X1)),X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f2388,f231])).
fof(f2388,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))),s(bool,X1)),X0)) ) | $spl2),
  inference(superposition,[],[f812,f466])).
fof(f2378,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f437,f430,f1329,f337])).
fof(f2386,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2379,f1287])).
fof(f2379,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f787,f430,f1329,f337])).
fof(f2381,plain,(
  sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f436,f430,f1329,f336])).
fof(f2382,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f436,f431,f1329,f336])).
fof(f2358,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f434,f430,f1328,f337])).
fof(f2359,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f435,f430,f1328,f337])).
fof(f2374,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2360,f1287])).
fof(f2360,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))),s(bool,X0)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f789,f430,f1328,f337])).
fof(f2346,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f2341,f420])).
fof(f2341,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,X1) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f710,f482])).
fof(f2326,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f436,f428,f1327,f336])).
fof(f2318,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f442,f426,f1326,f335])).
fof(f2321,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f442,f1313,f1326,f335])).
fof(f2304,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1320,f332])).
fof(f2299,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1314,f332])).
fof(f2276,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f442,f432,f1313,f335])).
fof(f2279,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f231,f442,f433,f1313,f335])).
fof(f2269,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1307,f333])).
fof(f2210,plain,(
  ( ! [X0] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl2),
  inference(superposition,[],[f844,f466])).
fof(f2174,plain,(
  ( ! [X0] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl2),
  inference(superposition,[],[f843,f466])).
fof(f2180,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl2),
  inference(forward_demodulation,[],[f2162,f466])).
fof(f2162,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f441,f427,f843,f337])).
fof(f2136,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f842,f466])).
fof(f2097,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f841,f466])).
fof(f2105,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | $spl2),
  inference(forward_demodulation,[],[f2081,f466])).
fof(f2081,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f443,f430,f841,f337])).
fof(f2049,plain,(
  ( ! [X0] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl2),
  inference(superposition,[],[f790,f466])).
fof(f2020,plain,(
  ( ! [X0] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl2),
  inference(superposition,[],[f789,f466])).
fof(f2032,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl2),
  inference(forward_demodulation,[],[f2016,f466])).
fof(f2016,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))),s(bool,X0)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f441,f427,f789,f337])).
fof(f1775,plain,(
  ( ! [X0] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f438,f427,f432,f1253,f335])).
fof(f1777,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f426,f433,f1253,f335])).
fof(f1779,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f439,f427,f433,f1253,f335])).
fof(f1809,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f466,f1253,f351])).
fof(f1810,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f482,f1253,f351])).
fof(f1813,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f466,f1253,f351])).
fof(f1814,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f482,f1253,f351])).
fof(f1818,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f463,f1253,f352])).
fof(f1820,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f463,f1253,f352])).
fof(f1825,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f445,f1253,f357])).
fof(f1830,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f445,f1253,f357])).
fof(f1831,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1253,f445,f1253,f357])).
fof(f1920,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f461,f1253,f394])).
fof(f1949,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f457,f1253,f405])).
fof(f1951,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f459,f1253,f405])).
fof(f1953,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f451,f1253,f406])).
fof(f1955,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f453,f1253,f406])).
fof(f1956,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1253,f550])).
fof(f1957,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1253,f560])).
fof(f1958,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1253,f592])).
fof(f1959,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f1253,f911])).
fof(f1960,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f1253,f1013])).
fof(f1961,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f1253,f1232])).
fof(f1725,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1716,f231])).
fof(f1716,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl2),
  inference(superposition,[],[f324,f466])).
fof(f1694,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(forward_demodulation,[],[f1677,f277])).
fof(f1677,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'F_0'),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(superposition,[],[f280,f482])).
fof(f1693,plain,(
  s('type.extreal.extreal',sK14) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(forward_demodulation,[],[f1676,f276])).
fof(f1676,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12))) | $spl2),
  inference(superposition,[],[f280,f466])).
fof(f1626,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f443,f1077,f430,f420,f337])).
fof(f1627,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f443,f778,f426,f420,f337])).
fof(f1641,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f441,f427,f434,f420,f337])).
fof(f1653,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f441,f427,f435,f420,f337])).
fof(f1661,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f443,f430,f437,f420,f337])).
fof(f1663,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f777,f428,f437,f420,f337])).
fof(f1577,plain,(
  sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f443,f430,f436,f420,f336])).
fof(f1581,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f443,f431,f436,f420,f336])).
fof(f1588,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f439,f428,f434,f420,f336])).
fof(f1592,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f439,f429,f435,f420,f336])).
fof(f1594,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f780,f430,f435,f420,f336])).
fof(f1531,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f442,f775,f430,f231,f335])).
fof(f1537,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f442,f426,f433,f231,f335])).
fof(f1538,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f442,f427,f433,f231,f335])).
fof(f1541,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f442,f431,f436,f231,f335])).
fof(f1542,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f442,f765,f437,f231,f335])).
fof(f1544,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f442,f431,f437,f231,f335])).
fof(f1287,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f285])).
fof(f1289,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f291])).
fof(f1290,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f292])).
fof(f1291,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f306])).
fof(f1293,plain,(
  ~sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f306])).
fof(f1294,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f1272,f306])).
fof(f1295,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f307])).
fof(f1296,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f310])).
fof(f1299,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f334])).
fof(f1300,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f334])).
fof(f1301,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f338])).
fof(f1307,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f338])).
fof(f1308,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f338])).
fof(f1313,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f339])).
fof(f1314,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f342])).
fof(f1320,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f342])).
fof(f1321,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f342])).
fof(f1326,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f343])).
fof(f1327,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f347])).
fof(f1328,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f348])).
fof(f1329,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f349])).
fof(f1330,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f466,f1272,f351])).
fof(f1331,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f482,f1272,f351])).
fof(f1332,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f276,f1272,f351])).
fof(f1334,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f466,f1272,f351])).
fof(f1335,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f482,f1272,f351])).
fof(f1336,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f276,f1272,f351])).
fof(f1338,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f462,f1272,f352])).
fof(f1339,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f463,f1272,f352])).
fof(f1340,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f462,f1272,f352])).
fof(f1341,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f463,f1272,f352])).
fof(f1342,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f355])).
fof(f1343,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f356])).
fof(f1346,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f231,f445,f1272,f357])).
fof(f1351,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f231,f445,f1272,f357])).
fof(f1352,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1272,f445,f1272,f357])).
fof(f1354,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f361])).
fof(f1360,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f361])).
fof(f1361,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f361])).
fof(f1366,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f362])).
fof(f1367,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f1272,f367])).
fof(f1372,plain,(
  sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f1272,f367])).
fof(f1374,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f1272,f367])).
fof(f1377,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f231,f1272,f367])).
fof(f1378,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f231,f1272,f367])).
fof(f1379,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f1272,f367])).
fof(f1380,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f1272,f1272,f367])).
fof(f1381,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f370])).
fof(f1385,plain,(
  ~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f370])).
fof(f1386,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f370])).
fof(f1389,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f371])).
fof(f1394,plain,(
  ( ! [X0] : (~sP9('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f372])).
fof(f1403,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f1272,f378])).
fof(f1408,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f379])).
fof(f1409,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f380])).
fof(f1410,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f383])).
fof(f1414,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f1272,f383])).
fof(f1415,plain,(
  ( ! [X0] : (sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f384])).
fof(f1419,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f1272,f384])).
fof(f1420,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f1272,f385])).
fof(f1421,plain,(
  ( ! [X0] : (~sP11('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f1272,f389])).
fof(f1441,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f461,f1272,f394])).
fof(f1442,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f231,f1272,f395])).
fof(f1448,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f395])).
fof(f1449,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f395])).
fof(f1454,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f397])).
fof(f1455,plain,(
  sP11('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f420,f1272,f399])).
fof(f1458,plain,(
  ( ! [X0] : (sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f1272,f401])).
fof(f1462,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f1272,f401])).
fof(f1463,plain,(
  ~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f402])).
fof(f1467,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f420,f1272,f402])).
fof(f1468,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | $spl1),
  inference(unit_resulting_resolution,[],[f1272,f420,f1272,f402])).
fof(f1469,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f456,f1272,f405])).
fof(f1470,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f457,f1272,f405])).
fof(f1471,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f458,f1272,f405])).
fof(f1472,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f459,f1272,f405])).
fof(f1473,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'T_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f450,f1272,f406])).
fof(f1474,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f451,f1272,f406])).
fof(f1475,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'T_0',X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f452,f1272,f406])).
fof(f1476,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f453,f1272,f406])).
fof(f1477,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1272,f550])).
fof(f1478,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1272,f560])).
fof(f1479,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1272,f592])).
fof(f1480,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f420,f1272,f911])).
fof(f1481,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f420,f1272,f1013])).
fof(f1482,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f420,f1272,f1232])).
fof(f1272,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) | $spl1),
  inference(subsumption_resolution,[],[f1259,f420])).
fof(f1259,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)))) | $spl1),
  inference(superposition,[],[f281,f482])).
fof(f1250,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f231,f420,f1233,f394])).
fof(f1243,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X11)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X11,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X9),s(bool,X10)))) ) | $spl2),
  inference(superposition,[],[f1232,f277])).
fof(f1242,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X8)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X8,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7)))) ) | $spl2),
  inference(superposition,[],[f1232,f276])).
fof(f1246,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(subsumption_resolution,[],[f1238,f231])).
fof(f1238,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f1232,f466])).
fof(f1233,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f231,f420,f1232])).
fof(f1232,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f1223,f231])).
fof(f1223,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0)) ) | $spl2),
  inference(superposition,[],[f402,f466])).
fof(f1200,plain,(
  ( ! [X0] : (sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1154,f333])).
fof(f1170,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f461,f231,f420,f394])).
fof(f1154,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f788,f466])).
fof(f1139,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f1078,f239])).
fof(f1116,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f420,f420,f1106,f377])).
fof(f1106,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f420,f420,f1105])).
fof(f1105,plain,(
  ( ! [X0,X1] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1096,f231])).
fof(f1096,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0)) ) | $spl2),
  inference(superposition,[],[f385,f466])).
fof(f1077,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f787,f466])).
fof(f1043,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f455,f420,f420,f377])).
fof(f1031,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f420,f1024,f407])).
fof(f1024,plain,(
  ~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f420,f1013])).
fof(f1017,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f231,f420,f911])).
fof(f1013,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1004,f231])).
fof(f1004,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f370,f466])).
fof(f965,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f231,f231,f357])).
fof(f911,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f903,f231])).
fof(f903,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(superposition,[],[f334,f466])).
fof(f775,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f420,f466,f353])).
fof(f777,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f482,f353])).
fof(f778,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f466,f353])).
fof(f780,plain,(
  ~sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f482,f353])).
fof(f765,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f463,f352])).
fof(f768,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f463,f352])).
fof(f751,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f231,f466,f351])).
fof(f752,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f231,f482,f351])).
fof(f753,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f466,f351])).
fof(f754,plain,(
  sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f231,f482,f351])).
fof(f742,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f420,f463,f350])).
fof(f744,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f420,f463,f350])).
fof(f738,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f451,f631])).
fof(f740,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f453,f631])).
fof(f716,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f457,f610])).
fof(f718,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f459,f610])).
fof(f642,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f592])).
fof(f639,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f447,f420,f407])).
fof(f641,plain,(
  ( ! [X0] : (sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f449,f420,f407])).
fof(f631,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f629,f231])).
fof(f629,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f406,f466])).
fof(f626,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f451,f231,f406])).
fof(f628,plain,(
  ( ! [X0] : (sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f453,f231,f406])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f608,f231])).
fof(f608,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f405,f466])).
fof(f605,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f457,f231,f405])).
fof(f607,plain,(
  ( ! [X0] : (sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f459,f231,f405])).
fof(f596,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f578])).
fof(f592,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f590,f231])).
fof(f590,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f401,f466])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f576,f231])).
fof(f576,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0)) ) | $spl2),
  inference(superposition,[],[f400,f466])).
fof(f572,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f560])).
fof(f564,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f231,f550])).
fof(f560,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f558,f231])).
fof(f558,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f384,f466])).
fof(f550,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f548,f231])).
fof(f548,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,X0)) ) | $spl2),
  inference(superposition,[],[f383,f466])).
fof(f514,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f420,f500])).
fof(f500,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f499,f231])).
fof(f499,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(superposition,[],[f310,f466])).
fof(f495,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f420,f492])).
fof(f492,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f491,f231])).
fof(f491,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(superposition,[],[f307,f466])).
fof(f466,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f250])).
fof(f463,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f249])).
fof(f461,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f398])).
fof(f459,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f397])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f396])).
fof(f455,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f381])).
fof(f453,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f380])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f379])).
fof(f449,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f364])).
fof(f447,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f363])).
fof(f445,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f362])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f349])).
fof(f441,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f348])).
fof(f439,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f347])).
fof(f437,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f345])).
fof(f435,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f344])).
fof(f433,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f343])).
fof(f431,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f341])).
fof(f429,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f415,f340])).
fof(f427,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f339])).
fof(f425,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f292])).
fof(f423,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f416,f291])).
fof(f16604,plain,(
  $false | ($spl1 | $spl2 | $spl32 | $spl78)),
  inference(global_subsumption,[],[f3366,f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666,f11664,f11581,f415])).
fof(f16603,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666,f11664,f11581,f415])).
fof(f16602,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666,f11664,f11581])).
fof(f16601,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666,f11664])).
fof(f16600,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349,f11666])).
fof(f16599,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347,f11349])).
fof(f16598,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346,f11347])).
fof(f16597,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121,f11346])).
fof(f16596,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132,f12121])).
fof(f16595,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144,f12132])).
fof(f16594,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155,f12144])).
fof(f16593,plain,(
  $false | ($spl1 | $spl2 | $spl78)),
  inference(global_subsumption,[],[f230,f229,f228,f231,f239,f250,f249,f276,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f416,f423,f422,f425,f424,f427,f426,f429,f428,f431,f430,f433,f432,f435,f434,f437,f436,f439,f438,f441,f440,f443,f442,f445,f444,f447,f446,f449,f448,f451,f450,f453,f452,f455,f454,f457,f456,f459,f458,f461,f460,f463,f462,f466,f486,f487,f490,f492,f495,f498,f500,f509,f514,f515,f518,f521,f529,f544,f545,f547,f550,f557,f560,f564,f565,f571,f572,f575,f578,f589,f592,f596,f607,f605,f610,f628,f626,f631,f641,f639,f642,f651,f710,f718,f716,f740,f738,f744,f743,f742,f741,f754,f753,f752,f751,f768,f767,f765,f764,f780,f778,f777,f775,f790,f789,f788,f787,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f845,f911,f907,f920,f936,f945,f951,f960,f965,f964,f974,f983,f1013,f1008,f1017,f1024,f1031,f1043,f1042,f1052,f1061,f1077,f1079,f1081,f1105,f1100,f1106,f1116,f1139,f1151,f1154,f1156,f1158,f1170,f1169,f1179,f1188,f1200,f1232,f1227,f1233,f1246,f1242,f1243,f1250,f1253,f1272,f1482,f1481,f1480,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1472,f1471,f1470,f1469,f1468,f1467,f1463,f1462,f1458,f1455,f1454,f1449,f1448,f1442,f1441,f1421,f1420,f1419,f1415,f1414,f1410,f1409,f1408,f1403,f1394,f1389,f1386,f1385,f1381,f1380,f1379,f1378,f1377,f1374,f1372,f1367,f1366,f1361,f1360,f1354,f1352,f1351,f1346,f1343,f1342,f1341,f1340,f1339,f1338,f1336,f1335,f1334,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1321,f1320,f1314,f1313,f1308,f1307,f1301,f1300,f1299,f1296,f1295,f1294,f1293,f1291,f1290,f1289,f1287,f1544,f1542,f1541,f1538,f1537,f1535,f1532,f1531,f1594,f1592,f1588,f1587,f1581,f1577,f1663,f1661,f1653,f1641,f1627,f1626,f1693,f1694,f1725,f1961,f1960,f1959,f1958,f1957,f1956,f1955,f1954,f1953,f1952,f1951,f1950,f1949,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1920,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1831,f1830,f1825,f1822,f1821,f1820,f1819,f1818,f1817,f1815,f1814,f1813,f1811,f1810,f1809,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1779,f1777,f1775,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2032,f2015,f2020,f2022,f2045,f2049,f2051,f2105,f2089,f2097,f2099,f2128,f2136,f2138,f2180,f2166,f2174,f2176,f2202,f2210,f2212,f2233,f2255,f2269,f2279,f2276,f2299,f2304,f2321,f2318,f2326,f2346,f2374,f2359,f2358,f2382,f2381,f2386,f2378,f2387,f2395,f2411,f2419,f2441,f2445,f2468,f2472,f2505,f2519,f2530,f2568,f2570,f2572,f2574,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2748,f2749,f2750,f2751,f2752,f2753,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f3365,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f3389,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3421,f3604,f3605,f3607,f3610,f3613,f3616,f3619,f3621,f3623,f3625,f3391,f3390,f2631,f1695,f1273,f324,f281,f280,f2638,f2905,f3388,f3387,f2854,f2852,f2851,f2849,f2794,f2792,f2791,f2789,f2747,f2586,f2584,f2583,f2581,f2531,f2506,f2460,f2433,f2385,f2373,f2306,f2285,f2213,f2211,f2198,f2181,f2177,f2175,f2163,f2161,f2141,f2139,f2137,f2125,f2103,f2102,f2100,f2098,f2086,f2084,f2054,f2052,f2050,f2033,f2025,f2023,f2021,f2017,f1816,f1812,f1747,f1746,f1664,f1639,f1337,f1333,f1228,f1159,f1157,f1155,f1101,f1082,f1080,f1078,f1009,f908,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f859,f844,f843,f842,f841,f801,f482,f420,f277,f254,f253,f6848,f6863,f6865,f6868,f6861,f6862,f6891,f6906,f6908,f6911,f6904,f6905,f6918,f6933,f6935,f6938,f6951,f6966,f6968,f6971,f6964,f6965,f7003,f7018,f7020,f7023,f7016,f7017,f7054,f7069,f7071,f7074,f7067,f7068,f7081,f7096,f7098,f7101,f7112,f7127,f7129,f7132,f7125,f7126,f7142,f7161,f7162,f7282,f7284,f7285,f7286,f7287,f7288,f7289,f7290,f7291,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7207,f7209,f7211,f7213,f7307,f7308,f7311,f7224,f7226,f7227,f7229,f7313,f7315,f7317,f7319,f7320,f7321,f7323,f7324,f7351,f7352,f7353,f7395,f7397,f7399,f7404,f7405,f7406,f7407,f7408,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7429,f7430,f7432,f7434,f7436,f7438,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7449,f7451,f7452,f7454,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7477,f7478,f7481,f7482,f7535,f7561,f7562,f7563,f7674,f7677,f7678,f7679,f7680,f7681,f7682,f7683,f7684,f7685,f7594,f7596,f7598,f7600,f7694,f7695,f7698,f7613,f7615,f7616,f7618,f7707,f7708,f7711,f7729,f7736,f7760,f7763,f7764,f7770,f7796,f7798,f7801,f7804,f7805,f7789,f7790,f7791,f7809,f7835,f7837,f7840,f7843,f7844,f7828,f7829,f7830,f7854,f7877,f7879,f7882,f7885,f7886,f7872,f7873,f7874,f7901,f7924,f7926,f7929,f7932,f7933,f7919,f7920,f7921,f7950,f7968,f7970,f7973,f7967,f7999,f8017,f8019,f8022,f8016,f8045,f8063,f8065,f8068,f8071,f8072,f8060,f8061,f8062,f8081,f8099,f8101,f8104,f8107,f8108,f8096,f8097,f8098,f8122,f8125,f8130,f8133,f8137,f8138,f8139,f8140,f8141,f8142,f8143,f8144,f8145,f8194,f8195,f8198,f8199,f8200,f8201,f8221,f415,f11344,f12204,f12203,f12202,f12201,f12305,f12306,f12337,f12339,f12390,f12391,f12392,f12393,f12394,f12395,f12396,f12397,f12398,f12399,f12400,f12401,f12402,f12403,f12404,f12405,f12416,f12417,f12418,f12419,f12420,f12421,f12422,f12423,f12424,f12425,f12426,f12427,f12428,f12429,f12238,f12240,f12242,f12244,f12438,f12439,f12442,f12257,f12259,f12260,f12262,f12451,f12452,f12455,f12456,f12457,f12458,f12459,f12464,f12465,f12477,f12487,f12490,f12880,f12886,f13085,f13140,f13143,f13166,f13168,f13170,f13171,f13173,f13176,f13178,f13179,f13181,f13182,f13193,f13196,f13219,f13221,f13223,f13224,f13226,f13229,f13231,f13232,f13235,f13246,f13277,f13279,f13282,f13276,f13329,f13331,f13334,f13328,f13382,f13384,f13387,f13381,f13440,f13442,f13445,f13439,f13475,f13476,f13477,f13478,f13479,f13499,f13500,f13514,f13684,f13695,f13707,f13728,f13739,f13751,f14107,f14162,f14166,f14209,f14248,f14290,f14295,f15140,f15174,f15185,f15197,f15218,f15259,f15270,f15282,f15299,f15328,f15339,f15351,f15363,f15366,f15396,f15399,f15429,f15432,f15443,f15455,f15468,f15471,f15482,f15494,f15507,f15510,f15540,f15543,f15554,f15566,f11669,f11604,f11603,f11602,f11601,f11599,f11598,f11597,f11595,f11594,f11593,f11591,f11589,f11588,f11587,f11586,f11585,f11584,f11583,f11578,f11577,f11573,f11571,f11569,f11566,f11561,f11560,f11556,f11555,f11554,f11549,f11532,f11528,f11527,f11525,f11522,f11521,f11517,f11516,f11515,f11514,f11513,f11509,f11507,f11501,f11500,f11495,f11494,f11488,f11455,f11454,f11453,f11452,f11446,f11445,f11439,f11438,f11437,f11432,f11427,f11422,f11421,f11420,f11419,f11418,f11416,f11414,f11409,f11408,f11407,f11406,f11399,f11398,f11390,f11389,f11388,f11383,f11382,f11381,f11376,f11364,f11357,f11354,f11353,f11351,f11345,f11342,f12155])).
fof(f16592,plain,(
  $false | ($spl4 | $spl80 | $spl84)),
  inference(subsumption_resolution,[],[f16528,f8224])).
fof(f16528,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f787,f8222,f337])).
fof(f11012,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f349])).
fof(f16591,plain,(
  $false | ($spl4 | $spl80 | $spl84)),
  inference(subsumption_resolution,[],[f16529,f8224])).
fof(f16529,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f841,f8222,f337])).
fof(f16588,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16538,f432])).
fof(f16538,plain,(
  ~sP3('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'T_0') | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f11700,f8222,f336])).
fof(f11700,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl4),
  inference(unit_resulting_resolution,[],[f10990,f333])).
fof(f10990,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f231,f420,f418,f338])).
fof(f16587,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16539,f11009])).
fof(f11009,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f343])).
fof(f16539,plain,(
  ~sP3('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f12183,f8222,f336])).
fof(f12183,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl4),
  inference(unit_resulting_resolution,[],[f10991,f333])).
fof(f10991,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f418,f420,f418,f338])).
fof(f16586,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16542,f438])).
fof(f16542,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','T_0') | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11700,f432,f8222,f336])).
fof(f16585,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16543,f420])).
fof(f16543,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f11700,f432,f8222,f336])).
fof(f16584,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16549,f11010])).
fof(f11010,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f347])).
fof(f16549,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f12183,f11009,f8222,f336])).
fof(f16583,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16550,f420])).
fof(f16550,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f12183,f11009,f8222,f336])).
fof(f16582,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16553,f11700])).
fof(f16553,plain,(
  sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl80),
  inference(unit_resulting_resolution,[],[f420,f438,f436,f8222,f336])).
fof(f16581,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16555,f12183])).
fof(f16555,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11010,f436,f8222,f336])).
fof(f16580,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16556,f438])).
fof(f16556,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','T_0') | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11700,f436,f8222,f336])).
fof(f16579,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16557,f420])).
fof(f16557,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f11700,f436,f8222,f336])).
fof(f16578,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16559,f11010])).
fof(f16559,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f12183,f436,f8222,f336])).
fof(f16577,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16560,f420])).
fof(f16560,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f12183,f436,f8222,f336])).
fof(f16576,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16575,f231])).
fof(f16575,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f16562,f10964])).
fof(f10964,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | $spl4),
  inference(unit_resulting_resolution,[],[f231,f418,f285])).
fof(f16562,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f11700,f432,f8222,f335])).
fof(f16574,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16573,f231])).
fof(f16573,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f16563,f10964])).
fof(f16563,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f12183,f11009,f8222,f335])).
fof(f16572,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16571,f231])).
fof(f16571,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f16564,f10964])).
fof(f16564,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f11700,f436,f8222,f335])).
fof(f16570,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16569,f231])).
fof(f16569,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f16565,f10964])).
fof(f16565,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f11012,f12183,f436,f8222,f335])).
fof(f16568,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f16567,f231])).
fof(f16567,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f16566,f10964])).
fof(f16566,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | $spl80),
  inference(unit_resulting_resolution,[],[f231,f420,f8222,f338])).
fof(f16561,plain,(
  $false | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f12183,f436,f8222,f336])).
fof(f16558,plain,(
  $false | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f11700,f436,f8222,f336])).
fof(f16551,plain,(
  $false | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f12183,f11009,f8222,f336])).
fof(f16544,plain,(
  $false | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f420,f11012,f11700,f432,f8222,f336])).
fof(f16525,plain,(
  $false | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f10990,f8222])).
fof(f16521,plain,(
  $false | ($spl3 | $spl4 | $spl78)),
  inference(subsumption_resolution,[],[f16520,f420])).
fof(f16520,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl78)),
  inference(forward_demodulation,[],[f16519,f5338])).
fof(f5338,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f4670,f350])).
fof(f4670,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f3989,f331])).
fof(f3989,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f420,f420,f417,f346])).
fof(f16519,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | ($spl4 | $spl78)),
  inference(subsumption_resolution,[],[f16498,f231])).
fof(f16498,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | ($spl4 | $spl78)),
  inference(superposition,[],[f12047,f8221])).
fof(f12047,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f11951,f231])).
fof(f11951,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(superposition,[],[f324,f10964])).
fof(f16516,plain,(
  $false | ($spl3 | $spl4 | $spl78)),
  inference(subsumption_resolution,[],[f16515,f420])).
fof(f16515,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl78)),
  inference(forward_demodulation,[],[f16514,f5338])).
fof(f16514,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | ($spl4 | $spl78)),
  inference(forward_demodulation,[],[f16513,f12306])).
fof(f16513,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl4 | $spl78)),
  inference(subsumption_resolution,[],[f16512,f231])).
fof(f16512,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl4 | $spl78)),
  inference(forward_demodulation,[],[f16494,f8221])).
fof(f16494,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | ($spl4 | $spl78)),
  inference(superposition,[],[f12047,f12306])).
fof(f13510,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) | $spl118),
  inference(cnf_transformation,[],[f13510_D])).
fof(f13510_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f13512,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl121),
  inference(cnf_transformation,[],[f13512_D])).
fof(f13512_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f13500,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK15(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1747,f651])).
fof(f13499,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1747,f651])).
fof(f13485,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl112),
  inference(cnf_transformation,[],[f13485_D])).
fof(f13485_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f13481,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) | $spl108),
  inference(cnf_transformation,[],[f13481_D])).
fof(f13481_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f13488,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl117),
  inference(cnf_transformation,[],[f13488_D])).
fof(f13488_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f13479,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1746,f239])).
fof(f13486,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl114),
  inference(cnf_transformation,[],[f13486_D])).
fof(f13486_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f13478,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK15(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1746,f651])).
fof(f13477,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1746,f651])).
fof(f13480,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl106),
  inference(cnf_transformation,[],[f13480_D])).
fof(f13480_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13483,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl111),
  inference(cnf_transformation,[],[f13483_D])).
fof(f13483_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f13476,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK15(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1746,f651])).
fof(f13475,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1746,f651])).
fof(f13443,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) | $spl104),
  inference(cnf_transformation,[],[f13443_D])).
fof(f13443_D,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3330,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f3330_D])).
fof(f3330_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13445,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f13444])).
fof(f13444,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13429,f710])).
fof(f13429,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f877,f651])).
fof(f13442,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f13441])).
fof(f13441,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13428,f710])).
fof(f13428,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f877,f651])).
fof(f13385,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) | $spl102),
  inference(cnf_transformation,[],[f13385_D])).
fof(f13385_D,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f13387,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f13386])).
fof(f13386,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13371,f710])).
fof(f13371,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f876,f651])).
fof(f13384,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f13383])).
fof(f13383,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13370,f710])).
fof(f13370,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f876,f651])).
fof(f13332,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) | $spl100),
  inference(cnf_transformation,[],[f13332_D])).
fof(f13332_D,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f13334,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f13333])).
fof(f13333,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13318,f710])).
fof(f13318,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f819,f651])).
fof(f13331,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f13330])).
fof(f13330,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13317,f710])).
fof(f13317,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f819,f651])).
fof(f13280,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) | $spl98),
  inference(cnf_transformation,[],[f13280_D])).
fof(f13280_D,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f13282,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f13281])).
fof(f13281,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13266,f710])).
fof(f13266,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f818,f651])).
fof(f13279,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f13278])).
fof(f13278,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13265,f710])).
fof(f13265,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f818,f651])).
fof(f10904,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10841,f3979])).
fof(f3979,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f344])).
fof(f10841,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10184])).
fof(f10184,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl0),
  inference(forward_demodulation,[],[f10175,f8611])).
fof(f8611,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f414,f285])).
fof(f10175,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f787,f430,f8652,f337])).
fof(f8652,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f349])).
fof(f10733,plain,(
  s('type.extreal.extreal',sK14) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(forward_demodulation,[],[f10637,f276])).
fof(f10637,plain,(
  s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))) | ($spl0 | $spl3 | $spl5)),
  inference(superposition,[],[f280,f8930])).
fof(f8930,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f8929,f420])).
fof(f8929,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(forward_demodulation,[],[f8853,f8570])).
fof(f8570,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | ($spl3 | $spl5)),
  inference(forward_demodulation,[],[f8539,f5338])).
fof(f8539,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f3938,f419,f710])).
fof(f3938,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f249])).
fof(f8853,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl0),
  inference(superposition,[],[f414,f1695])).
fof(f10903,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10840,f3979])).
fof(f10840,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10183])).
fof(f10183,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl0),
  inference(forward_demodulation,[],[f10176,f8611])).
fof(f10176,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f841,f430,f8652,f337])).
fof(f10902,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10838,f5301])).
fof(f5301,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f3998,f331])).
fof(f3998,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f420,f417,f346])).
fof(f10838,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10178])).
fof(f10178,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f420,f436,f3973,f8652,f336])).
fof(f3973,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f341])).
fof(f10901,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10837,f4665])).
fof(f4665,plain,(
  ~sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f3981,f331])).
fof(f3981,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f420,f420,f417,f346])).
fof(f10837,plain,(
  sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10177])).
fof(f10177,plain,(
  sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f436,f430,f8652,f336])).
fof(f10900,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10835,f5285])).
fof(f5285,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f3990,f331])).
fof(f3990,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f417,f420,f417,f346])).
fof(f10835,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10173])).
fof(f10173,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f420,f3980,f430,f8652,f337])).
fof(f3980,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f345])).
fof(f10899,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10834,f3980])).
fof(f10834,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10160])).
fof(f10160,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl0),
  inference(forward_demodulation,[],[f10145,f8611])).
fof(f10145,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),s(bool,X0)),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f789,f430,f8651,f337])).
fof(f8651,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f348])).
fof(f10898,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10833,f3980])).
fof(f10833,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10159])).
fof(f10159,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl0),
  inference(forward_demodulation,[],[f10146,f8611])).
fof(f10146,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f843,f430,f8651,f337])).
fof(f10897,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10831,f5285])).
fof(f10831,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10143])).
fof(f10143,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f420,f3979,f430,f8651,f337])).
fof(f10896,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10830,f4670])).
fof(f10830,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10142])).
fof(f10142,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f434,f430,f8651,f337])).
fof(f10895,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10828,f5301])).
fof(f10828,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10112])).
fof(f10112,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f420,f436,f3972,f8650,f336])).
fof(f8650,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f347])).
fof(f3972,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f340])).
fof(f10894,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10827,f4675])).
fof(f4675,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f3997,f331])).
fof(f3997,plain,(
  ~sP2('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f420,f417,f346])).
fof(f10827,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10109])).
fof(f10109,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f436,f428,f8650,f336])).
fof(f10893,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10824,f4629])).
fof(f4629,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f3974,f332])).
fof(f3974,plain,(
  ~sP3('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f231,f417,f342])).
fof(f10824,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10092])).
fof(f10092,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f442,f426,f8649,f335])).
fof(f8649,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f343])).
fof(f10892,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10823,f4638])).
fof(f4638,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f440,f430,f3979,f337])).
fof(f10823,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10075])).
fof(f10075,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f8644,f332])).
fof(f8644,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f342])).
fof(f10891,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10822,f4656])).
fof(f4656,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f231,f442,f430,f3980,f335])).
fof(f10822,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10044])).
fof(f10044,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f8639,f332])).
fof(f8639,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f342])).
fof(f10890,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10821,f9268])).
fof(f9268,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f3938,f7722,f352])).
fof(f7722,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12)))) | $spl3),
  inference(unit_resulting_resolution,[],[f3938,f1273])).
fof(f10821,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10022])).
fof(f10022,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f231,f442,f9257,f8638,f335])).
fof(f8638,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f339])).
fof(f9257,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK12)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f7722,f343])).
fof(f10889,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10820,f4608])).
fof(f4608,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f3967,f333])).
fof(f3967,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f231,f417,f338])).
fof(f10820,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f10019])).
fof(f10019,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f442,f432,f8638,f335])).
fof(f10888,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10819,f4614])).
fof(f4614,plain,(
  sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f438,f436,f3972,f336])).
fof(f10819,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f9993])).
fof(f9993,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f8633,f333])).
fof(f8633,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f338])).
fof(f10887,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10813,f4030])).
fof(f4030,plain,(
  ( ! [X0] : (sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f366])).
fof(f10813,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8838])).
fof(f8838,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f1009])).
fof(f10886,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10812,f4030])).
fof(f10812,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8832])).
fof(f8832,plain,(
  ( ! [X0] : (~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f1008])).
fof(f10885,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10803,f4038])).
fof(f4038,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f372])).
fof(f10803,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8779])).
fof(f8779,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f452,f414,f406])).
fof(f10884,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10802,f4037])).
fof(f4037,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f371])).
fof(f10802,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'T_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8778])).
fof(f8778,plain,(
  ( ! [X0] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f450,f414,f406])).
fof(f10883,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10801,f4093])).
fof(f4093,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f389])).
fof(f10801,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0',X0)) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8777])).
fof(f8777,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f458,f414,f405])).
fof(f10882,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10800,f4091])).
fof(f4091,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f417,f388])).
fof(f10800,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8776])).
fof(f8776,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f4107,f414,f405])).
fof(f4107,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f396])).
fof(f10881,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10799,f4088])).
fof(f4088,plain,(
  ( ! [X0] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f388])).
fof(f10799,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8775])).
fof(f8775,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f456,f414,f405])).
fof(f10880,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10797,f4112])).
fof(f4112,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f420,f231,f417,f399])).
fof(f10797,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8773])).
fof(f8773,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f402])).
fof(f10879,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10796,f4114])).
fof(f4114,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f400])).
fof(f10796,plain,(
  ~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8770])).
fof(f8770,plain,(
  ~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f402])).
fof(f10878,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10794,f4090])).
fof(f4090,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f388])).
fof(f10794,plain,(
  sP11('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8759])).
fof(f8759,plain,(
  sP11('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f420,f414,f399])).
fof(f10877,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10792,f4107])).
fof(f10792,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8754])).
fof(f8754,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f420,f414,f395])).
fof(f10876,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10791,f4107])).
fof(f10791,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8753])).
fof(f8753,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f395])).
fof(f10875,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10790,f4108])).
fof(f4108,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f398])).
fof(f10790,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8748])).
fof(f8748,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f231,f414,f395])).
fof(f10874,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10789,f4091])).
fof(f10789,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8746])).
fof(f8746,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f420,f4108,f414,f394])).
fof(f10873,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10787,f4078])).
fof(f4078,plain,(
  sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f420,f420,f417,f382])).
fof(f10787,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8729])).
fof(f8729,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f420,f414,f385])).
fof(f10872,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10782,f4067])).
fof(f4067,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f381])).
fof(f10782,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8712])).
fof(f8712,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f420,f414,f378])).
fof(f10871,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10779,f4024])).
fof(f4024,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f365])).
fof(f10779,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8703])).
fof(f8703,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f414,f420,f414,f370])).
fof(f10870,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10778,f4024])).
fof(f10778,plain,(
  ~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8702])).
fof(f8702,plain,(
  ~sP7('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f370])).
fof(f10869,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10777,f4030])).
fof(f10777,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8698])).
fof(f8698,plain,(
  ~sP7('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f370])).
fof(f10868,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10773,f4036])).
fof(f4036,plain,(
  ~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f231,f231,f417,f370])).
fof(f10773,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8694])).
fof(f8694,plain,(
  sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f231,f414,f367])).
fof(f10867,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10772,f4011])).
fof(f4011,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f355])).
fof(f10772,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8690])).
fof(f8690,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f414,f414,f367])).
fof(f10866,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10771,f4011])).
fof(f10771,plain,(
  sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8688])).
fof(f8688,plain,(
  sP7('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f231,f414,f367])).
fof(f10865,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10770,f4014])).
fof(f4014,plain,(
  ( ! [X0] : (~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f356])).
fof(f10770,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8682])).
fof(f8682,plain,(
  sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f231,f414,f367])).
fof(f10864,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10768,f4022])).
fof(f4022,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f363])).
fof(f10768,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8677])).
fof(f8677,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f414,f420,f414,f361])).
fof(f10863,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10767,f4022])).
fof(f10767,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8676])).
fof(f8676,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f361])).
fof(f10862,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10766,f4023])).
fof(f4023,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f417,f364])).
fof(f10766,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8671])).
fof(f8671,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f361])).
fof(f10861,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10765,f4121])).
fof(f4121,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f448,f417,f407])).
fof(f10765,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0',X0)) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8662])).
fof(f8662,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f414,f356])).
fof(f10860,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10764,f4119])).
fof(f4119,plain,(
  ( ! [X0] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f446,f417,f407])).
fof(f10764,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'F_0')) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8661])).
fof(f8661,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f414,f355])).
fof(f10859,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10763,f4658])).
fof(f4658,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f231,f442,f3973,f3980,f335])).
fof(f10763,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8660])).
fof(f8660,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f3938,f414,f352])).
fof(f10858,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10762,f4005])).
fof(f4005,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f462,f417,f350])).
fof(f10762,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8659])).
fof(f8659,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f462,f414,f352])).
fof(f10857,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10760,f4007])).
fof(f4007,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f462,f417,f350])).
fof(f10760,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8657])).
fof(f8657,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f462,f414,f352])).
fof(f10856,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10758,f4010])).
fof(f4010,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f276,f417,f353])).
fof(f10758,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8655])).
fof(f8655,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f276,f414,f351])).
fof(f10855,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10756,f4009])).
fof(f4009,plain,(
  ( ! [X0,X1] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f276,f417,f353])).
fof(f10756,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8653])).
fof(f8653,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f276,f414,f351])).
fof(f10854,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10751,f3979])).
fof(f10751,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8645])).
fof(f8645,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f414,f420,f414,f342])).
fof(f10853,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10750,f3979])).
fof(f10750,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8644])).
fof(f10852,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10749,f3980])).
fof(f10749,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8639])).
fof(f10851,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10747,f3972])).
fof(f10747,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8634])).
fof(f8634,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f414,f420,f414,f338])).
fof(f10850,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10746,f3972])).
fof(f10746,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8633])).
fof(f10849,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10745,f3973])).
fof(f10745,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0','F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8628])).
fof(f8628,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f338])).
fof(f10848,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10744,f3949])).
fof(f3949,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f231,f417,f334])).
fof(f10744,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8627])).
fof(f8627,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f440,f426,f432,f414,f335])).
fof(f10847,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10742,f3955])).
fof(f3955,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f438,f428,f434,f417,f336])).
fof(f10742,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8623])).
fof(f8623,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f420,f414,f334])).
fof(f10846,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10741,f3945])).
fof(f3945,plain,(
  ~sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f305])).
fof(f10741,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8620])).
fof(f8620,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f420,f414,f310])).
fof(f10845,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10740,f3943])).
fof(f3943,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f420,f417,f305])).
fof(f10740,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8619])).
fof(f8619,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f420,f414,f307])).
fof(f10844,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10738,f3947])).
fof(f3947,plain,(
  sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f307])).
fof(f10738,plain,(
  ~sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8617])).
fof(f8617,plain,(
  ~sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f231,f414,f306])).
fof(f10843,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10737,f3948])).
fof(f3948,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f231,f417,f310])).
fof(f10737,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8615])).
fof(f8615,plain,(
  ~sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f231,f414,f306])).
fof(f10842,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f10734,f3938])).
fof(f10734,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))) | ($spl0 | $spl3 | $spl5)),
  inference(backward_demodulation,[],[f10733,f8611])).
fof(f8221,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) | $spl78),
  inference(cnf_transformation,[],[f8221_D])).
fof(f8221_D,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8230,plain,(
  ( ! [X17,X16] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X16,X17)) ) | $spl96),
  inference(cnf_transformation,[],[f8230_D])).
fof(f8230_D,plain,(
  ( ! [X17,X16] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X16,X17)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8145,plain,(
  ( ! [X17,X16] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X16,X17) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f461,f1695])).
fof(f8229,plain,(
  ( ! [X14,X15] : (sP10(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) | $spl94),
  inference(cnf_transformation,[],[f8229_D])).
fof(f8229_D,plain,(
  ( ! [X14,X15] : (sP10(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8144,plain,(
  ( ! [X14,X15] : (sP10(X14,X15,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f457,f1695])).
fof(f8228,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X12,X13)) ) | $spl92),
  inference(cnf_transformation,[],[f8228_D])).
fof(f8228_D,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X12,X13)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8143,plain,(
  ( ! [X12,X13] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X12,X13) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f455,f1695])).
fof(f8227,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X10,X11)) ) | $spl90),
  inference(cnf_transformation,[],[f8227_D])).
fof(f8227_D,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X10,X11)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8142,plain,(
  ( ! [X10,X11] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X10,X11) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f449,f1695])).
fof(f8226,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X9)) ) | $spl88),
  inference(cnf_transformation,[],[f8226_D])).
fof(f8226_D,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X9)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8141,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X9) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f447,f1695])).
fof(f8225,plain,(
  ( ! [X6,X7] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X6,X7)) ) | $spl86),
  inference(cnf_transformation,[],[f8225_D])).
fof(f8225_D,plain,(
  ( ! [X6,X7] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X6,X7)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8140,plain,(
  ( ! [X6,X7] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X6,X7) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f437,f1695])).
fof(f8224,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X5)) ) | $spl84),
  inference(cnf_transformation,[],[f8224_D])).
fof(f8224_D,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X5)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8139,plain,(
  ( ! [X4,X5] : (sP3(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X5) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f435,f1695])).
fof(f8223,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X2,X3)) ) | $spl82),
  inference(cnf_transformation,[],[f8223_D])).
fof(f8223_D,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X2,X3)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8138,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X2,X3) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f431,f1695])).
fof(f8222,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1)) ) | $spl80),
  inference(cnf_transformation,[],[f8222_D])).
fof(f8222_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8137,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)),X1) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) ) | $spl1),
  inference(superposition,[],[f429,f1695])).
fof(f8102,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0') | p(s(bool,X2))) ) | $spl76),
  inference(cnf_transformation,[],[f8102_D])).
fof(f8102_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8104,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0')) )),
  inference(condensation,[],[f8103])).
fof(f8103,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f8088,f710])).
fof(f8088,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1228,f651])).
fof(f8101,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0')) )),
  inference(condensation,[],[f8100])).
fof(f8100,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f8087,f710])).
fof(f8087,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1228,f651])).
fof(f8066,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0') | p(s(bool,X2))) ) | $spl74),
  inference(cnf_transformation,[],[f8066_D])).
fof(f8066_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8068,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0')) )),
  inference(condensation,[],[f8067])).
fof(f8067,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f8052,f710])).
fof(f8052,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1227,f651])).
fof(f8065,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0')) )),
  inference(condensation,[],[f8064])).
fof(f8064,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f8051,f710])).
fof(f8051,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1227,f651])).
fof(f8020,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) ) | $spl72),
  inference(cnf_transformation,[],[f8020_D])).
fof(f8020_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8022,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) )),
  inference(condensation,[],[f8021])).
fof(f8021,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,'T_0')),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f8006,f710])).
fof(f8006,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,'T_0')),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1101,f651])).
fof(f8019,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) )),
  inference(condensation,[],[f8018])).
fof(f8018,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,'T_0')),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f8005,f710])).
fof(f8005,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,'T_0')),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1101,f651])).
fof(f7971,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) ) | $spl70),
  inference(cnf_transformation,[],[f7971_D])).
fof(f7971_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7973,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) )),
  inference(condensation,[],[f7972])).
fof(f7972,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X32)),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7957,f710])).
fof(f7957,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X32)),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1100,f651])).
fof(f7970,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) )),
  inference(condensation,[],[f7969])).
fof(f7969,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X27)),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7956,f710])).
fof(f7956,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X27)),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1100,f651])).
fof(f7927,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) ) | $spl68),
  inference(cnf_transformation,[],[f7927_D])).
fof(f7927_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7929,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) )),
  inference(condensation,[],[f7928])).
fof(f7928,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7911,f710])).
fof(f7911,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1009,f651])).
fof(f7926,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)))) )),
  inference(condensation,[],[f7925])).
fof(f7925,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7910,f710])).
fof(f7910,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1009,f651])).
fof(f7880,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) ) | $spl66),
  inference(cnf_transformation,[],[f7880_D])).
fof(f7880_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7882,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) )),
  inference(condensation,[],[f7881])).
fof(f7881,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7864,f710])).
fof(f7864,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1008,f651])).
fof(f7879,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)))) )),
  inference(condensation,[],[f7878])).
fof(f7878,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7863,f710])).
fof(f7863,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP7(X26,'T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1008,f651])).
fof(f7838,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2) | p(s(bool,X2))) ) | $spl64),
  inference(cnf_transformation,[],[f7838_D])).
fof(f7838_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2) | p(s(bool,X2))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7840,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2)) )),
  inference(condensation,[],[f7839])).
fof(f7839,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7820,f710])).
fof(f7820,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f908,f651])).
fof(f7837,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2)) )),
  inference(condensation,[],[f7836])).
fof(f7836,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7819,f710])).
fof(f7819,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f908,f651])).
fof(f7799,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2) | p(s(bool,X2))) ) | $spl62),
  inference(cnf_transformation,[],[f7799_D])).
fof(f7799_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2) | p(s(bool,X2))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7801,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2)) )),
  inference(condensation,[],[f7800])).
fof(f7800,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7781,f710])).
fof(f7781,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f907,f651])).
fof(f7798,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2)) )),
  inference(condensation,[],[f7797])).
fof(f7797,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7780,f710])).
fof(f7780,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f907,f651])).
fof(f7761,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) ) | $spl58),
  inference(cnf_transformation,[],[f7761_D])).
fof(f7761_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7762,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl60),
  inference(cnf_transformation,[],[f7762_D])).
fof(f7762_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7763,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X9))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X10)))) )),
  inference(subsumption_resolution,[],[f7744,f249])).
fof(f7744,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X10)))) )),
  inference(superposition,[],[f1273,f651])).
fof(f7760,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) )),
  inference(subsumption_resolution,[],[f7743,f249])).
fof(f7743,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1273,f651])).
fof(f7130,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl56),
  inference(cnf_transformation,[],[f7130_D])).
fof(f7130_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7132,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f7131])).
fof(f7131,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7119,f710])).
fof(f7119,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f875,f651])).
fof(f7129,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f7128])).
fof(f7128,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7118,f710])).
fof(f7118,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f875,f651])).
fof(f7099,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) | $spl54),
  inference(cnf_transformation,[],[f7099_D])).
fof(f7099_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7101,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f7100])).
fof(f7100,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7088,f710])).
fof(f7088,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f874,f651])).
fof(f7098,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f7097])).
fof(f7097,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7087,f710])).
fof(f7087,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f874,f651])).
fof(f7072,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl52),
  inference(cnf_transformation,[],[f7072_D])).
fof(f7072_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7074,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f7073])).
fof(f7073,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7061,f710])).
fof(f7061,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f873,f651])).
fof(f7071,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f7070])).
fof(f7070,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7060,f710])).
fof(f7060,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f873,f651])).
fof(f7021,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) | $spl50),
  inference(cnf_transformation,[],[f7021_D])).
fof(f7021_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7023,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f7022])).
fof(f7022,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7010,f710])).
fof(f7010,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f872,f651])).
fof(f7020,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f7019])).
fof(f7019,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f7009,f710])).
fof(f7009,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f872,f651])).
fof(f6969,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl48),
  inference(cnf_transformation,[],[f6969_D])).
fof(f6969_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6971,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f6970])).
fof(f6970,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6958,f710])).
fof(f6958,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f817,f651])).
fof(f6968,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f6967])).
fof(f6967,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6957,f710])).
fof(f6957,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f817,f651])).
fof(f6936,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f6936_D])).
fof(f6936_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6938,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f6937])).
fof(f6937,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6925,f710])).
fof(f6925,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f816,f651])).
fof(f6935,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f6934])).
fof(f6934,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6924,f710])).
fof(f6924,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f816,f651])).
fof(f6909,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl44),
  inference(cnf_transformation,[],[f6909_D])).
fof(f6909_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6911,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f6910])).
fof(f6910,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6898,f710])).
fof(f6898,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f815,f651])).
fof(f6908,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f6907])).
fof(f6907,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6897,f710])).
fof(f6897,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f815,f651])).
fof(f6866,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) | $spl42),
  inference(cnf_transformation,[],[f6866_D])).
fof(f6866_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f6868,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f6867])).
fof(f6867,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6855,f710])).
fof(f6855,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f814,f651])).
fof(f6865,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f6864])).
fof(f6864,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6854,f710])).
fof(f6854,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f814,f651])).
fof(f3617,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X1)) ) | $spl40),
  inference(cnf_transformation,[],[f3617_D])).
fof(f3617_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X1)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3625,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X1)) )),
  inference(condensation,[],[f3624])).
fof(f3624,plain,(
  ( ! [X163,X161,X162,X160] : (p(s(bool,X161)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X162)),X161) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(subsumption_resolution,[],[f2786,f710])).
fof(f2786,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,X160)) | p(s(bool,X161)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X162)),X161) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f812,f651])).
fof(f3614,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) ) | $spl38),
  inference(cnf_transformation,[],[f3614_D])).
fof(f3614_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3623,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) )),
  inference(condensation,[],[f3622])).
fof(f3622,plain,(
  ( ! [X167,X165,X166,X164] : (p(s(bool,X165)) | sP1(X165,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(subsumption_resolution,[],[f2787,f710])).
fof(f2787,plain,(
  ( ! [X167,X165,X166,X164] : (~p(s(bool,X164)) | p(s(bool,X165)) | sP1(X165,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f813,f651])).
fof(f3611,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X1)) ) | $spl36),
  inference(cnf_transformation,[],[f3611_D])).
fof(f3611_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3621,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X1)) )),
  inference(condensation,[],[f3620])).
fof(f3620,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X201)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,'T_0')),X201) | s(bool,X200) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f2798,f710])).
fof(f2798,plain,(
  ( ! [X202,X200,X203,X201] : (~p(s(bool,X200)) | p(s(bool,X201)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,'T_0')),X201) | s(bool,X200) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f870,f651])).
fof(f3608,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')))) ) | $spl34),
  inference(cnf_transformation,[],[f3608_D])).
fof(f3608_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3619,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')))) )),
  inference(condensation,[],[f3618])).
fof(f3618,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X205)) | sP1(X205,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X206),s(bool,'T_0'))) | s(bool,X204) = s(bool,X207) | s(bool,'T_0') = s(bool,X207)) )),
  inference(subsumption_resolution,[],[f2799,f710])).
fof(f2799,plain,(
  ( ! [X206,X204,X207,X205] : (~p(s(bool,X204)) | p(s(bool,X205)) | sP1(X205,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X206),s(bool,'T_0'))) | s(bool,X204) = s(bool,X207) | s(bool,'T_0') = s(bool,X207)) )),
  inference(superposition,[],[f871,f651])).
fof(f3616,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X1)) )),
  inference(condensation,[],[f3615])).
fof(f3615,plain,(
  ( ! [X163,X161,X162,X160] : (p(s(bool,X161)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X162)),X161) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(subsumption_resolution,[],[f2956,f710])).
fof(f2956,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,X160)) | p(s(bool,X161)) | sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X162)),X161) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f812,f651])).
fof(f3613,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) )),
  inference(condensation,[],[f3612])).
fof(f3612,plain,(
  ( ! [X167,X165,X166,X164] : (p(s(bool,X165)) | sP1(X165,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(subsumption_resolution,[],[f2957,f710])).
fof(f2957,plain,(
  ( ! [X167,X165,X166,X164] : (~p(s(bool,X164)) | p(s(bool,X165)) | sP1(X165,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X166))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f813,f651])).
fof(f3610,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X1)) )),
  inference(condensation,[],[f3609])).
fof(f3609,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,X201)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,'T_0')),X201) | s(bool,X200) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f2968,f710])).
fof(f2968,plain,(
  ( ! [X202,X200,X203,X201] : (~p(s(bool,X200)) | p(s(bool,X201)) | sP1('const.bool.COND_3'(s(bool,'F_0'),s(bool,X202),s(bool,'T_0')),X201) | s(bool,X200) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f870,f651])).
fof(f3607,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')))) )),
  inference(condensation,[],[f3606])).
fof(f3606,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,X205)) | sP1(X205,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X206),s(bool,'T_0'))) | s(bool,X204) = s(bool,X207) | s(bool,'T_0') = s(bool,X207)) )),
  inference(subsumption_resolution,[],[f2969,f710])).
fof(f2969,plain,(
  ( ! [X206,X204,X207,X205] : (~p(s(bool,X204)) | p(s(bool,X205)) | sP1(X205,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X206),s(bool,'T_0'))) | s(bool,X204) = s(bool,X207) | s(bool,'T_0') = s(bool,X207)) )),
  inference(superposition,[],[f871,f651])).
fof(f3603,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496,f3499,f3501,f3602])).
fof(f3602,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) ) | $spl8),
  inference(forward_demodulation,[],[f3524,f276])).
fof(f3524,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'T_0'),s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f280])).
fof(f3501,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2460])).
fof(f3499,plain,(
  ( ! [X0] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2433])).
fof(f3496,plain,(
  ( ! [X0,X1] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2306])).
fof(f3495,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2285])).
fof(f3491,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2198])).
fof(f3485,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2163])).
fof(f3484,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2161])).
fof(f3483,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2141])).
fof(f3479,plain,(
  ( ! [X0,X1] : (sP4('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2125])).
fof(f3476,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X11),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2102])).
fof(f3472,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'T_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2086])).
fof(f3471,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2084])).
fof(f3470,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2054])).
fof(f3466,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2025])).
fof(f3463,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f2017])).
fof(f3462,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X2),s(bool,X1)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1816])).
fof(f3461,plain,(
  ( ! [X2,X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X2)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1812])).
fof(f3457,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1639])).
fof(f3453,plain,(
  ( ! [X14,X17,X15,X16] : (~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15)),X17,X16) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1228])).
fof(f3452,plain,(
  ( ! [X10,X11,X9] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1159])).
fof(f3449,plain,(
  ( ! [X14,X17,X15,X16] : (~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15)),X17,X16) | ~p(s(bool,X15)) | p(s(bool,X16)) | p(s(bool,X17))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1101])).
fof(f3448,plain,(
  ( ! [X10,X11,X9] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1082])).
fof(f3445,plain,(
  ( ! [X14,X17,X15,X16] : (~sP7(X17,X16,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f1009])).
fof(f3444,plain,(
  ( ! [X14,X17,X15,X16] : (~sP5(X16,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X14),s(bool,X15)),X17) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f908])).
fof(f3551,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~p(s(bool,X52))) ) | $spl8),
  inference(forward_demodulation,[],[f3443,f3330])).
fof(f3443,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | ~p(s(bool,X52)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f877])).
fof(f3549,plain,(
  ( ! [X47,X50,X48,X49] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~p(s(bool,X48))) ) | $spl8),
  inference(forward_demodulation,[],[f3442,f3330])).
fof(f3442,plain,(
  ( ! [X47,X50,X48,X49] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | ~p(s(bool,X48)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X49,X50)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f876])).
fof(f3441,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46) | ~p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f875])).
fof(f3440,plain,(
  ( ! [X39,X41,X42,X40] : (sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41) | ~p(s(bool,X40)) | p(s(bool,X41))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f874])).
fof(f3439,plain,(
  ( ! [X37,X35,X38,X36] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36)),X37,X38) | ~p(s(bool,X36)) | ~p(s(bool,X37))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f873])).
fof(f3438,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33) | ~p(s(bool,X32)) | ~p(s(bool,X33))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f872])).
fof(f3437,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X28),s(bool,X29))) | ~p(s(bool,X29)) | p(s(bool,X30))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f871])).
fof(f3436,plain,(
  ( ! [X26,X27,X25] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f870])).
fof(f3435,plain,(
  ( ! [X24,X23,X22] : (sP0('const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)),X24) | ~p(s(bool,X23))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f869])).
fof(f3434,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | ~p(s(bool,X20))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f868])).
fof(f3432,plain,(
  ( ! [X0,X1] : (sP5('T_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f844])).
fof(f3431,plain,(
  ( ! [X0,X1] : (~sP5('F_0',X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f843])).
fof(f3430,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f842])).
fof(f3429,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f841])).
fof(f3425,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f277])).
fof(f3533,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f3423,f3330])).
fof(f3423,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f3330,f253])).
fof(f3601,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3502,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496,f3499,f3501])).
fof(f3502,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0'))) | p(s(bool,X0))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2468])).
fof(f3600,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496,f3499,f3501])).
fof(f3599,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3500,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496,f3499])).
fof(f3500,plain,(
  ( ! [X0,X1] : (sP1('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')),X0) | p(s(bool,X0))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2441])).
fof(f3598,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496,f3499])).
fof(f3597,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3498,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496])).
fof(f3498,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2385])).
fof(f3596,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3497,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496])).
fof(f3497,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2373])).
fof(f3595,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495,f3496])).
fof(f3594,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491,f3495])).
fof(f3593,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3493,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491])).
fof(f3493,plain,(
  ( ! [X1] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2211])).
fof(f3592,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3492,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491])).
fof(f3492,plain,(
  ( ! [X0] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2210])).
fof(f3591,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485,f3491])).
fof(f3590,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3490,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485])).
fof(f3490,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2181])).
fof(f3589,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3489,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485])).
fof(f3489,plain,(
  ( ! [X0] : (~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2180])).
fof(f3588,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3487,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485])).
fof(f3487,plain,(
  ( ! [X1] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2175])).
fof(f3587,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3486,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485])).
fof(f3486,plain,(
  ( ! [X0] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2174])).
fof(f3586,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484,f3485])).
fof(f3585,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483,f3484])).
fof(f3584,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479,f3483])).
fof(f3583,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3481,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479])).
fof(f3481,plain,(
  ( ! [X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2137])).
fof(f3582,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3480,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479])).
fof(f3480,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2136])).
fof(f3581,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476,f3479])).
fof(f3580,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3478,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476])).
fof(f3478,plain,(
  ( ! [X0] : (~sP3('const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),'F_0')) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2105])).
fof(f3579,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3477,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476])).
fof(f3477,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2103])).
fof(f3578,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472,f3476])).
fof(f3577,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3474,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472])).
fof(f3474,plain,(
  ( ! [X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2098])).
fof(f3576,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3473,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472])).
fof(f3473,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2097])).
fof(f3575,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471,f3472])).
fof(f3574,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470,f3471])).
fof(f3573,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466,f3470])).
fof(f3572,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3468,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466])).
fof(f3468,plain,(
  ( ! [X1] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2050])).
fof(f3571,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3467,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466])).
fof(f3467,plain,(
  ( ! [X0] : (~sP2('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2033])).
fof(f3570,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463,f3466])).
fof(f3569,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3464,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463])).
fof(f3464,plain,(
  ( ! [X1] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f2021])).
fof(f3568,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462,f3463])).
fof(f3567,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461,f3462])).
fof(f3566,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457,f3461])).
fof(f3565,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3458,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457])).
fof(f3458,plain,(
  ( ! [X0] : (~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X0)),'F_0')) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1664])).
fof(f3564,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453,f3457])).
fof(f3563,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3456,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453])).
fof(f3456,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1337])).
fof(f3562,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3455,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453])).
fof(f3455,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1333])).
fof(f3561,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f3454,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453])).
fof(f3454,plain,(
  ( ! [X10,X11,X9] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X11,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X9),s(bool,X10))) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1243])).
fof(f3560,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452,f3453])).
fof(f3559,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449,f3452])).
fof(f3558,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3450,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449])).
fof(f3450,plain,(
  ( ! [X1] : (sP5('T_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1155])).
fof(f3557,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448,f3449])).
fof(f3556,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445,f3448])).
fof(f3555,plain,(
  $false | ($spl1 | $spl8)),
  inference(global_subsumption,[],[f3446,f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445])).
fof(f3446,plain,(
  ( ! [X1] : (~sP5('F_0','const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) ) | ($spl1 | $spl8)),
  inference(backward_demodulation,[],[f3330,f1078])).
fof(f3554,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444,f3445])).
fof(f3553,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551,f3444])).
fof(f3552,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549,f3551])).
fof(f3550,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441,f3549])).
fof(f3548,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440,f3441])).
fof(f3547,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439,f3440])).
fof(f3546,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438,f3439])).
fof(f3545,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437,f3438])).
fof(f3544,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436,f3437])).
fof(f3543,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435,f3436])).
fof(f3542,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434,f3435])).
fof(f3541,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432,f3434])).
fof(f3540,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431,f3432])).
fof(f3539,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430,f3431])).
fof(f3538,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429,f3430])).
fof(f3537,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425,f3429])).
fof(f3536,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3426,f231])).
fof(f3426,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f3330,f420])).
fof(f3535,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533,f3425])).
fof(f3534,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f230,f229,f228,f231,f420,f239,f250,f249,f254,f253,f276,f277,f280,f281,f285,f282,f292,f291,f421,f310,f307,f306,f305,f323,f324,f337,f336,f335,f334,f333,f332,f331,f341,f340,f339,f338,f345,f344,f343,f342,f349,f348,f347,f346,f353,f352,f351,f350,f407,f357,f356,f355,f354,f364,f363,f362,f361,f370,f367,f366,f365,f377,f376,f375,f406,f373,f372,f371,f381,f380,f379,f378,f385,f384,f383,f382,f394,f393,f392,f405,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f422,f424,f426,f428,f430,f432,f434,f436,f438,f440,f442,f444,f446,f448,f450,f452,f454,f456,f458,f460,f462,f486,f487,f490,f498,f509,f515,f518,f521,f529,f544,f545,f547,f557,f565,f571,f575,f589,f651,f710,f743,f741,f767,f764,f790,f789,f788,f787,f801,f810,f811,f812,f813,f814,f815,f816,f817,f818,f819,f844,f843,f842,f841,f845,f859,f868,f869,f870,f871,f872,f873,f874,f875,f876,f877,f907,f908,f920,f936,f945,f951,f960,f964,f974,f983,f1008,f1009,f1042,f1052,f1061,f1079,f1080,f1081,f1082,f1100,f1101,f1151,f1156,f1157,f1158,f1159,f1169,f1179,f1188,f1227,f1228,f1253,f1273,f1535,f1532,f1587,f1639,f1695,f1746,f1747,f1954,f1952,f1950,f1948,f1947,f1946,f1942,f1941,f1937,f1934,f1933,f1928,f1927,f1921,f1900,f1899,f1898,f1894,f1893,f1889,f1888,f1887,f1882,f1873,f1868,f1865,f1864,f1860,f1859,f1858,f1857,f1856,f1853,f1851,f1846,f1845,f1840,f1839,f1833,f1822,f1821,f1819,f1817,f1816,f1815,f1812,f1811,f1808,f1807,f1806,f1805,f1800,f1799,f1793,f1792,f1787,f1786,f1780,f1773,f1771,f1770,f1767,f1766,f1765,f1764,f1762,f1761,f1760,f1759,f1758,f2017,f2015,f2022,f2023,f2025,f2045,f2051,f2052,f2054,f2086,f2084,f2089,f2099,f2100,f2102,f2125,f2128,f2138,f2139,f2141,f2163,f2161,f2166,f2176,f2177,f2198,f2202,f2212,f2213,f2233,f2255,f2285,f2306,f2387,f2411,f2433,f2460,f2505,f2506,f2519,f2530,f2531,f2568,f2570,f2572,f2574,f2581,f2583,f2584,f2586,f2631,f2638,f2641,f2740,f2741,f2742,f2743,f2744,f2745,f3328,f3332,f3335,f3338,f3341,f3344,f3347,f3350,f3353,f3356,f3359,f3362,f2776,f2778,f2780,f2782,f2789,f2791,f2792,f2794,f3367,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3378,f3379,f3380,f3381,f3382,f2836,f2838,f2840,f2842,f3383,f3384,f3385,f3386,f2849,f2851,f2852,f2854,f3387,f3388,f3390,f3391,f2905,f2908,f3397,f3399,f3401,f3403,f3405,f3407,f3409,f3411,f3413,f3415,f3417,f3419,f3330,f3533])).
fof(f3366,plain,(
  ( ! [X1] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,'T_0') | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f3366_D])).
fof(f3366_D,plain,(
  ( ! [X1] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,'T_0') | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3421,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,'T_0')) ) | $spl2),
  inference(condensation,[],[f3420])).
fof(f3420,plain,(
  ( ! [X218,X216,X217] : (p(s(bool,X217)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X217,'T_0') | s(bool,X216) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) ) | $spl2),
  inference(subsumption_resolution,[],[f2974,f710])).
fof(f2974,plain,(
  ( ! [X218,X216,X217] : (~p(s(bool,X216)) | p(s(bool,X217)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X217,'T_0') | s(bool,X216) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) ) | $spl2),
  inference(superposition,[],[f1232,f651])).
fof(f3363,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f3363_D])).
fof(f3363_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3419,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3418])).
fof(f3418,plain,(
  ( ! [X123,X121,X124,X122] : (~sP8('T_0',X122,X123) | sP9('T_0',X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f2944,f710])).
fof(f2944,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X121)) | ~sP8('T_0',X122,X123) | sP9('T_0',X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f406,f651])).
fof(f3360,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f3360_D])).
fof(f3360_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3417,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3416])).
fof(f3416,plain,(
  ( ! [X118,X120,X119,X117] : (~sP10('T_0',X118,X119) | sP11('T_0',X118,X119) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f2943,f710])).
fof(f2943,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,X117)) | ~sP10('T_0',X118,X119) | sP11('T_0',X118,X119) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f405,f651])).
fof(f3357,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f3357_D])).
fof(f3357_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3415,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3414])).
fof(f3414,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | p(s(bool,X115)) | ~sP11('T_0',X115,X114) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f2942,f710])).
fof(f2942,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | p(s(bool,X115)) | ~sP11('T_0',X115,X114) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f402,f651])).
fof(f3354,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f3354_D])).
fof(f3354_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3413,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3412])).
fof(f3412,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X110)) | sP11('T_0',X110,X111) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f2941,f710])).
fof(f2941,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | ~p(s(bool,X110)) | sP11('T_0',X110,X111) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f401,f651])).
fof(f3351,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f3351_D])).
fof(f3351_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3411,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3410])).
fof(f3410,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,X106)) | sP11('T_0',X107,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f2940,f710])).
fof(f2940,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,X105)) | p(s(bool,X106)) | sP11('T_0',X107,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f400,f651])).
fof(f3348,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) ) | $spl20),
  inference(cnf_transformation,[],[f3348_D])).
fof(f3348_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3409,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3408])).
fof(f3408,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | p(s(bool,X103)) | ~sP9('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f2939,f710])).
fof(f2939,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | ~sP9('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f385,f651])).
fof(f3345,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f3345_D])).
fof(f3345_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3407,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3406])).
fof(f3406,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X98)) | sP9('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f2938,f710])).
fof(f2938,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X97)) | ~p(s(bool,X98)) | sP9('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f384,f651])).
fof(f3342,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f3342_D])).
fof(f3342_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3405,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3404])).
fof(f3404,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X94)) | sP9('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f2937,f710])).
fof(f2937,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | ~p(s(bool,X94)) | sP9('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f383,f651])).
fof(f3339,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f3339_D])).
fof(f3339_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3403,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f3402])).
fof(f3402,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X90)) | p(s(bool,X91)) | ~sP7(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f2936,f710])).
fof(f2936,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | ~sP7(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f370,f651])).
fof(f3336,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP5(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f3336_D])).
fof(f3336_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP5(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3401,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f3400])).
fof(f3400,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP5(X86,'T_0',X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f2935,f710])).
fof(f2935,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP5(X86,'T_0',X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f334,f651])).
fof(f3333,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f3333_D])).
fof(f3333_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3399,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3398])).
fof(f3398,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X83)) | sP1(X83,'T_0') | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f2934,f710])).
fof(f2934,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | p(s(bool,X83)) | sP1(X83,'T_0') | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f310,f651])).
fof(f3329,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f3329_D])).
fof(f3329_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3397,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3396])).
fof(f3396,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X80)) | sP1('T_0',X80) | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f2933,f710])).
fof(f2933,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP1('T_0',X80) | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f307,f651])).
fof(f3365,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X1,'T_0')) ) | $spl2),
  inference(condensation,[],[f3364])).
fof(f3364,plain,(
  ( ! [X218,X216,X217] : (p(s(bool,X217)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X217,'T_0') | s(bool,X216) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) ) | $spl2),
  inference(subsumption_resolution,[],[f2804,f710])).
fof(f2804,plain,(
  ( ! [X218,X216,X217] : (~p(s(bool,X216)) | p(s(bool,X217)) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)),X217,'T_0') | s(bool,X216) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) ) | $spl2),
  inference(superposition,[],[f1232,f651])).
fof(f3362,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3361])).
fof(f3361,plain,(
  ( ! [X123,X121,X124,X122] : (~sP8('T_0',X122,X123) | sP9('T_0',X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f2774,f710])).
fof(f2774,plain,(
  ( ! [X123,X121,X124,X122] : (~p(s(bool,X121)) | ~sP8('T_0',X122,X123) | sP9('T_0',X122,X123) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f406,f651])).
fof(f3359,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3358])).
fof(f3358,plain,(
  ( ! [X118,X120,X119,X117] : (~sP10('T_0',X118,X119) | sP11('T_0',X118,X119) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f2773,f710])).
fof(f2773,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,X117)) | ~sP10('T_0',X118,X119) | sP11('T_0',X118,X119) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f405,f651])).
fof(f3356,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3355])).
fof(f3355,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | p(s(bool,X115)) | ~sP11('T_0',X115,X114) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f2772,f710])).
fof(f2772,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | p(s(bool,X115)) | ~sP11('T_0',X115,X114) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f402,f651])).
fof(f3353,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3352])).
fof(f3352,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X110)) | sP11('T_0',X110,X111) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f2771,f710])).
fof(f2771,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,X109)) | ~p(s(bool,X110)) | sP11('T_0',X110,X111) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f401,f651])).
fof(f3350,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3349])).
fof(f3349,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,X106)) | sP11('T_0',X107,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f2770,f710])).
fof(f2770,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,X105)) | p(s(bool,X106)) | sP11('T_0',X107,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f400,f651])).
fof(f3347,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3346])).
fof(f3346,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | p(s(bool,X103)) | ~sP9('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f2769,f710])).
fof(f2769,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X101)) | p(s(bool,X102)) | p(s(bool,X103)) | ~sP9('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f385,f651])).
fof(f3344,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3343])).
fof(f3343,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X98)) | sP9('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f2768,f710])).
fof(f2768,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X97)) | ~p(s(bool,X98)) | sP9('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f384,f651])).
fof(f3341,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3340])).
fof(f3340,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X94)) | sP9('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f2767,f710])).
fof(f2767,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | ~p(s(bool,X94)) | sP9('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f383,f651])).
fof(f3338,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f3337])).
fof(f3337,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X90)) | p(s(bool,X91)) | ~sP7(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f2766,f710])).
fof(f2766,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | ~sP7(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f370,f651])).
fof(f3335,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f3334])).
fof(f3334,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP5(X86,'T_0',X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f2765,f710])).
fof(f2765,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP5(X86,'T_0',X87) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f334,f651])).
fof(f3332,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3331])).
fof(f3331,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X83)) | sP1(X83,'T_0') | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f2764,f710])).
fof(f2764,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X82)) | p(s(bool,X83)) | sP1(X83,'T_0') | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f310,f651])).
fof(f3328,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3327])).
fof(f3327,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X80)) | sP1('T_0',X80) | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f2763,f710])).
fof(f2763,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP1('T_0',X80) | s(bool,X79) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f307,f651])).
fof(f414,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl0),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f418,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | $spl4),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f416,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | $spl2),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f228,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))))),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))))) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f153])).
fof(f153,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) | (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) | (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) <~> (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))))) <=> (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) <=> (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))) <=> (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))))))),
  file('thm.extreal.le_max','thm.extreal.le_max')).
fof(f415,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) | $spl1),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f419,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | $spl5),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f229,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))))),
  inference(cnf_transformation,[],[f154])).
fof(f417,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | $spl3),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f230,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK14)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_max_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))))),
  inference(cnf_transformation,[],[f154])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1352
% Passive clauses: 3664
% Generated clauses: 21546
% Final active clauses: 503
% Final passive clauses: 790
% Input formulas: 43
% Initial clauses: 177
% 
% Duplicate literals: 326
% Trivial inequalities: 33
% Fw subsumption resolutions: 848
% Fw demodulations: 2833
% Bw demodulations: 655
% Condensations: 744
% Global subsumptions: 86
% 
% Simple tautologies: 83
% Equational tautologies: 97
% Forward subsumptions: 12565
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 8290
% Factoring: 3
% Forward superposition: 4167
% Backward superposition: 2850
% Self superposition: 43
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 95
% Split components: 196
% SAT solver clauses: 4114
% SAT solver unit clauses: 255
% SAT solver binary clauses: 1876
% 
% Sat splits: 95
% Sat splitting refutations: 172
% 
% Memory used [KB]: 8315
% Time elapsed: 0.742 s
% ------------------------------
% ------------------------------
% Success in time 2.06 s
