Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.int_arith
% (63712)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 917
% Passive clauses: 9085
% Generated clauses: 23277
% Final active clauses: 798
% Final passive clauses: 7483
% Input formulas: 60
% Initial clauses: 223
% 
% Duplicate literals: 609
% Trivial inequalities: 121
% Fw subsumption resolutions: 2447
% Fw demodulations: 560
% Bw demodulations: 111
% 
% Simple tautologies: 283
% Equational tautologies: 65
% Forward subsumptions: 9424
% 
% Binary resolution: 11026
% Factoring: 1560
% Forward superposition: 3557
% Backward superposition: 2576
% Self superposition: 65
% Equality factoring: 28
% Equality resolution: 41
% 
% Split clauses: 708
% Split components: 1423
% SAT solver clauses: 1277
% SAT solver unit clauses: 19
% SAT solver binary clauses: 159
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 708
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1276
% TWLsolver calls for satisfiability: 568
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.int_arith
% (63729)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 473
% Passive clauses: 2497
% Generated clauses: 17007
% Final active clauses: 243
% Final passive clauses: 1121
% Input formulas: 60
% Initial clauses: 223
% 
% Duplicate literals: 190
% Trivial inequalities: 47
% Fw subsumption resolutions: 1180
% Fw demodulations: 1988
% Bw demodulations: 1019
% Condensations: 305
% Global subsumptions: 22
% 
% Simple tautologies: 291
% Equational tautologies: 42
% Forward subsumptions: 6074
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1790
% Factoring: 63
% Forward superposition: 2147
% Backward superposition: 3201
% Self superposition: 127
% Equality factoring: 32
% Equality resolution: 43
% 
% Split clauses: 228
% Split components: 467
% SAT solver clauses: 3585
% SAT solver unit clauses: 214
% SAT solver binary clauses: 128
% SAT solver learnt clauses: 53
% SAT solver learnt literals: 40
% 
% Sat splits: 228
% Sat splitting refutations: 104
% 
% TWLsolver clauses: 3488
% TWLsolver calls for satisfiability: 1799
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.int_arith
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.int_arith
% SZS output start Proof for thm.int_arith
fof(f7634,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f305,f541,f539,f304,f540,f538,f1586,f1595,f1602,f1601,f1937,f1938,f1978,f1981,f1980,f1979,f1983,f1982,f2508,f2512,f2516,f2519,f2522,f2525,f2528,f2531,f2534,f2537,f2540,f2543,f2546,f2549,f2552,f2555,f2558,f2588,f2509,f2590,f2513,f2593,f2517,f2595,f2520,f2597,f2523,f2599,f2526,f2601,f2529,f2603,f2532,f2605,f2535,f2607,f2538,f2609,f2541,f2611,f2544,f2613,f2547,f2615,f2550,f2617,f2553,f2619,f2556,f2621,f2559,f2645,f2646,f3496,f3565,f3679,f3686,f3690,f3701,f3705,f4052,f4053,f3625,f3564,f4192,f4311,f4312,f4313,f4314,f4315,f4316,f4319,f4320,f4321,f4322,f4324,f4325,f4326,f4327,f4328,f4329,f4330,f4331,f4332,f4333,f4334,f4444,f4447,f4692,f4693,f4727,f4740,f4777,f4741,f4778,f4775,f4781,f4780,f4874,f4879,f4880,f4881,f4882,f4883,f4884,f4885,f4886,f4887,f4889,f4890,f4891,f4892,f4893,f4894,f4895,f4896,f4897,f4902,f4904,f4907,f4905,f4919,f5030,f5031,f5032,f5033,f5142,f5147,f5479,f5480,f5487,f5484,f5481,f5482,f5490,f5483,f5485,f5489,f5551,f5552,f5569,f4776,f5635,f5636,f5648,f5646,f5637,f5638,f5650,f5639,f5492,f5649,f5681,f5682,f4779,f5688,f5793,f5818,f6627,f6630,f6628,f6673,f6676,f6674,f6720,f6723,f6721,f6769,f6772,f2510,f6770,f7242,f7244,f7254,f7255,f7256,f5541,f4075,f7259,f7429,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f4318,f4317,f7442,f7444,f7445,f7446,f7447,f7448,f7449,f7450,f7451,f7452,f7457,f7488,f7490,f7499,f7501,f7503,f7507,f7509,f7511,f7513,f7515,f7517,f7521,f7524,f7538,f7542,f7544,f7546,f7548,f7550,f7551,f7552,f7553,f7554,f7620,f7621,f7622,f7630,f7631,f7632,f7633])).
fof(f7633,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7426,f756])).
fof(f756,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f579,f542,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f281])).
fof(f281,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f542,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f309,f306])).
fof(f306,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.TRUTH')).
fof(f309,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.NOT_CLAUSES|split|2')).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f383,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f383,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X7))))),
  file('thm.int_arith.eq_justify_multiplication','thm.integer.INT_LT_REFL')).
fof(f7426,plain,(
  ~sP7('F_0',g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7215])).
fof(f7215,plain,(
  ~sP7('F_0',g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | $spl84),
  inference(unit_resulting_resolution,[],[f7195,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f282])).
fof(f7195,plain,(
  ~sP6(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f542,f6724])).
fof(f6724,plain,(
  ( ! [X22] : (~sP6(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22) | p(s(bool,X22))) ) | $spl84),
  inference(subsumption_resolution,[],[f6713,f306])).
fof(f6713,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | ~sP6(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22)) ) | $spl84),
  inference(superposition,[],[f989,f5489])).
fof(f989,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(subsumption_resolution,[],[f983,f306])).
fof(f983,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(superposition,[],[f482,f591])).
fof(f591,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f303,f327])).
fof(f327,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.EQ_CLAUSES|split|1')).
fof(f303,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) & (s('type.integer.int',sK13) = s('type.integer.int',sK14) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14)))) & (s('type.integer.int',sK13) != s('type.integer.int',sK14) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f200])).
fof(f200,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & (s('type.integer.int',X1) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) & (s('type.integer.int',X1) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & ((s('type.integer.int',X1) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) & (s('type.integer.int',X1) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & (s('type.integer.int',X1) = s('type.integer.int',X2) <~> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) => (s('type.integer.int',X1) = s('type.integer.int',X2) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) => (s('type.integer.int',X7) = s('type.integer.int',X8) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) => (s('type.integer.int',X7) = s('type.integer.int',X8) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.int_arith.eq_justify_multiplication')).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f284])).
fof(f7632,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7424,f579])).
fof(f7424,plain,(
  ~sP6(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7195])).
fof(f7631,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7420,f7456])).
fof(f7456,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(subsumption_resolution,[],[f7455,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f383,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f275])).
fof(f275,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f7455,plain,(
  ( ! [X0,X1] : (~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(forward_demodulation,[],[f7454,f7259])).
fof(f7454,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(forward_demodulation,[],[f7453,f7259])).
fof(f7453,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(subsumption_resolution,[],[f7323,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f383,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f272])).
fof(f272,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f7323,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f1392])).
fof(f1392,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f1391,f561])).
fof(f561,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f303,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1391,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0) | ~sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(subsumption_resolution,[],[f1379,f306])).
fof(f1379,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0) | ~sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(superposition,[],[f456,f591])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f269])).
fof(f269,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f7420,plain,(
  ~sP5('F_0',g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7161])).
fof(f7161,plain,(
  ~sP5('F_0',g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | $spl84),
  inference(unit_resulting_resolution,[],[f7141,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f270])).
fof(f7141,plain,(
  ~sP3(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f542,f6677])).
fof(f6677,plain,(
  ( ! [X22] : (~sP3(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22) | p(s(bool,X22))) ) | $spl84),
  inference(subsumption_resolution,[],[f6666,f306])).
fof(f6666,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | ~sP3(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22)) ) | $spl84),
  inference(superposition,[],[f947,f5489])).
fof(f947,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(subsumption_resolution,[],[f941,f306])).
fof(f941,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(superposition,[],[f463,f591])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f7630,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7418,f577])).
fof(f7418,plain,(
  ~sP3(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7141])).
fof(f7622,plain,(
  $false | ($spl83 | $spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7566,f7555])).
fof(f7555,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7416,f575])).
fof(f7416,plain,(
  ( ! [X0] : (~sP4(g(s(bool,X0)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7119])).
fof(f7119,plain,(
  ( ! [X0] : (~sP4(g(s(bool,X0)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl84),
  inference(superposition,[],[f7105,f314])).
fof(f314,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X9))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.BOOL_CASES_AX')).
fof(f7105,plain,(
  ~sP4(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f542,f6631])).
fof(f6631,plain,(
  ( ! [X22] : (~sP4(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22) | p(s(bool,X22))) ) | $spl84),
  inference(subsumption_resolution,[],[f6620,f306])).
fof(f6620,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | ~sP4(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X22)) ) | $spl84),
  inference(superposition,[],[f923,f5489])).
fof(f923,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(subsumption_resolution,[],[f917,f306])).
fof(f917,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(superposition,[],[f459,f591])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f273])).
fof(f7566,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl83 | $spl84 | $spl110)),
  inference(backward_demodulation,[],[f7555,f5487])).
fof(f7621,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7558,f306])).
fof(f7558,plain,(
  ~p(s(bool,'T_0')) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7555,f542])).
fof(f7620,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7619,f7443])).
fof(f7443,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl110),
  inference(subsumption_resolution,[],[f7304,f575])).
fof(f7304,plain,(
  ( ! [X0] : (~sP4('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f1035])).
fof(f1035,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f1030,f306])).
fof(f1030,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f923,f591])).
fof(f7619,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7556,f7555])).
fof(f7556,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7555,f330])).
fof(f330,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.EQ_CLAUSES|split|3')).
fof(f7554,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7415,f7456])).
fof(f7415,plain,(
  ~sP5(g(s(bool,'F_0')),'F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7118])).
fof(f7118,plain,(
  ~sP5(g(s(bool,'F_0')),'F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | $spl84),
  inference(unit_resulting_resolution,[],[f7105,f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f270])).
fof(f7553,plain,(
  $false | ($spl84 | $spl110)),
  inference(subsumption_resolution,[],[f7413,f575])).
fof(f7413,plain,(
  ~sP4(g(s(bool,'F_0')),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl84 | $spl110)),
  inference(backward_demodulation,[],[f7259,f7105])).
fof(f7552,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7403,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP10('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f383,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f299])).
fof(f299,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f7403,plain,(
  ~sP10('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','T_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f6881])).
fof(f6881,plain,(
  ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f6876,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f296])).
fof(f296,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f6876,plain,(
  ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f1278])).
fof(f1278,plain,(
  ( ! [X0,X1] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1272,f306])).
fof(f1272,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(superposition,[],[f523,f591])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f180,f197,f196])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X19)) => p(s(bool,X18)))) <=> ((p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X18))) & (~p(s(bool,X19)) | p(s(bool,X18)) | ~p(s(bool,X20)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.sat.sat_27_prove_ltb')).
fof(f7551,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7402,f1219])).
fof(f1219,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f583,f306,f542,f515])).
fof(f7402,plain,(
  ~sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','T_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f6876])).
fof(f7550,plain,(
  $false | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7549,f7359])).
fof(f7359,plain,(
  ( ! [X0] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK12)))) ) | ($spl66 | $spl110)),
  inference(backward_demodulation,[],[f7259,f5760])).
fof(f5760,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl66),
  inference(condensation,[],[f5759])).
fof(f5759,plain,(
  ( ! [X14,X12,X10,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X13)))) ) | $spl66),
  inference(subsumption_resolution,[],[f5742,f4778])).
fof(f5742,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X13)))) )),
  inference(superposition,[],[f1584,f5081])).
fof(f5081,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) <=> (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7,X8,X17] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X17))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X17))) <=> (s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X7) = s('type.integer.int',X8)))),
  file('thm.int_arith.eq_justify_multiplication','thm.integer.INT_EQ_RMUL')).
fof(f1584,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))),s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) )),
  inference(superposition,[],[f437,f390])).
fof(f390,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X8,X7] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X7)))),
  file('thm.int_arith.eq_justify_multiplication','thm.integer.INT_MUL_SYM')).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f7549,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7383,f7360])).
fof(f7360,plain,(
  ( ! [X6] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',X6)))) ) | ($spl66 | $spl110)),
  inference(backward_demodulation,[],[f7259,f5762])).
fof(f5762,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) ) | $spl66),
  inference(backward_demodulation,[],[f5760,f4663])).
fof(f4663,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) )),
  inference(superposition,[],[f1547,f390])).
fof(f1547,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(equality_resolution,[],[f438])).
fof(f7383,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f6148])).
fof(f6148,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6147,f5760])).
fof(f6147,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6135,f390])).
fof(f6135,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f1600,f5820,f1587])).
fof(f1587,plain,(
  ( ! [X2,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int',X3)) )),
  inference(superposition,[],[f437,f390])).
fof(f5820,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl66 | $spl72)),
  inference(forward_demodulation,[],[f5768,f5762])).
fof(f5768,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))) | ($spl66 | $spl72)),
  inference(backward_demodulation,[],[f5760,f4781])).
fof(f1600,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(cnf_transformation,[],[f1600_D])).
fof(f1600_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f7548,plain,(
  $false | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7547,f7359])).
fof(f7547,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int',sK12))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7382,f7360])).
fof(f7382,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f6146])).
fof(f6146,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6136,f5762])).
fof(f6136,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5784,f5820,f1587])).
fof(f5784,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ($spl5 | $spl66)),
  inference(backward_demodulation,[],[f5762,f5090])).
fof(f5090,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f1600,f1600,f437])).
fof(f7546,plain,(
  $false | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7545,f7360])).
fof(f7545,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7381,f7360])).
fof(f7381,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f6144])).
fof(f6144,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6143,f5760])).
fof(f6143,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6138,f390])).
fof(f6138,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) | ($spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f1600,f5820,f437])).
fof(f7544,plain,(
  $false | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7543,f7360])).
fof(f7543,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7380,f7360])).
fof(f7380,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f6142])).
fof(f6142,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6139,f5762])).
fof(f6139,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5784,f5820,f437])).
fof(f7542,plain,(
  $false | ($spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7541,f7359])).
fof(f7541,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7379,f7360])).
fof(f7379,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f6141])).
fof(f6141,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl66 | $spl72)),
  inference(forward_demodulation,[],[f6140,f5762])).
fof(f6140,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5820,f5820,f437])).
fof(f7538,plain,(
  $false | ($spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7369,f7360])).
fof(f7369,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7259,f5820])).
fof(f7524,plain,(
  $false | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7523,f7359])).
fof(f7523,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7485,f7359])).
fof(f7485,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int',sK12))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f7078])).
fof(f7078,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5820,f4500,f437])).
fof(f4500,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ($spl2 | $spl5 | $spl53)),
  inference(forward_demodulation,[],[f4468,f540])).
fof(f4468,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ($spl5 | $spl53)),
  inference(unit_resulting_resolution,[],[f3563,f1600,f437])).
fof(f3563,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK14) | $spl53),
  inference(cnf_transformation,[],[f3563_D])).
fof(f3563_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK14) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f7521,plain,(
  $false | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7520,f7359])).
fof(f7520,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7483,f7360])).
fof(f7483,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f7075])).
fof(f7075,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl2 | $spl5 | $spl53 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5820,f4500,f1587])).
fof(f7517,plain,(
  $false | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7516,f7360])).
fof(f7516,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7481,f7360])).
fof(f7481,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6358])).
fof(f6358,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) | ($spl1 | $spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6357,f390])).
fof(f6357,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl1 | $spl5 | $spl66 | $spl72)),
  inference(forward_demodulation,[],[f6354,f390])).
fof(f6354,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl1 | $spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5820,f4467,f437])).
fof(f4467,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f539,f1600,f437])).
fof(f7515,plain,(
  $false | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7514,f7359])).
fof(f7514,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int',sK12))) | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7480,f7360])).
fof(f7480,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) | ($spl1 | $spl5 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6338])).
fof(f6338,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl1 | $spl5 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5820,f4467,f1587])).
fof(f7513,plain,(
  $false | ($spl2 | $spl14 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7512,f7359])).
fof(f7512,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl2 | $spl14 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7479,f7359])).
fof(f7479,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl2 | $spl14 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6126])).
fof(f6126,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl2 | $spl14 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5515,f5820,f437])).
fof(f5515,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl2 | $spl14)),
  inference(superposition,[],[f1983,f540])).
fof(f7511,plain,(
  $false | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7510,f7359])).
fof(f7510,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7478,f7359])).
fof(f7478,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X0))),s('type.integer.int',sK12)))) ) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6125])).
fof(f6125,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X0))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) | ($spl14 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5511,f5820,f437])).
fof(f5511,plain,(
  ( ! [X0] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X0)))) ) | $spl14),
  inference(superposition,[],[f1983,f390])).
fof(f7509,plain,(
  $false | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7508,f7359])).
fof(f7508,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7477,f7359])).
fof(f7477,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))),s('type.integer.int',sK12)))) ) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6124])).
fof(f6124,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) | ($spl14 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f1983,f5820,f437])).
fof(f7507,plain,(
  $false | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7506,f7360])).
fof(f7506,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7505,f7360])).
fof(f7505,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7504,f390])).
fof(f7504,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7476,f390])).
fof(f7476,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6121])).
fof(f6121,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl1 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f539,f5820,f437])).
fof(f7503,plain,(
  $false | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7502,f7359])).
fof(f7502,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7475,f7360])).
fof(f7475,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X0)))))) ) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6105])).
fof(f6105,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X0)))))) ) | ($spl14 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f5511,f5820,f1587])).
fof(f7501,plain,(
  $false | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7500,f7359])).
fof(f7500,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7474,f7360])).
fof(f7474,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14)))))) ) | ($spl14 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6104])).
fof(f6104,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14)))))) ) | ($spl14 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f1983,f5820,f1587])).
fof(f7499,plain,(
  $false | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7498,f7360])).
fof(f7498,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7497,f7360])).
fof(f7497,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7473,f390])).
fof(f7473,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6101])).
fof(f6101,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl1 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f539,f5820,f1587])).
fof(f7490,plain,(
  $false | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7489,f7359])).
fof(f7489,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7461,f7360])).
fof(f7461,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6122])).
fof(f6122,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl53 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f3563,f5820,f437])).
fof(f7488,plain,(
  $false | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f7487,f7360])).
fof(f7487,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(forward_demodulation,[],[f7460,f7360])).
fof(f7460,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) | ($spl53 | $spl66 | $spl72 | $spl110)),
  inference(backward_demodulation,[],[f7360,f6102])).
fof(f6102,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | ($spl53 | $spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f3563,f5820,f1587])).
fof(f7457,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7329,f1175])).
fof(f1175,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f610,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f140,f188,f187,f186,f185])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> s(bool,X18) = s(bool,X19)) <=> ((p(s(bool,X20)) | p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X20)) | ~p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X20))) & (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.sat.sat_2|49|_prove_ltb')).
fof(f610,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f383,f331])).
fof(f331,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f7329,plain,(
  sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1469])).
fof(f1469,plain,(
  sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f559,f554,f550,f542,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f270])).
fof(f550,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f542,f461])).
fof(f554,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f542,f465])).
fof(f559,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f303,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f7452,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7322,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (sP3('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f383,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f276])).
fof(f7322,plain,(
  ~sP3('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1377])).
fof(f1377,plain,(
  ~sP3('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f560,f551,f882,f306,f456])).
fof(f882,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f542,f591,f474])).
fof(f551,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f542,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f273])).
fof(f560,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f306,f470])).
fof(f7451,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7319,f724])).
fof(f724,plain,(
  ( ! [X0,X1] : (sP11('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f582,f306,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f297])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X2)))) )),
  inference(unit_resulting_resolution,[],[f383,f517])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f300])).
fof(f7319,plain,(
  ~sP11('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f1353])).
fof(f1353,plain,(
  ~sP11('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f1343,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f297])).
fof(f1343,plain,(
  ~sP10('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f542,f306,f1240])).
fof(f1240,plain,(
  ( ! [X0,X1] : (~sP10(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1233,f306])).
fof(f1233,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(superposition,[],[f516,f591])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f7450,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7318,f582])).
fof(f7318,plain,(
  ~sP10('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f1343])).
fof(f7449,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7315,f1079])).
fof(f1079,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f581,f542,f542,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f383,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f291])).
fof(f291,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f7315,plain,(
  ~sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1157])).
fof(f1157,plain,(
  ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1144,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f289])).
fof(f1144,plain,(
  ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f542,f1101])).
fof(f1101,plain,(
  ( ! [X0,X1] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1095,f306])).
fof(f1095,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(superposition,[],[f499,f591])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f292])).
fof(f7448,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7314,f581])).
fof(f7314,plain,(
  ~sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1144])).
fof(f7447,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7310,f756])).
fof(f7310,plain,(
  ~sP7('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f1086])).
fof(f1086,plain,(
  ~sP7('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f1065,f475])).
fof(f1065,plain,(
  ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f989])).
fof(f7446,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7308,f579])).
fof(f7308,plain,(
  ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1065])).
fof(f7445,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7306,f577])).
fof(f7306,plain,(
  ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1046])).
fof(f1046,plain,(
  ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f947])).
fof(f7444,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7305,f1178])).
fof(f1178,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f306,f610,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f279])).
fof(f7305,plain,(
  ~sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f1039])).
fof(f1039,plain,(
  ~sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f1025,f454])).
fof(f1025,plain,(
  ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f923])).
fof(f7442,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7303,f575])).
fof(f7303,plain,(
  ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f1025])).
fof(f4317,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl58),
  inference(cnf_transformation,[],[f4317_D])).
fof(f4317_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4318,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl60),
  inference(cnf_transformation,[],[f4318_D])).
fof(f4318_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7441,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(global_subsumption,[],[f305,f304,f303,f306,f314,f327,f326,f383,f390,f394,f391,f405,f404,f543,f423,f420,f419,f418,f436,f439,f458,f457,f456,f455,f454,f453,f452,f462,f461,f460,f459,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f528,f478,f477,f476,f475,f485,f484,f483,f482,f491,f488,f487,f486,f498,f497,f496,f527,f494,f493,f492,f502,f501,f500,f499,f506,f505,f504,f503,f515,f514,f513,f526,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f544,f546,f548,f550,f551,f552,f554,f555,f556,f558,f560,f562,f564,f565,f566,f568,f570,f571,f572,f574,f583,f582,f581,f580,f579,f578,f577,f576,f575,f587,f586,f616,f618,f622,f630,f641,f647,f650,f653,f656,f663,f664,f666,f675,f684,f690,f696,f709,f724,f758,f756,f769,f831,f855,f854,f875,f873,f900,f899,f897,f896,f916,f940,f951,f957,f966,f970,f982,f993,f1079,f1078,f1094,f1105,f1178,f1177,f1176,f1175,f1219,f1218,f1232,f1244,f1376,f1371,f1370,f1364,f1428,f1413,f1471,f1518,f1520,f2092,f2130,f2193,f2236,f2508,f2512,f2516,f2519,f2522,f2525,f2528,f2531,f2534,f2537,f2540,f2543,f2546,f2549,f2552,f2555,f2558,f2560,f2561,f2562,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2305,f2588,f2590,f2593,f2595,f2597,f2599,f2601,f2603,f2605,f2607,f2609,f2611,f2613,f2615,f2617,f2619,f2621,f2304,f2303,f2192,f2191,f386,f2581,f387,f2582,f2190,f2302,f1566,f1565,f1499,f1487,f610,f542,f331,f330,f3214,f3215,f3217,f3253,f1547,f4692,f4693,f4740,f4741,f4491,f4476,f4455,f1593,f1587,f1584,f1545,f1474,f1469,f1452,f1430,f1425,f1421,f1414,f1392,f1377,f1367,f1365,f1353,f1343,f1278,f1240,f1157,f1144,f1143,f1135,f1101,f1086,f1075,f1065,f1056,f1046,f1039,f1035,f1025,f1024,f989,f971,f947,f923,f884,f882,f863,f862,f761,f759,f748,f746,f744,f728,f726,f715,f711,f698,f693,f691,f683,f681,f677,f668,f646,f632,f627,f624,f591,f573,f569,f567,f563,f561,f559,f557,f553,f549,f547,f545,f438,f437,f5081,f4694,f4690,f4684,f4663,f4653,f5446,f5464,f5479,f5480,f5481,f5482,f5483,f5528,f5530,f5532,f5537,f5539,f5542,f5544,f5546,f5549,f5551,f5552,f5553,f5556,f5589,f5591,f5593,f5596,f5598,f5600,f5603,f5605,f5607,f5612,f5635,f5636,f5637,f5638,f5639,f5652,f5665,f5667,f5669,f5672,f5674,f5676,f5681,f5682,f5706,f5708,f5710,f5719,f5721,f5723,f5726,f5728,f5730,f5735,f5740,f5745,f5835,f5847,f5849,f5852,f5854,f6250,f6263,f6276,f6625,f6627,f6630,f6671,f6673,f6676,f6718,f6720,f6723,f6767,f6769,f6772,f6876,f6881,f7181,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7274,f7275,f7276,f7277,f7278,f7279,f7280,f7281,f7282,f7283,f7284,f7285,f7286,f7287,f7431,f7289,f7290,f7433,f7438,f7439,f7300,f7440,f7302])).
fof(f7302,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f1024])).
fof(f7300,plain,(
  sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','T_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f971])).
fof(f7433,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(subsumption_resolution,[],[f7432,f581])).
fof(f7432,plain,(
  ( ! [X0,X1] : (~sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(forward_demodulation,[],[f7291,f7259])).
fof(f7291,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f748])).
fof(f7290,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f746])).
fof(f7289,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f744])).
fof(f7431,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(subsumption_resolution,[],[f7430,f583])).
fof(f7430,plain,(
  ( ! [X0,X1] : (~sP10('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(forward_demodulation,[],[f7288,f7259])).
fof(f7288,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f728])).
fof(f7287,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f726])).
fof(f7286,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'F_0')) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f715])).
fof(f7285,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f711])).
fof(f7284,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f698])).
fof(f7283,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f693])).
fof(f7282,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0',X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f691])).
fof(f7281,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f683])).
fof(f7280,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'T_0')) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f681])).
fof(f7279,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f677])).
fof(f7278,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | ~p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f668])).
fof(f7277,plain,(
  sP1('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f646])).
fof(f7276,plain,(
  ( ! [X0] : (sP1(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f632])).
fof(f7275,plain,(
  sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f627])).
fof(f7274,plain,(
  ( ! [X0] : (sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f624])).
fof(f7272,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f573])).
fof(f7271,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f569])).
fof(f7270,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f567])).
fof(f7269,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f563])).
fof(f7268,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f561])).
fof(f7267,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f559])).
fof(f7266,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f557])).
fof(f7265,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f553])).
fof(f7264,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f549])).
fof(f7263,plain,(
  ( ! [X0] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f547])).
fof(f7262,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f545])).
fof(f7261,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f438])).
fof(f7260,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f437])).
fof(f7181,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f582,f728])).
fof(f6767,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6755,f306])).
fof(f6755,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | ~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1024,f769])).
fof(f6718,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6705,f306])).
fof(f6705,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | ~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f989,f769])).
fof(f6671,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6658,f306])).
fof(f6658,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | ~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f947,f769])).
fof(f6625,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6612,f306])).
fof(f6612,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f923,f769])).
fof(f6276,plain,(
  ( ! [X2,X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6266,f306])).
fof(f6266,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f711,f769])).
fof(f6263,plain,(
  ( ! [X2,X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6253,f306])).
fof(f6253,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f677,f769])).
fof(f6250,plain,(
  ( ! [X2,X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6240,f306])).
fof(f6240,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f668,f769])).
fof(f5854,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))),s('type.integer.int',X12))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))),s('type.integer.int',X12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))),s('type.integer.int',X12))),s('type.integer.int',X14))) | s('type.integer.int',X13) = s('type.integer.int',X15)) )),
  inference(superposition,[],[f1593,f1518])).
fof(f5852,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X3))) | s('type.integer.int',X2) = s('type.integer.int',X4)) )),
  inference(superposition,[],[f1593,f390])).
fof(f5849,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))),s('type.integer.int',X12))),s('type.integer.int',X14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))),s('type.integer.int',X12))),s('type.integer.int',X14))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))),s('type.integer.int',X12))))) | s('type.integer.int',X13) = s('type.integer.int',X15)) )),
  inference(superposition,[],[f1593,f1518])).
fof(f5847,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X3))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X3))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | s('type.integer.int',X2) = s('type.integer.int',X4)) )),
  inference(superposition,[],[f1593,f390])).
fof(f5835,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))),s('type.integer.int',X12))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X13))) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f1593,f1518])).
fof(f5745,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26))),s('type.integer.int',X28))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X30),s('type.integer.int',X31))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X27),s('type.integer.int',X26))),s('type.integer.int',X28))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X26),s('type.integer.int',X27))),s('type.integer.int',X28))) | s('type.integer.int',X29) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30)))) )),
  inference(superposition,[],[f1584,f1518])).
fof(f5740,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3)))) )),
  inference(superposition,[],[f1584,f390])).
fof(f5735,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) )),
  inference(superposition,[],[f1584,f5081])).
fof(f5730,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))))),s('type.integer.int',X24))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X23) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))),s('type.integer.int',X25)))) )),
  inference(superposition,[],[f1584,f1518])).
fof(f5728,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',X13))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X14)))) )),
  inference(superposition,[],[f1584,f5081])).
fof(f5726,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f1584,f390])).
fof(f5723,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))),s('type.integer.int',X25))),s('type.integer.int',X24))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X23) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22)))))) )),
  inference(superposition,[],[f1584,f1518])).
fof(f5721,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X14))),s('type.integer.int',X13))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1584,f5081])).
fof(f5719,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) )),
  inference(superposition,[],[f1584,f390])).
fof(f5710,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int',X25))),s('type.integer.int',X23))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))),s('type.integer.int',X23))) | s('type.integer.int',X23) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X25),s('type.integer.int',X24))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22)))) )),
  inference(superposition,[],[f1584,f1518])).
fof(f5708,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X14))),s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int',X13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1584,f5081])).
fof(f5706,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3)))) )),
  inference(superposition,[],[f1584,f390])).
fof(f5676,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))),s('type.integer.int',X23))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))),s('type.integer.int',X23)))))) )),
  inference(superposition,[],[f1545,f1518])).
fof(f5674,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X12))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X12)))))) )),
  inference(superposition,[],[f1545,f5081])).
fof(f5672,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2)))))) )),
  inference(superposition,[],[f1545,f390])).
fof(f5669,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))),s('type.integer.int',X23))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22)))))))) )),
  inference(superposition,[],[f1545,f1518])).
fof(f5667,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X12))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1545,f5081])).
fof(f5665,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) )),
  inference(superposition,[],[f1545,f390])).
fof(f5652,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK17(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f5651])).
fof(f5651,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f5642,f314])).
fof(f5642,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1566,f769])).
fof(f5612,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int',X10) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))),s('type.integer.int',X11)))) )),
  inference(superposition,[],[f1520,f5081])).
fof(f5607,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int',X23) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X25))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))),s('type.integer.int',X24))),s('type.integer.int',X25)))) )),
  inference(superposition,[],[f1520,f1518])).
fof(f5605,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X13))),s('type.integer.int',X14)))) )),
  inference(superposition,[],[f1520,f5081])).
fof(f5603,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X4))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X3))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f1520,f390])).
fof(f5600,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int',X23) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))),s('type.integer.int',X24))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X25))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))))),s('type.integer.int',X25)))) )),
  inference(superposition,[],[f1520,f1518])).
fof(f5598,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X13))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int',X14)))) )),
  inference(superposition,[],[f1520,f5081])).
fof(f5596,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X4))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f1520,f390])).
fof(f5593,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X23),s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X21),s('type.integer.int',X20))),s('type.integer.int',X22))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int',X23))),s('type.integer.int',X25))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X21))),s('type.integer.int',X22))),s('type.integer.int',X25)))) )),
  inference(superposition,[],[f1520,f1518])).
fof(f5591,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X12))),s('type.integer.int',X14))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X14)))) )),
  inference(superposition,[],[f1520,f5081])).
fof(f5589,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X4))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))),s('type.integer.int',X4)))) )),
  inference(superposition,[],[f1520,f390])).
fof(f5556,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int',X2)) )),
  inference(superposition,[],[f1587,f390])).
fof(f5553,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',X16))),s('type.integer.int',X18))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',X16))),s('type.integer.int',X18))),s('type.integer.int',X19))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X17))),s('type.integer.int',X18))) | s('type.integer.int',X19) = s('type.integer.int',X20)) )),
  inference(superposition,[],[f1587,f1518])).
fof(f5549,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X2) = s('type.integer.int',X3)) )),
  inference(superposition,[],[f1587,f390])).
fof(f5546,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',X16))),s('type.integer.int',X18))),s('type.integer.int',X19))) | s('type.integer.int',X19) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X20) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X17))),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f1587,f1518])).
fof(f5544,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1587,f5081])).
fof(f5542,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f1587,f390])).
fof(f5539,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X13))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int',X12))),s('type.integer.int',X14))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X13))) | s('type.integer.int',X14) = s('type.integer.int',X15)) )),
  inference(superposition,[],[f1587,f1518])).
fof(f5537,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X6))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1587,f5081])).
fof(f5532,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',X19))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',X16))),s('type.integer.int',X18))))) | s('type.integer.int',X19) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X20) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X16),s('type.integer.int',X17))),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f1587,f1518])).
fof(f5530,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int',X10))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1587,f5081])).
fof(f5528,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f1587,f390])).
fof(f5464,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f5459,f306])).
fof(f5459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1499,f769])).
fof(f5446,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f5441,f306])).
fof(f5441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool)) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1487,f769])).
fof(f4653,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1547,f439])).
fof(f4684,plain,(
  ( ! [X12,X10,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f1518,f1547])).
fof(f4690,plain,(
  ( ! [X30,X28,X31,X29] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X29),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X30))) | s('type.integer.int',X30) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X31) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f437,f1547])).
fof(f4694,plain,(
  ( ! [X47,X45,X46,X44] : (s('type.integer.int',X46) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X45),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X46),s('type.integer.int',X47))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X44),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X47)))) )),
  inference(superposition,[],[f439,f1547])).
fof(f545,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f404])).
fof(f547,plain,(
  ( ! [X0] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(unit_resulting_resolution,[],[f303,f405])).
fof(f549,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f460])).
fof(f553,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f464])).
fof(f557,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f468])).
fof(f563,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f483])).
fof(f567,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f303,f500])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f303,f501])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f303,f518])).
fof(f624,plain,(
  ( ! [X0] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f623,f306])).
fof(f623,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f420,f591])).
fof(f627,plain,(
  sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f542,f624])).
fof(f632,plain,(
  ( ! [X0] : (sP1(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f631,f306])).
fof(f631,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(superposition,[],[f423,f591])).
fof(f646,plain,(
  sP1('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f542,f632])).
fof(f668,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(subsumption_resolution,[],[f667,f306])).
fof(f667,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(superposition,[],[f504,f591])).
fof(f677,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f676,f306])).
fof(f676,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(superposition,[],[f505,f591])).
fof(f681,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f668])).
fof(f683,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(subsumption_resolution,[],[f682,f306])).
fof(f682,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(superposition,[],[f668,f591])).
fof(f691,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f677])).
fof(f693,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f692,f306])).
fof(f692,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f677,f591])).
fof(f698,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f697,f306])).
fof(f697,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(superposition,[],[f521,f591])).
fof(f711,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f710,f306])).
fof(f710,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(superposition,[],[f522,f591])).
fof(f715,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f698])).
fof(f726,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(unit_resulting_resolution,[],[f573,f306,f526])).
fof(f728,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f727,f306])).
fof(f727,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(superposition,[],[f526,f591])).
fof(f744,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f567,f306,f527])).
fof(f746,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(unit_resulting_resolution,[],[f569,f306,f527])).
fof(f748,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f747,f306])).
fof(f747,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) )),
  inference(superposition,[],[f527,f591])).
fof(f759,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f711])).
fof(f761,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f760,f306])).
fof(f760,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f711,f591])).
fof(f862,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f306,f591,f472])).
fof(f863,plain,(
  sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f306,f591,f472])).
fof(f884,plain,(
  ~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f542,f591,f474])).
fof(f971,plain,(
  sP7('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f563,f306,f306,f478])).
fof(f1024,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(subsumption_resolution,[],[f1018,f306])).
fof(f1018,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(superposition,[],[f491,f591])).
fof(f1056,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP3('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f1051,f306])).
fof(f1051,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP3('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f947,f591])).
fof(f1075,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP6('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(subsumption_resolution,[],[f1070,f306])).
fof(f1070,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP6('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) )),
  inference(superposition,[],[f989,f591])).
fof(f1135,plain,(
  ( ! [X0,X1] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1129,f306])).
fof(f1129,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1,X0)) )),
  inference(superposition,[],[f506,f591])).
fof(f1143,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP7(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(subsumption_resolution,[],[f1139,f306])).
fof(f1139,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~sP7(X0,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(superposition,[],[f1024,f591])).
fof(f1365,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f560,f553,f548,f306,f456])).
fof(f1367,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f560,f553,f549,f306,f456])).
fof(f1414,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f557,f553,f550,f542,f457])).
fof(f1421,plain,(
  ( ! [X0] : (sP5('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f557,f553,f575,f542,f457])).
fof(f1425,plain,(
  sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f561,f555,f551,f542,f457])).
fof(f1430,plain,(
  ( ! [X0] : (sP5('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f561,f555,f576,f542,f457])).
fof(f1452,plain,(
  ( ! [X0] : (sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(unit_resulting_resolution,[],[f561,f578,f551,f542,f458])).
fof(f1474,plain,(
  ( ! [X0] : (sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f559,f577,f575,f542,f458])).
fof(f1545,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3)))))) )),
  inference(superposition,[],[f438,f390])).
fof(f1593,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s('type.integer.int',X4) = s('type.integer.int',X5)) )),
  inference(superposition,[],[f437,f390])).
fof(f4455,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X10))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X10)))))) )),
  inference(superposition,[],[f438,f1518])).
fof(f4476,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))),s('type.integer.int',X10))),s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f437,f1518])).
fof(f4491,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))),s('type.integer.int',X10))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X8))),s('type.integer.int',X10))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X9))),s('type.integer.int',X10))) | s('type.integer.int',X11) = s('type.integer.int',X12)) )),
  inference(superposition,[],[f437,f1518])).
fof(f3253,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s('type.integer.int',X67) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X65),s('type.integer.int',X64))),s('type.integer.int',X66))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X67),s('type.integer.int',X68))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X64),s('type.integer.int',X65))),s('type.integer.int',X66))),s('type.integer.int',X68)))) )),
  inference(superposition,[],[f439,f1518])).
fof(f3217,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X8)))) )),
  inference(superposition,[],[f1518,f390])).
fof(f3215,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))))),s('type.integer.int',X9))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X8))),s('type.integer.int',X9)))) )),
  inference(superposition,[],[f1518,f1518])).
fof(f3214,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))),s('type.integer.int',X3)))) )),
  inference(unit_resulting_resolution,[],[f1518,f439])).
fof(f1487,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f386,f314])).
fof(f1499,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f387,f314])).
fof(f1565,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,X2,bool))) )),
  inference(superposition,[],[f436,f314])).
fof(f1566,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f436,f314])).
fof(f2302,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') != s(bool,f(s(X117,sK17(X118,X119,bool,X117)))) | s(fun(X117,bool),X118) = s(fun(X117,bool),X119) | s(bool,X120) = s(bool,g(s(X117,sK17(X118,X119,bool,X117)))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f436,f769])).
fof(f2190,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,X124) != s(bool,f(s(X121,sK17(X122,X123,bool,X121)))) | s(fun(X121,bool),X122) = s(fun(X121,bool),X123) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,g(s(X121,sK17(X122,X123,bool,X121))))) )),
  inference(superposition,[],[f436,f769])).
fof(f2582,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | s(bool,X113) = s(bool,happ(s(fun(X110,bool),X111),s(X110,sK16(X111,X110)))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f2301,f306])).
fof(f2301,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X110,bool),X111),s(X110,X112)))) | s(bool,X113) = s(bool,happ(s(fun(X110,bool),X111),s(X110,sK16(X111,X110)))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f387,f769])).
fof(f387,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X13] : ! [X14] : (~! [X7] : p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X7)))) <=> ? [X7] : ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X7)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.NOT_FORALL_THM')).
fof(f2581,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,X108)))) | s(bool,X109) = s(bool,happ(s(fun(X106,bool),X107),s(X106,sK15(X107,X106)))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f2300,f306])).
fof(f2300,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X106,bool),X107),s(X106,X108)))) | s(bool,X109) = s(bool,happ(s(fun(X106,bool),X107),s(X106,sK15(X107,X106)))) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f386,f769])).
fof(f386,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f2191,plain,(
  ( ! [X127,X125,X126] : (~p(s(bool,happ(s(fun(bool,bool),X125),s(bool,X126)))) | p(s(bool,happ(s(fun(bool,bool),X125),s(bool,X127)))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,sK15(X125,bool))) )),
  inference(superposition,[],[f386,f769])).
fof(f2192,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,happ(s(fun(bool,bool),X128),s(bool,X129)))) | p(s(bool,happ(s(fun(bool,bool),X128),s(bool,X130)))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,sK16(X128,bool))) )),
  inference(superposition,[],[f387,f769])).
fof(f2303,plain,(
  ( ! [X123,X121,X122] : (~p(s(bool,happ(s(fun(bool,bool),X121),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X121),s(bool,X122)))) | s(bool,X123) = s(bool,sK15(X121,bool)) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f386,f769])).
fof(f2304,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,happ(s(fun(bool,bool),X124),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X124),s(bool,X125)))) | s(bool,X126) = s(bool,sK16(X124,bool)) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f387,f769])).
fof(f2305,plain,(
  ( ! [X127,X130,X128,X129] : (s(X129,f(s(bool,'T_0'))) != s(X129,g(s(bool,'T_0'))) | s(fun(bool,X129),X127) = s(fun(bool,X129),X128) | s(bool,X130) = s(bool,sK17(X127,X128,X129,bool)) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f436,f769])).
fof(f2580,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2291,f306])).
fof(f2291,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f527,f769])).
fof(f2579,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2290,f306])).
fof(f2290,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f526,f769])).
fof(f2578,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2289,f306])).
fof(f2289,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f523,f769])).
fof(f2577,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2288,f306])).
fof(f2288,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f522,f769])).
fof(f2576,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2287,f306])).
fof(f2287,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f521,f769])).
fof(f2575,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2286,f306])).
fof(f2286,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f516,f769])).
fof(f2574,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2285,f306])).
fof(f2285,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f506,f769])).
fof(f2573,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2284,f306])).
fof(f2284,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f505,f769])).
fof(f2572,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2283,f306])).
fof(f2283,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f504,f769])).
fof(f2571,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2282,f306])).
fof(f2282,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f499,f769])).
fof(f2570,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2281,f306])).
fof(f2281,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f491,f769])).
fof(f2569,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2280,f306])).
fof(f2280,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f482,f769])).
fof(f2568,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2279,f306])).
fof(f2279,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f463,f769])).
fof(f2567,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2278,f306])).
fof(f2278,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f459,f769])).
fof(f2566,plain,(
  ( ! [X33,X31,X34,X32] : (sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2277,f306])).
fof(f2277,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f456,f769])).
fof(f2565,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2276,f306])).
fof(f2276,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f423,f769])).
fof(f2564,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2275,f306])).
fof(f2275,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f420,f769])).
fof(f2563,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2274,f306])).
fof(f2274,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f405,f769])).
fof(f2562,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2273,f306])).
fof(f2273,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f404,f769])).
fof(f2561,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2272,f306])).
fof(f2272,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f394,f769])).
fof(f2560,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2268,f306])).
fof(f2268,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f330,f769])).
fof(f2236,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f326,f769])).
fof(f2193,plain,(
  ( ! [X134,X132,X133,X131] : (s(X133,f(s(bool,X134))) != s(X133,g(s(bool,X134))) | s(fun(bool,X133),X131) = s(fun(bool,X133),X132) | s(bool,'T_0') = s(bool,X134) | s(bool,'T_0') = s(bool,sK17(X131,X132,X133,bool))) )),
  inference(superposition,[],[f436,f769])).
fof(f2130,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f769,f769])).
fof(f2092,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f2088,f542])).
fof(f2088,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f831,f610])).
fof(f1520,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X5))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))),s('type.integer.int',X5)))) )),
  inference(superposition,[],[f439,f390])).
fof(f1518,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2)))) )),
  inference(unit_resulting_resolution,[],[f390,f439])).
fof(f1471,plain,(
  ( ! [X0] : (~sP2('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f554,f1175,f550,f542,f458])).
fof(f1413,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f556,f552,f550,f542,f457])).
fof(f1428,plain,(
  ( ! [X0] : (~sP2('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f577,f1177,f551,f542,f457])).
fof(f1364,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f560,f552,f548,f306,f456])).
fof(f1370,plain,(
  ( ! [X0] : (sP5('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f560,f555,f576,f306,f456])).
fof(f1371,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f560,f578,f576,f306,f456])).
fof(f1376,plain,(
  ( ! [X0] : (~sP4('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f556,f552,f896,f306,f456])).
fof(f1244,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1232,f511])).
fof(f1232,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f306,f542,f306,f516])).
fof(f1218,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f574,f306,f542,f515])).
fof(f1176,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f306,f610,f472])).
fof(f1177,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f542,f610,f474])).
fof(f1105,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1094,f494])).
fof(f1094,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f542,f306,f499])).
fof(f1078,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f570,f542,f542,f498])).
fof(f993,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f982,f475])).
fof(f982,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f306,f306,f482])).
fof(f970,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f562,f306,f306,f478])).
fof(f966,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f957,f452])).
fof(f957,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f542,f542,f467])).
fof(f951,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f940,f453])).
fof(f940,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f306,f306,f463])).
fof(f916,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f306,f306,f459])).
fof(f896,plain,(
  ( ! [X0] : (~sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f587,f473])).
fof(f897,plain,(
  ( ! [X0] : (sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f542,f587,f471])).
fof(f899,plain,(
  ( ! [X0] : (~sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f306,f587,f473])).
fof(f900,plain,(
  ( ! [X0] : (sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f542,f587,f471])).
fof(f873,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f306,f586,f473])).
fof(f875,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f306,f586,f473])).
fof(f854,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f542,f586,f471])).
fof(f855,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f586,f471])).
fof(f831,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f827,f542])).
fof(f827,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f391,f314])).
fof(f769,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f314,f314])).
fof(f758,plain,(
  ( ! [X0,X1] : (sP7('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f580,f542,f528])).
fof(f709,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f306,f522])).
fof(f696,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f306,f521])).
fof(f690,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f542,f510])).
fof(f684,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f542,f509])).
fof(f675,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f306,f505])).
fof(f666,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f306,f504])).
fof(f664,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f542,f493])).
fof(f663,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f542,f492])).
fof(f656,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f542,f542,f487])).
fof(f653,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f542,f486])).
fof(f650,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f542,f477])).
fof(f647,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f306,f542,f476])).
fof(f641,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f542,f543])).
fof(f630,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f542,f306,f423])).
fof(f622,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f306,f420])).
fof(f618,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f306,f306,f419])).
fof(f616,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f542,f542,f418])).
fof(f586,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f542,f326])).
fof(f587,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(unit_resulting_resolution,[],[f383,f326])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f383,f462])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (sP6('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f383,f485])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f542,f519])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f306,f518])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f542,f517])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f542,f502])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f306,f501])).
fof(f566,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f500])).
fof(f565,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f542,f485])).
fof(f564,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f542,f484])).
fof(f562,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f483])).
fof(f558,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f306,f469])).
fof(f556,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f468])).
fof(f555,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f542,f466])).
fof(f552,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f464])).
fof(f548,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f460])).
fof(f546,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f306,f405])).
fof(f544,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f306,f404])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f302])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f300])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f297])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f297])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f297])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f297])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f144,f194,f193])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X19)) | p(s(bool,X18)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | p(s(bool,X18)) | ~p(s(bool,X20)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.sat.sat_25_prove_ltb')).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f294])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f294])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f294])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f292])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f292])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f289])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f289])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f289])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f142,f191,f190])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X19)) & p(s(bool,X18)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X18)) | ~p(s(bool,X20)))))),
  file('thm.int_arith.eq_justify_multiplication','thm.sat.sat_23_prove_ltb')).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f286])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f284])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f282])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f282])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f282])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f279])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f279])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f278])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f276])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f273])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f270])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f270])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f270])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f263])).
fof(f436,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.int_arith.eq_justify_multiplication',aHL_EXT)).
fof(f418,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f419,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f420,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f423,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f543,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f403,f542])).
fof(f403,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f404,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f405,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f391,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.int_arith.eq_justify_multiplication','thm.bool.IMP_ANTISYM_AX')).
fof(f394,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f164])).
fof(f326,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f7440,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(subsumption_resolution,[],[f7301,f579])).
fof(f7301,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f989])).
fof(f7439,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(subsumption_resolution,[],[f7299,f577])).
fof(f7299,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f947])).
fof(f7438,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(subsumption_resolution,[],[f7298,f575])).
fof(f7298,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl110),
  inference(backward_demodulation,[],[f7259,f923])).
fof(f7437,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7297,f900])).
fof(f7297,plain,(
  ~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f884])).
fof(f7436,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7296,f897])).
fof(f7296,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f882])).
fof(f7435,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7295,f899])).
fof(f7295,plain,(
  sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f863])).
fof(f7434,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7294,f896])).
fof(f7294,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | $spl110),
  inference(backward_demodulation,[],[f7259,f862])).
fof(f7429,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f7273,f587])).
fof(f7273,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | $spl110),
  inference(backward_demodulation,[],[f7259,f591])).
fof(f7259,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl110),
  inference(cnf_transformation,[],[f7259_D])).
fof(f7259_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4075,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK14) = s('type.integer.int',X10)) ) | $spl56),
  inference(cnf_transformation,[],[f4075_D])).
fof(f4075_D,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK14) = s('type.integer.int',X10)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5541,plain,(
  ( ! [X20] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X20),s('type.integer.int',sK12))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',sK14) = s('type.integer.int',X20)) ) | $spl2),
  inference(superposition,[],[f1587,f540])).
fof(f7256,plain,(
  $false | ($spl1 | $spl56)),
  inference(subsumption_resolution,[],[f7239,f390])).
fof(f7239,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl56)),
  inference(unit_resulting_resolution,[],[f539,f4075])).
fof(f7255,plain,(
  $false | ($spl1 | $spl56)),
  inference(subsumption_resolution,[],[f7241,f539])).
fof(f7241,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | $spl56),
  inference(unit_resulting_resolution,[],[f390,f4075])).
fof(f7254,plain,(
  $false | ($spl1 | $spl56)),
  inference(subsumption_resolution,[],[f7243,f539])).
fof(f7243,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | $spl56),
  inference(unit_resulting_resolution,[],[f390,f4075])).
fof(f7244,plain,(
  $false | ($spl1 | $spl56)),
  inference(unit_resulting_resolution,[],[f539,f390,f4075])).
fof(f7242,plain,(
  $false | ($spl1 | $spl56)),
  inference(unit_resulting_resolution,[],[f539,f390,f4075])).
fof(f6770,plain,(
  ( ! [X1] : (~sP7(X1,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | p(s(bool,X1))) ) | $spl108),
  inference(cnf_transformation,[],[f6770_D])).
fof(f6770_D,plain,(
  ( ! [X1] : (~sP7(X1,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | p(s(bool,X1))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2510,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f2510_D])).
fof(f2510_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6772,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7(X1,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(condensation,[],[f6771])).
fof(f6771,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP7(X18,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f6761,f831])).
fof(f6761,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP7(X18,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f1024,f769])).
fof(f6769,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7(X1,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) )),
  inference(condensation,[],[f6768])).
fof(f6768,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP7(X15,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f6760,f831])).
fof(f6760,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP7(X15,'T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1024,f769])).
fof(f6721,plain,(
  ( ! [X1] : (~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) | $spl106),
  inference(cnf_transformation,[],[f6721_D])).
fof(f6721_D,plain,(
  ( ! [X1] : (~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6723,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6722])).
fof(f6722,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f6711,f831])).
fof(f6711,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f989,f769])).
fof(f6720,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6719])).
fof(f6719,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f6710,f831])).
fof(f6710,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f989,f769])).
fof(f6674,plain,(
  ( ! [X1] : (~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) | $spl104),
  inference(cnf_transformation,[],[f6674_D])).
fof(f6674_D,plain,(
  ( ! [X1] : (~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f6676,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6675])).
fof(f6675,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f6664,f831])).
fof(f6664,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f947,f769])).
fof(f6673,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6672])).
fof(f6672,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f6663,f831])).
fof(f6663,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f947,f769])).
fof(f6628,plain,(
  ( ! [X1] : (~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) | $spl102),
  inference(cnf_transformation,[],[f6628_D])).
fof(f6628_D,plain,(
  ( ! [X1] : (~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1) | p(s(bool,X1))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6630,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6629])).
fof(f6629,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f6618,f831])).
fof(f6618,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f923,f769])).
fof(f6627,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X1)) )),
  inference(condensation,[],[f6626])).
fof(f6626,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f6617,f831])).
fof(f6617,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f923,f769])).
fof(f5818,plain,(
  $false | ($spl66 | $spl72 | $spl90)),
  inference(subsumption_resolution,[],[f5768,f5817])).
fof(f5817,plain,(
  ( ! [X10,X11] : (s('type.integer.int',X10) = s('type.integer.int',X11)) ) | ($spl66 | $spl90)),
  inference(subsumption_resolution,[],[f5816,f5760])).
fof(f5816,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X10) = s('type.integer.int',X11)) ) | ($spl66 | $spl90)),
  inference(forward_demodulation,[],[f5773,f5762])).
fof(f5773,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int',X11)) ) | ($spl66 | $spl90)),
  inference(backward_demodulation,[],[f5760,f5569])).
fof(f5793,plain,(
  $false | ($spl64 | $spl66 | $spl72)),
  inference(subsumption_resolution,[],[f5768,f5792])).
fof(f5792,plain,(
  ( ! [X39,X38] : (s('type.integer.int',X38) = s('type.integer.int',X39)) ) | ($spl64 | $spl66)),
  inference(subsumption_resolution,[],[f5791,f5760])).
fof(f5791,plain,(
  ( ! [X39,X38] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) ) | ($spl64 | $spl66)),
  inference(forward_demodulation,[],[f5766,f5760])).
fof(f5766,plain,(
  ( ! [X39,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) ) | ($spl64 | $spl66)),
  inference(backward_demodulation,[],[f5760,f4777])).
fof(f5688,plain,(
  ( ! [X10,X8,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))))) ) | $spl100),
  inference(cnf_transformation,[],[f5688_D])).
fof(f5688_D,plain,(
  ( ! [X10,X8,X11] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4779,plain,(
  ( ! [X41] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl68),
  inference(cnf_transformation,[],[f4779_D])).
fof(f4779_D,plain,(
  ( ! [X41] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5682,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X14),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))))) )),
  inference(superposition,[],[f1545,f5081])).
fof(f5681,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))))) )),
  inference(superposition,[],[f1545,f5081])).
fof(f5649,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK17(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X6,X7,bool,bool))))) ) | $spl96),
  inference(cnf_transformation,[],[f5649_D])).
fof(f5649_D,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK17(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X6,X7,bool,bool))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5492,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl89),
  inference(cnf_transformation,[],[f5492_D])).
fof(f5492_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f5639,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X12,X13,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X12,X13,bool,bool))) )),
  inference(superposition,[],[f1566,f314])).
fof(f5650,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl98),
  inference(cnf_transformation,[],[f5650_D])).
fof(f5650_D,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5638,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X11))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X9,X10,bool,bool)))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,sK17(X9,X10,bool,bool))) )),
  inference(superposition,[],[f1566,f769])).
fof(f5637,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X6,X7,bool,bool)))) | s(bool,'T_0') = s(bool,sK17(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1566,f769])).
fof(f5646,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) | $spl92),
  inference(cnf_transformation,[],[f5646_D])).
fof(f5646_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f5648,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl95),
  inference(cnf_transformation,[],[f5648_D])).
fof(f5648_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f5636,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK17(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1566,f769])).
fof(f5635,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1566,f769])).
fof(f4776,plain,(
  ( ! [X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X36),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl62),
  inference(cnf_transformation,[],[f4776_D])).
fof(f4776_D,plain,(
  ( ! [X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X36),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5569,plain,(
  ( ! [X10,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int',X11)) ) | $spl90),
  inference(cnf_transformation,[],[f5569_D])).
fof(f5569_D,plain,(
  ( ! [X10,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int',X11)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5552,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X14))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X14) = s('type.integer.int',X15)) )),
  inference(superposition,[],[f1587,f5081])).
fof(f5551,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X10) = s('type.integer.int',X11)) )),
  inference(superposition,[],[f1587,f5081])).
fof(f5489,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl84),
  inference(cnf_transformation,[],[f5489_D])).
fof(f5489_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5485,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool))) ) | $spl80),
  inference(cnf_transformation,[],[f5485_D])).
fof(f5485_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5483,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK17(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1565,f314])).
fof(f5490,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl86),
  inference(cnf_transformation,[],[f5490_D])).
fof(f5490_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5482,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK17(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1565,f769])).
fof(f5481,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK17(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1565,f769])).
fof(f5484,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f5484_D])).
fof(f5484_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5487,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl83),
  inference(cnf_transformation,[],[f5487_D])).
fof(f5487_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f5480,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK17(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1565,f769])).
fof(f5479,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1565,f769])).
fof(f5147,plain,(
  $false | ($spl1 | $spl2 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f5146,f540])).
fof(f5146,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f5145,f390])).
fof(f5145,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f5096,f390])).
fof(f5096,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl1 | $spl68 | $spl70)),
  inference(unit_resulting_resolution,[],[f539,f5034,f437])).
fof(f5034,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl68 | $spl70)),
  inference(forward_demodulation,[],[f4779,f4780])).
fof(f5142,plain,(
  $false | ($spl1 | $spl2 | $spl68 | $spl70)),
  inference(subsumption_resolution,[],[f5141,f540])).
fof(f5141,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f5140,f390])).
fof(f5140,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl68 | $spl70)),
  inference(forward_demodulation,[],[f5100,f390])).
fof(f5100,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl1 | $spl68 | $spl70)),
  inference(unit_resulting_resolution,[],[f539,f5034,f437])).
fof(f5033,plain,(
  $false | ($spl53 | $spl64 | $spl66 | $spl70)),
  inference(subsumption_resolution,[],[f5025,f4918])).
fof(f4918,plain,(
  ( ! [X39,X38] : (s('type.integer.int',X38) = s('type.integer.int',X39)) ) | ($spl64 | $spl66 | $spl70)),
  inference(subsumption_resolution,[],[f4917,f4789])).
fof(f4789,plain,(
  ( ! [X43,X42] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int',sK12)))) ) | ($spl66 | $spl70)),
  inference(forward_demodulation,[],[f4778,f4780])).
fof(f4917,plain,(
  ( ! [X39,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int',sK12))) | s('type.integer.int',X38) = s('type.integer.int',X39)) ) | ($spl64 | $spl70)),
  inference(forward_demodulation,[],[f4777,f4780])).
fof(f5025,plain,(
  ( ! [X128] : (s('type.integer.int',sK12) != s('type.integer.int',X128)) ) | ($spl53 | $spl64 | $spl66 | $spl70)),
  inference(superposition,[],[f3563,f4918])).
fof(f5032,plain,(
  $false | ($spl1 | $spl64 | $spl66 | $spl70)),
  inference(subsumption_resolution,[],[f5024,f4918])).
fof(f5024,plain,(
  ( ! [X127] : (s('type.integer.int',sK13) != s('type.integer.int',X127)) ) | ($spl1 | $spl64 | $spl66 | $spl70)),
  inference(superposition,[],[f539,f4918])).
fof(f5031,plain,(
  $false | ($spl53 | $spl64 | $spl66 | $spl70)),
  inference(subsumption_resolution,[],[f4978,f4918])).
fof(f4978,plain,(
  ( ! [X128] : (s('type.integer.int',sK12) != s('type.integer.int',X128)) ) | ($spl53 | $spl64 | $spl66 | $spl70)),
  inference(superposition,[],[f3563,f4918])).
fof(f5030,plain,(
  $false | ($spl1 | $spl64 | $spl66 | $spl70)),
  inference(subsumption_resolution,[],[f4977,f4918])).
fof(f4977,plain,(
  ( ! [X127] : (s('type.integer.int',sK13) != s('type.integer.int',X127)) ) | ($spl1 | $spl64 | $spl66 | $spl70)),
  inference(superposition,[],[f539,f4918])).
fof(f4919,plain,(
  $false | ($spl1 | $spl64 | $spl66 | $spl70)),
  inference(unit_resulting_resolution,[],[f539,f4918])).
fof(f4905,plain,(
  ( ! [X10] : (s('type.integer.int',sK14) = s('type.integer.int',X10)) ) | $spl74),
  inference(cnf_transformation,[],[f4905_D])).
fof(f4905_D,plain,(
  ( ! [X10] : (s('type.integer.int',sK14) = s('type.integer.int',X10)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4907,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl77),
  inference(cnf_transformation,[],[f4907_D])).
fof(f4907_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f4904,plain,(
  ( ! [X10] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK14) = s('type.integer.int',X10)) ) | ($spl56 | $spl62 | $spl70)),
  inference(backward_demodulation,[],[f4870,f4075])).
fof(f4870,plain,(
  ( ! [X141] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X141),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f4780])).
fof(f4790,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl62 | $spl70)),
  inference(forward_demodulation,[],[f4776,f4780])).
fof(f4902,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4860,f1175])).
fof(f4860,plain,(
  sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1469])).
fof(f4897,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4853,f578])).
fof(f4853,plain,(
  ~sP3('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1377])).
fof(f4896,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4850,f724])).
fof(f4850,plain,(
  ~sP11('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1353])).
fof(f4895,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4849,f582])).
fof(f4849,plain,(
  ~sP10('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1343])).
fof(f4894,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4846,f1079])).
fof(f4846,plain,(
  ~sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1157])).
fof(f4893,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4845,f581])).
fof(f4845,plain,(
  ~sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1144])).
fof(f4892,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4841,f756])).
fof(f4841,plain,(
  ~sP7('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1086])).
fof(f4891,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4839,f579])).
fof(f4839,plain,(
  ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1065])).
fof(f4890,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4837,f577])).
fof(f4837,plain,(
  ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1046])).
fof(f4889,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4836,f1178])).
fof(f4836,plain,(
  ~sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1039])).
fof(f4887,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4834,f575])).
fof(f4834,plain,(
  ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1025])).
fof(f4886,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(global_subsumption,[],[f305,f304,f303,f306,f314,f327,f326,f383,f390,f394,f391,f405,f404,f543,f423,f420,f419,f418,f436,f439,f458,f457,f456,f455,f454,f453,f452,f462,f461,f460,f459,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f528,f478,f477,f476,f475,f485,f484,f483,f482,f491,f488,f487,f486,f498,f497,f496,f527,f494,f493,f492,f502,f501,f500,f499,f506,f505,f504,f503,f515,f514,f513,f526,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f544,f546,f548,f550,f551,f552,f554,f555,f556,f558,f560,f562,f564,f565,f566,f568,f570,f571,f572,f574,f583,f582,f581,f580,f579,f578,f577,f576,f575,f587,f586,f616,f618,f622,f630,f641,f647,f650,f653,f656,f663,f664,f666,f675,f684,f690,f696,f709,f724,f758,f756,f769,f831,f855,f854,f875,f873,f900,f899,f897,f896,f916,f940,f951,f957,f966,f970,f982,f993,f1079,f1078,f1094,f1105,f1178,f1177,f1176,f1175,f1219,f1218,f1232,f1244,f1376,f1371,f1370,f1364,f1428,f1413,f1471,f1518,f1520,f2092,f2130,f2193,f2236,f2508,f2512,f2516,f2519,f2522,f2525,f2528,f2531,f2534,f2537,f2540,f2543,f2546,f2549,f2552,f2555,f2558,f2560,f2561,f2562,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2305,f2588,f2590,f2593,f2595,f2597,f2599,f2601,f2603,f2605,f2607,f2609,f2611,f2613,f2615,f2617,f2619,f2621,f2304,f2303,f2192,f2191,f386,f2581,f387,f2582,f2190,f2302,f1566,f1565,f1499,f1487,f610,f542,f331,f330,f3214,f3215,f3217,f3253,f1593,f1587,f1584,f1547,f1545,f1474,f1469,f1452,f1430,f1425,f1421,f1414,f1392,f1377,f1367,f1365,f1353,f1343,f1278,f1240,f1157,f1144,f1143,f1135,f1101,f1086,f1075,f1065,f1056,f1046,f1039,f1035,f1025,f1024,f989,f971,f947,f923,f884,f882,f863,f862,f761,f759,f748,f746,f744,f728,f726,f715,f711,f698,f693,f691,f683,f681,f677,f668,f646,f632,f627,f624,f591,f573,f569,f567,f563,f561,f559,f557,f553,f549,f547,f545,f438,f437,f4455,f4476,f4491,f4653,f4663,f4684,f4690,f4692,f4693,f4694,f4740,f4741,f4780,f4784,f4786,f4787,f4790,f4791,f4792,f4793,f4794,f4795,f4796,f4797,f4798,f4799,f4800,f4801,f4802,f4803,f4805,f4806,f4807,f4808,f4809,f4810,f4811,f4812,f4813,f4814,f4815,f4816,f4817,f4818,f4876,f4820,f4821,f4878,f4883,f4884,f4831,f4885,f4833])).
fof(f4833,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f1024])).
fof(f4831,plain,(
  sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','T_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f971])).
fof(f4878,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4877,f581])).
fof(f4877,plain,(
  ( ! [X0,X1] : (~sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(forward_demodulation,[],[f4822,f4790])).
fof(f4822,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f748])).
fof(f4821,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f746])).
fof(f4820,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f744])).
fof(f4876,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4875,f583])).
fof(f4875,plain,(
  ( ! [X0,X1] : (~sP10('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(forward_demodulation,[],[f4819,f4790])).
fof(f4819,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f728])).
fof(f4818,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f726])).
fof(f4817,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'F_0')) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f715])).
fof(f4816,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f711])).
fof(f4815,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f698])).
fof(f4814,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f693])).
fof(f4813,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0',X0)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f691])).
fof(f4812,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f683])).
fof(f4811,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'T_0')) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f681])).
fof(f4810,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f677])).
fof(f4809,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | ~p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f668])).
fof(f4808,plain,(
  sP1('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f646])).
fof(f4807,plain,(
  ( ! [X0] : (sP1(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f632])).
fof(f4806,plain,(
  sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f627])).
fof(f4805,plain,(
  ( ! [X0] : (sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f624])).
fof(f4803,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f573])).
fof(f4802,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f569])).
fof(f4801,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f567])).
fof(f4800,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f563])).
fof(f4799,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f561])).
fof(f4798,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f559])).
fof(f4797,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f557])).
fof(f4796,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f553])).
fof(f4795,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f549])).
fof(f4794,plain,(
  ( ! [X0] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f547])).
fof(f4793,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f545])).
fof(f4792,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f438])).
fof(f4791,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f437])).
fof(f4787,plain,(
  ( ! [X47,X45,X46] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',X47))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X46),s('type.integer.int',X47))) | s('type.integer.int',X46) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X45),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl70),
  inference(backward_demodulation,[],[f4780,f4694])).
fof(f4786,plain,(
  ( ! [X30,X28,X31] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',X30))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X31),s('type.integer.int',X30))) | s('type.integer.int',X30) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X31) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl70),
  inference(backward_demodulation,[],[f4780,f4690])).
fof(f4784,plain,(
  ( ! [X6] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) ) | $spl70),
  inference(backward_demodulation,[],[f4780,f4663])).
fof(f4885,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4832,f579])).
fof(f4832,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f989])).
fof(f4884,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4830,f577])).
fof(f4830,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f947])).
fof(f4883,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4829,f575])).
fof(f4829,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f923])).
fof(f4882,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4828,f900])).
fof(f4828,plain,(
  ~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f884])).
fof(f4881,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4827,f897])).
fof(f4827,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f882])).
fof(f4880,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4826,f899])).
fof(f4826,plain,(
  sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f863])).
fof(f4879,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4825,f896])).
fof(f4825,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f862])).
fof(f4874,plain,(
  $false | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4804,f587])).
fof(f4804,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl62 | $spl70)),
  inference(backward_demodulation,[],[f4790,f591])).
fof(f4780,plain,(
  ( ! [X141] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X141),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl70),
  inference(cnf_transformation,[],[f4780_D])).
fof(f4780_D,plain,(
  ( ! [X141] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X141),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4781,plain,(
  ( ! [X142] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X142),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14)))) ) | $spl72),
  inference(cnf_transformation,[],[f4781_D])).
fof(f4781_D,plain,(
  ( ! [X142] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X142),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4775,plain,(
  ( ! [X142,X141] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X142),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X141),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f1602,f1547])).
fof(f4778,plain,(
  ( ! [X43,X42,X40] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl66),
  inference(cnf_transformation,[],[f4778_D])).
fof(f4778_D,plain,(
  ( ! [X43,X42,X40] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4741,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f438,f1547])).
fof(f4777,plain,(
  ( ! [X39,X37,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) ) | $spl64),
  inference(cnf_transformation,[],[f4777_D])).
fof(f4777_D,plain,(
  ( ! [X39,X37,X38] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4740,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X36),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) )),
  inference(superposition,[],[f437,f1547])).
fof(f4727,plain,(
  ( ! [X142,X141] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X142),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK14))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X141),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f1602,f1547])).
fof(f4693,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X42),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X40),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f438,f1547])).
fof(f4692,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X37),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X36),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X38) = s('type.integer.int',X39)) )),
  inference(superposition,[],[f437,f1547])).
fof(f4447,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl53)),
  inference(subsumption_resolution,[],[f4446,f540])).
fof(f4446,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl4 | $spl53)),
  inference(forward_demodulation,[],[f4445,f390])).
fof(f4445,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl4 | $spl53)),
  inference(forward_demodulation,[],[f4404,f390])).
fof(f4404,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl1 | $spl4 | $spl53)),
  inference(unit_resulting_resolution,[],[f539,f3563,f1605])).
fof(f1605,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int',X2) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f437])).
fof(f4444,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl53)),
  inference(subsumption_resolution,[],[f4443,f540])).
fof(f4443,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | ($spl1 | $spl4 | $spl53)),
  inference(forward_demodulation,[],[f4442,f390])).
fof(f4442,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ($spl1 | $spl4 | $spl53)),
  inference(forward_demodulation,[],[f4406,f390])).
fof(f4406,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ($spl1 | $spl4 | $spl53)),
  inference(unit_resulting_resolution,[],[f539,f3563,f1605])).
fof(f4334,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4267,f1175])).
fof(f4267,plain,(
  sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1674])).
fof(f1674,plain,(
  sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1469])).
fof(f4333,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4264,f1177])).
fof(f4264,plain,(
  sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1671])).
fof(f1671,plain,(
  sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f1425])).
fof(f4332,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4261,f578])).
fof(f4261,plain,(
  ~sP3('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1667])).
fof(f1667,plain,(
  ~sP3('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0','F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1377])).
fof(f4331,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4258,f724])).
fof(f4258,plain,(
  ~sP11('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1664])).
fof(f1664,plain,(
  ~sP11('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f1353])).
fof(f4330,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4257,f582])).
fof(f4257,plain,(
  ~sP10('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1663])).
fof(f1663,plain,(
  ~sP10('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f1343])).
fof(f4329,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4254,f1079])).
fof(f4254,plain,(
  ~sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1660])).
fof(f1660,plain,(
  ~sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0','F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1157])).
fof(f4328,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4253,f581])).
fof(f4253,plain,(
  ~sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0','F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1659])).
fof(f1659,plain,(
  ~sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0','F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1144])).
fof(f4327,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4249,f756])).
fof(f4249,plain,(
  ~sP7('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1655])).
fof(f1655,plain,(
  ~sP7('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f1086])).
fof(f4326,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4247,f579])).
fof(f4247,plain,(
  ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1653])).
fof(f1653,plain,(
  ~sP6('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1065])).
fof(f4325,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4245,f577])).
fof(f4245,plain,(
  ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1651])).
fof(f1651,plain,(
  ~sP3('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1046])).
fof(f4324,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4244,f1178])).
fof(f4244,plain,(
  ~sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1650])).
fof(f1650,plain,(
  ~sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f1039])).
fof(f4322,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4242,f575])).
fof(f4242,plain,(
  ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1648])).
fof(f1648,plain,(
  ~sP4('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1025])).
fof(f4321,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(global_subsumption,[],[f305,f304,f303,f306,f314,f327,f326,f383,f390,f394,f391,f405,f404,f543,f423,f420,f419,f418,f436,f439,f438,f437,f458,f457,f456,f455,f454,f453,f452,f462,f461,f460,f459,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f528,f478,f477,f476,f475,f485,f484,f483,f482,f491,f488,f487,f486,f498,f497,f496,f527,f494,f493,f492,f502,f501,f500,f499,f506,f505,f504,f503,f515,f514,f513,f526,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f545,f544,f547,f546,f549,f548,f550,f551,f553,f552,f554,f555,f557,f556,f559,f558,f561,f560,f563,f562,f564,f565,f567,f566,f569,f568,f570,f571,f573,f572,f574,f583,f582,f581,f580,f579,f578,f577,f576,f575,f587,f586,f591,f616,f618,f622,f624,f627,f630,f632,f641,f646,f647,f650,f653,f656,f663,f664,f666,f668,f675,f677,f681,f683,f684,f690,f691,f693,f696,f698,f709,f711,f715,f726,f724,f728,f746,f744,f748,f758,f756,f759,f761,f769,f831,f855,f854,f863,f862,f875,f873,f884,f882,f900,f899,f897,f896,f916,f923,f940,f947,f951,f957,f966,f971,f970,f982,f989,f993,f1024,f1025,f1035,f1039,f1046,f1056,f1065,f1075,f1079,f1078,f1086,f1094,f1101,f1105,f1135,f1143,f1144,f1157,f1178,f1177,f1176,f1175,f1219,f1218,f1232,f1240,f1244,f1278,f1343,f1353,f1377,f1376,f1371,f1370,f1367,f1365,f1364,f1392,f1430,f1428,f1425,f1421,f1414,f1413,f1474,f1471,f1469,f1452,f1518,f1520,f1545,f1547,f1584,f1587,f1593,f1937,f1938,f1978,f1979,f2092,f2130,f2193,f2236,f2508,f2512,f2516,f2519,f2522,f2525,f2528,f2531,f2534,f2537,f2540,f2543,f2546,f2549,f2552,f2555,f2558,f2560,f2561,f2562,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2305,f2588,f2590,f2593,f2595,f2597,f2599,f2601,f2603,f2605,f2607,f2609,f2611,f2613,f2615,f2617,f2619,f2621,f2304,f2303,f2192,f2191,f386,f2581,f387,f2582,f2190,f2302,f1566,f1565,f1499,f1487,f610,f542,f331,f330,f3214,f3215,f3217,f3253,f3298,f3297,f3296,f3295,f3294,f3293,f3292,f3291,f3290,f3289,f3288,f3287,f3286,f3285,f3283,f3246,f3173,f3170,f3133,f3128,f3125,f3121,f3097,f3094,f2060,f2057,f2054,f2049,f1935,f1905,f1885,f1875,f1684,f1682,f1681,f1680,f1679,f1678,f1677,f1676,f1675,f1674,f1673,f1672,f1671,f1670,f1669,f1667,f1666,f1665,f1664,f1663,f1662,f1661,f1660,f1659,f1658,f1657,f1656,f1655,f1654,f1653,f1652,f1651,f1650,f1649,f1648,f1647,f1646,f1645,f1644,f1643,f1642,f1641,f1640,f1639,f1638,f1637,f1635,f1634,f1632,f1631,f1630,f1629,f1628,f1627,f1626,f1625,f1624,f1623,f1622,f1621,f1620,f1619,f1618,f1617,f1616,f1615,f1614,f1613,f1612,f1611,f1610,f1609,f1608,f1607,f1606,f1605,f1601,f3564,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4208,f4209,f4210,f4211,f4212,f4213,f4215,f4216,f4217,f4218,f4219,f4220,f4221,f4222,f4223,f4224,f4225,f4226,f4227,f4228,f4229,f4230,f4231,f4232,f4316,f4319,f4239,f4320,f4241])).
fof(f4241,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1647])).
fof(f4239,plain,(
  sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0','T_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1645])).
fof(f4232,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1638])).
fof(f4231,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0',X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1637])).
fof(f4230,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1635])).
fof(f4229,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1634])).
fof(f4228,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1632])).
fof(f4227,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'F_0')) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1631])).
fof(f4226,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1630])).
fof(f4225,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1629])).
fof(f4224,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1628])).
fof(f4223,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0',X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1627])).
fof(f4222,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1626])).
fof(f4221,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,'T_0')) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1625])).
fof(f4220,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1624])).
fof(f4219,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1,X0) | ~p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1623])).
fof(f4218,plain,(
  sP1('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1622])).
fof(f4217,plain,(
  ( ! [X0] : (sP1(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1621])).
fof(f4216,plain,(
  sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'F_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1620])).
fof(f4215,plain,(
  ( ! [X0] : (sP1('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1619])).
fof(f4213,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1617])).
fof(f4212,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1616])).
fof(f4211,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1615])).
fof(f4210,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1614])).
fof(f4209,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1613])).
fof(f4208,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1612])).
fof(f4207,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1611])).
fof(f4206,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1610])).
fof(f4205,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1609])).
fof(f4204,plain,(
  ( ! [X0] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1608])).
fof(f4203,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1607])).
fof(f4202,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1606])).
fof(f4201,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK12) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1605])).
fof(f4200,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1601])).
fof(f1606,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',sK14) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f438])).
fof(f1607,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f545])).
fof(f1608,plain,(
  ( ! [X0] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f547])).
fof(f1609,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f549])).
fof(f1610,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f553])).
fof(f1611,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f557])).
fof(f1612,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f559])).
fof(f1613,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f561])).
fof(f1614,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f563])).
fof(f1615,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f567])).
fof(f1616,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f569])).
fof(f1617,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f573])).
fof(f1618,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f591])).
fof(f1619,plain,(
  ( ! [X0] : (sP1('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f624])).
fof(f1620,plain,(
  sP1('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f627])).
fof(f1621,plain,(
  ( ! [X0] : (sP1(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f632])).
fof(f1622,plain,(
  sP1('F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f646])).
fof(f1623,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f668])).
fof(f1624,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f677])).
fof(f1625,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,'T_0')) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f681])).
fof(f1626,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f683])).
fof(f1627,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0',X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f691])).
fof(f1628,plain,(
  ( ! [X0] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f693])).
fof(f1629,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f698])).
fof(f1630,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f711])).
fof(f1631,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,'F_0')) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f715])).
fof(f1632,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f726])).
fof(f1634,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f744])).
fof(f1635,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f746])).
fof(f1637,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0',X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f759])).
fof(f1638,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f761])).
fof(f1639,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f862])).
fof(f1640,plain,(
  sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f863])).
fof(f1641,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f882])).
fof(f1642,plain,(
  ~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1601,f884])).
fof(f1643,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f923])).
fof(f1644,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f947])).
fof(f1645,plain,(
  sP7('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0','T_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f971])).
fof(f1646,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f989])).
fof(f1647,plain,(
  ( ! [X0,X1] : (~sP7(X1,X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1024])).
fof(f1649,plain,(
  ( ! [X0] : (~sP4('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1035])).
fof(f1652,plain,(
  ( ! [X0] : (~sP3('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1056])).
fof(f1654,plain,(
  ( ! [X0] : (~sP6('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1075])).
fof(f1656,plain,(
  ( ! [X0,X1] : (~sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1101])).
fof(f1657,plain,(
  ( ! [X0,X1] : (~sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1135])).
fof(f1658,plain,(
  ( ! [X0] : (~sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1143])).
fof(f1661,plain,(
  ( ! [X0,X1] : (~sP10(X1,X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1240])).
fof(f1662,plain,(
  ( ! [X0,X1] : (~sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1278])).
fof(f1665,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0','T_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1365])).
fof(f1666,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'T_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1367])).
fof(f1669,plain,(
  sP5('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0','F_0') | $spl4),
  inference(backward_demodulation,[],[f1601,f1414])).
fof(f1670,plain,(
  ( ! [X0] : (sP5('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1421])).
fof(f1672,plain,(
  ( ! [X0] : (sP5('const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1430])).
fof(f1673,plain,(
  ( ! [X0] : (sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)),'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1452])).
fof(f1675,plain,(
  ( ! [X0] : (sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1474])).
fof(f1676,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3)))))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1545])).
fof(f1677,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',sK14)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1547])).
fof(f1678,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))),s('type.integer.int',X4))) | s('type.integer.int',sK14) = s('type.integer.int',X4) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1584])).
fof(f1679,plain,(
  ( ! [X2,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int',X2) | s('type.integer.int',X1) = s('type.integer.int',X3)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1587])).
fof(f1680,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) | s('type.integer.int',X4) = s('type.integer.int',X5)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1593])).
fof(f1681,plain,(
  ( ! [X0,X1] : (~sP10('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1)) ) | $spl4),
  inference(forward_demodulation,[],[f1633,f1601])).
fof(f1633,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f728])).
fof(f1682,plain,(
  ( ! [X0,X1] : (~sP8('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1)) ) | $spl4),
  inference(forward_demodulation,[],[f1636,f1601])).
fof(f1636,plain,(
  ( ! [X0,X1] : (sP9('const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0,X1) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0,X1)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f748])).
fof(f1684,plain,(
  ( ! [X0,X1] : (~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X0) | sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1)) ) | $spl4),
  inference(forward_demodulation,[],[f1683,f1601])).
fof(f1683,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12))) | ~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(forward_demodulation,[],[f1668,f1601])).
fof(f1668,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK12)),X1) | sP5(X0,X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))) | ~sP3(X1,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)),X0)) ) | $spl4),
  inference(backward_demodulation,[],[f1601,f1392])).
fof(f1875,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',sK14))),s('type.integer.int',X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1677,f439])).
fof(f1885,plain,(
  ( ! [X8,X7] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X7)))) ) | $spl4),
  inference(superposition,[],[f1677,f390])).
fof(f1905,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.integer.int',X19) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X18),s('type.integer.int',sK14))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X19),s('type.integer.int',X20))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',sK14))),s('type.integer.int',X20)))) ) | $spl4),
  inference(superposition,[],[f439,f1677])).
fof(f1935,plain,(
  ( ! [X103,X105,X102,X104] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X105),s('type.integer.int',X104))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X103),s('type.integer.int',sK14))),s('type.integer.int',X104))) | s('type.integer.int',sK14) = s('type.integer.int',X104) | s('type.integer.int',X105) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X102),s('type.integer.int',sK14)))) ) | $spl4),
  inference(superposition,[],[f1605,f1677])).
fof(f2049,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int',sK14) = s('type.integer.int',X2) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) | $spl4),
  inference(superposition,[],[f1679,f390])).
fof(f2054,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int',X2) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) | $spl4),
  inference(superposition,[],[f1679,f390])).
fof(f2057,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X2) = s('type.integer.int',X3)) ) | $spl4),
  inference(superposition,[],[f1679,f390])).
fof(f2060,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int',X1) | s('type.integer.int',X0) = s('type.integer.int',X2)) ) | $spl4),
  inference(superposition,[],[f1679,f390])).
fof(f3094,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl4),
  inference(superposition,[],[f1676,f390])).
fof(f3097,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2)))))) ) | $spl4),
  inference(superposition,[],[f1676,f390])).
fof(f3121,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4))),s('type.integer.int',X2))) | s('type.integer.int',sK14) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3)))) ) | $spl4),
  inference(superposition,[],[f1678,f390])).
fof(f3125,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',sK14) = s('type.integer.int',X3) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) ) | $spl4),
  inference(superposition,[],[f1678,f390])).
fof(f3128,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))),s('type.integer.int',X3))) | s('type.integer.int',sK14) = s('type.integer.int',X3) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X4)))) ) | $spl4),
  inference(superposition,[],[f1678,f390])).
fof(f3133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X4))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int',X3)))) ) | $spl4),
  inference(superposition,[],[f1678,f390])).
fof(f3170,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X3))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X3))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | s('type.integer.int',X2) = s('type.integer.int',X4)) ) | $spl4),
  inference(superposition,[],[f1680,f390])).
fof(f3173,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X3))) | s('type.integer.int',X2) = s('type.integer.int',X4)) ) | $spl4),
  inference(superposition,[],[f1680,f390])).
fof(f3246,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X36),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X34))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X34),s('type.integer.int',X33))),s('type.integer.int',X35))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X33),s('type.integer.int',X34))) | s('type.integer.int',X35) = s('type.integer.int',X36)) ) | $spl4),
  inference(superposition,[],[f1679,f1518])).
fof(f3283,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X183),s('type.integer.int',X182))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X180),s('type.integer.int',X179))),s('type.integer.int',X181))),s('type.integer.int',X182))) | s('type.integer.int',sK14) = s('type.integer.int',X182) | s('type.integer.int',X183) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X179),s('type.integer.int',X180))),s('type.integer.int',X181)))) ) | $spl4),
  inference(superposition,[],[f1605,f1518])).
fof(f3285,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X192),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X190),s('type.integer.int',X189))),s('type.integer.int',X191))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X193),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X190),s('type.integer.int',X189))),s('type.integer.int',X191))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X189),s('type.integer.int',X190))),s('type.integer.int',X191))) | s('type.integer.int',X192) = s('type.integer.int',X193)) ) | $spl4),
  inference(superposition,[],[f1605,f1518])).
fof(f3286,plain,(
  ( ! [X198,X196,X194,X197,X195] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X195),s('type.integer.int',X194))),s('type.integer.int',X196))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X197),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X194),s('type.integer.int',X195))),s('type.integer.int',X196))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X198),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X194),s('type.integer.int',X195))),s('type.integer.int',X196)))))) ) | $spl4),
  inference(superposition,[],[f1606,f1518])).
fof(f3287,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X202),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X200),s('type.integer.int',X199))),s('type.integer.int',X201))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X203),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X199),s('type.integer.int',X200))),s('type.integer.int',X201))),s('type.integer.int',X202))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X204),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X199),s('type.integer.int',X200))),s('type.integer.int',X201))),s('type.integer.int',X202)))))) ) | $spl4),
  inference(superposition,[],[f1676,f1518])).
fof(f3288,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X206),s('type.integer.int',X205))),s('type.integer.int',X207))),s('type.integer.int',X208))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X209),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X205),s('type.integer.int',X206))),s('type.integer.int',X207))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X210),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X205),s('type.integer.int',X206))),s('type.integer.int',X207)))))))) ) | $spl4),
  inference(superposition,[],[f1676,f1518])).
fof(f3289,plain,(
  ( ! [X212,X216,X215,X213,X211,X214] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X214),s('type.integer.int',X215))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X216),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X212),s('type.integer.int',X211))),s('type.integer.int',X213))))),s('type.integer.int',X215))) | s('type.integer.int',sK14) = s('type.integer.int',X215) | s('type.integer.int',X214) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X211),s('type.integer.int',X212))),s('type.integer.int',X213))),s('type.integer.int',X216)))) ) | $spl4),
  inference(superposition,[],[f1678,f1518])).
fof(f3290,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X220),s('type.integer.int',X221))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X218),s('type.integer.int',X217))),s('type.integer.int',X219))),s('type.integer.int',X222))),s('type.integer.int',X221))) | s('type.integer.int',sK14) = s('type.integer.int',X221) | s('type.integer.int',X220) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X222),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X217),s('type.integer.int',X218))),s('type.integer.int',X219)))))) ) | $spl4),
  inference(superposition,[],[f1678,f1518])).
fof(f3291,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X226),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X224),s('type.integer.int',X223))),s('type.integer.int',X225))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X227),s('type.integer.int',X228))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X224),s('type.integer.int',X223))),s('type.integer.int',X225))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X223),s('type.integer.int',X224))),s('type.integer.int',X225))) | s('type.integer.int',X226) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X228),s('type.integer.int',X227)))) ) | $spl4),
  inference(superposition,[],[f1678,f1518])).
fof(f3292,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X233),s('type.integer.int',X234))),s('type.integer.int',X232))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X230),s('type.integer.int',X229))),s('type.integer.int',X231))),s('type.integer.int',X232))) | s('type.integer.int',sK14) = s('type.integer.int',X232) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X234),s('type.integer.int',X233))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X229),s('type.integer.int',X230))),s('type.integer.int',X231)))) ) | $spl4),
  inference(superposition,[],[f1678,f1518])).
fof(f3293,plain,(
  ( ! [X239,X237,X235,X238,X236] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X239),s('type.integer.int',X238))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X238),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X236),s('type.integer.int',X235))),s('type.integer.int',X237))))) | s('type.integer.int',sK14) = s('type.integer.int',X238) | s('type.integer.int',X239) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X235),s('type.integer.int',X236))),s('type.integer.int',X237)))) ) | $spl4),
  inference(superposition,[],[f1679,f1518])).
fof(f3294,plain,(
  ( ! [X243,X241,X244,X242,X240] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X244),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X241),s('type.integer.int',X240))),s('type.integer.int',X242))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X241),s('type.integer.int',X240))),s('type.integer.int',X242))),s('type.integer.int',X243))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X240),s('type.integer.int',X241))),s('type.integer.int',X242))) | s('type.integer.int',X243) = s('type.integer.int',X244)) ) | $spl4),
  inference(superposition,[],[f1679,f1518])).
fof(f3295,plain,(
  ( ! [X249,X246,X248,X247,X245] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X248),s('type.integer.int',X249))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X246),s('type.integer.int',X245))),s('type.integer.int',X247))),s('type.integer.int',X248))) | s('type.integer.int',sK14) = s('type.integer.int',X248) | s('type.integer.int',X249) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X245),s('type.integer.int',X246))),s('type.integer.int',X247)))) ) | $spl4),
  inference(superposition,[],[f1679,f1518])).
fof(f3296,plain,(
  ( ! [X255,X253,X251,X254,X252,X250] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X253),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X254),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X251),s('type.integer.int',X250))),s('type.integer.int',X252))))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X255),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X254),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X251),s('type.integer.int',X250))),s('type.integer.int',X252))))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X250),s('type.integer.int',X251))),s('type.integer.int',X252))),s('type.integer.int',X254))) | s('type.integer.int',X253) = s('type.integer.int',X255)) ) | $spl4),
  inference(superposition,[],[f1680,f1518])).
fof(f3297,plain,(
  ( ! [X261,X257,X259,X260,X256,X258] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X259),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X257),s('type.integer.int',X256))),s('type.integer.int',X258))),s('type.integer.int',X260))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X261),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X257),s('type.integer.int',X256))),s('type.integer.int',X258))),s('type.integer.int',X260))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X260),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X256),s('type.integer.int',X257))),s('type.integer.int',X258))))) | s('type.integer.int',X259) = s('type.integer.int',X261)) ) | $spl4),
  inference(superposition,[],[f1680,f1518])).
fof(f3298,plain,(
  ( ! [X263,X265,X267,X262,X264,X266] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X267),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X265),s('type.integer.int',X266))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X263),s('type.integer.int',X262))),s('type.integer.int',X264))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X265),s('type.integer.int',X266))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X266),s('type.integer.int',X265))) | s('type.integer.int',X267) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X262),s('type.integer.int',X263))),s('type.integer.int',X264)))) ) | $spl4),
  inference(superposition,[],[f1680,f1518])).
fof(f4320,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4240,f579])).
fof(f4240,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1646])).
fof(f4319,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4238,f577])).
fof(f4238,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1644])).
fof(f4316,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4237,f575])).
fof(f4237,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1643])).
fof(f4315,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4236,f900])).
fof(f4236,plain,(
  ~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1642])).
fof(f4314,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4235,f897])).
fof(f4235,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1641])).
fof(f4313,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4234,f899])).
fof(f4234,plain,(
  sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1640])).
fof(f4312,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4233,f896])).
fof(f4233,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),'T_0') | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1639])).
fof(f4311,plain,(
  $false | ($spl4 | $spl52)),
  inference(subsumption_resolution,[],[f4214,f587])).
fof(f4214,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ($spl4 | $spl52)),
  inference(backward_demodulation,[],[f3564,f1618])).
fof(f4192,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f4080])).
fof(f4080,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f538,f541])).
fof(f3564,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK14) | $spl52),
  inference(cnf_transformation,[],[f3564_D])).
fof(f3564_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK14) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3625,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK12) = s('type.integer.int',sK14) | s('type.integer.int',sK14) = s('type.integer.int',X10)) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f1679,f540])).
fof(f4053,plain,(
  $false | ($spl1 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f4003,f3618])).
fof(f3618,plain,(
  ( ! [X112,X113] : (s('type.integer.int',X112) = s('type.integer.int',X113)) ) | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f1981,f1982])).
fof(f4003,plain,(
  ( ! [X329] : (s('type.integer.int',sK13) != s('type.integer.int',X329)) ) | ($spl1 | $spl10 | $spl12)),
  inference(superposition,[],[f539,f3618])).
fof(f4052,plain,(
  $false | ($spl1 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f3840,f3618])).
fof(f3840,plain,(
  ( ! [X329] : (s('type.integer.int',sK13) != s('type.integer.int',X329)) ) | ($spl1 | $spl10 | $spl12)),
  inference(superposition,[],[f539,f3618])).
fof(f3705,plain,(
  $false | ($spl1 | $spl4 | $spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f539,f539,f3618,f1679])).
fof(f3701,plain,(
  $false | ($spl1 | $spl4 | $spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f539,f539,f3618,f1605])).
fof(f3690,plain,(
  $false | ($spl1 | $spl4 | $spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f539,f539,f3618,f1679])).
fof(f3686,plain,(
  $false | ($spl1 | $spl4 | $spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f539,f539,f3618,f1605])).
fof(f3679,plain,(
  $false | ($spl1 | $spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f539,f3618])).
fof(f3565,plain,(
  ( ! [X10] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK13) = s('type.integer.int',X10)) ) | $spl54),
  inference(cnf_transformation,[],[f3565_D])).
fof(f3565_D,plain,(
  ( ! [X10] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK13) = s('type.integer.int',X10)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3496,plain,(
  ( ! [X10] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X10),s('type.integer.int',sK12))) | s('type.integer.int',sK12) = s('type.integer.int',sK14) | s('type.integer.int',sK13) = s('type.integer.int',X10)) ) | ($spl2 | $spl4 | $spl8)),
  inference(superposition,[],[f1679,f1986])).
fof(f1986,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ($spl2 | $spl4 | $spl8)),
  inference(backward_demodulation,[],[f1980,f1882])).
fof(f1882,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14)))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f1677,f540])).
fof(f2646,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2624,f306])).
fof(f2624,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(backward_demodulation,[],[f2510,f542])).
fof(f2645,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f305,f304,f303,f306,f542,f314,f327,f326,f331,f330,f383,f387,f386,f390,f394,f391,f405,f404,f543,f423,f420,f419,f418,f436,f439,f438,f437,f458,f457,f456,f455,f454,f453,f452,f462,f461,f460,f459,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f528,f478,f477,f476,f475,f485,f484,f483,f482,f491,f488,f487,f486,f498,f497,f496,f527,f494,f493,f492,f502,f501,f500,f499,f506,f505,f504,f503,f515,f514,f513,f526,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f545,f544,f547,f546,f549,f548,f550,f551,f553,f552,f554,f555,f557,f556,f559,f558,f561,f560,f563,f562,f564,f565,f567,f566,f569,f568,f570,f571,f573,f572,f574,f583,f582,f581,f580,f579,f578,f577,f576,f575,f587,f586,f591,f610,f616,f618,f622,f624,f627,f630,f632,f641,f646,f647,f650,f653,f656,f663,f664,f666,f668,f675,f677,f681,f683,f684,f690,f691,f693,f696,f698,f709,f711,f715,f726,f724,f728,f746,f744,f748,f758,f756,f759,f761,f769,f831,f855,f854,f863,f862,f875,f873,f884,f882,f900,f899,f897,f896,f916,f923,f940,f947,f951,f957,f966,f971,f970,f982,f989,f993,f1024,f1025,f1035,f1039,f1046,f1056,f1065,f1075,f1079,f1078,f1086,f1094,f1101,f1105,f1135,f1143,f1144,f1157,f1178,f1177,f1176,f1175,f1219,f1218,f1232,f1240,f1244,f1278,f1343,f1353,f1377,f1376,f1371,f1370,f1367,f1365,f1364,f1392,f1430,f1428,f1425,f1421,f1414,f1413,f1474,f1471,f1469,f1452,f1487,f1499,f1518,f1520,f1545,f1547,f1565,f1566,f1584,f1587,f1593,f2092,f2130,f2190,f2191,f2192,f2193,f2236,f2508,f2512,f2516,f2519,f2522,f2525,f2528,f2531,f2534,f2537,f2540,f2543,f2546,f2549,f2552,f2555,f2558,f2560,f2561,f2562,f2563,f2564,f2565,f2566,f2567,f2568,f2569,f2570,f2571,f2572,f2573,f2574,f2575,f2576,f2577,f2578,f2579,f2580,f2581,f2582,f2302,f2303,f2304,f2305,f2588,f2590,f2593,f2595,f2597,f2599,f2601,f2603,f2605,f2607,f2609,f2611,f2613,f2615,f2617,f2619,f2621,f2510,f2644])).
fof(f2644,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(subsumption_resolution,[],[f2622,f2510])).
fof(f2622,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl18),
  inference(backward_demodulation,[],[f2510,f330])).
fof(f2559,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f2559_D])).
fof(f2559_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2621,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2620])).
fof(f2620,plain,(
  ( ! [X94,X95,X93,X96] : (~sP8('T_0',X94,X95) | sP9('T_0',X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f2333,f831])).
fof(f2333,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | ~sP8('T_0',X94,X95) | sP9('T_0',X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f527,f769])).
fof(f2556,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f2556_D])).
fof(f2556_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2619,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2618])).
fof(f2618,plain,(
  ( ! [X92,X90,X91,X89] : (~sP10('T_0',X90,X91) | sP11('T_0',X90,X91) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f2332,f831])).
fof(f2332,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | ~sP10('T_0',X90,X91) | sP11('T_0',X90,X91) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f526,f769])).
fof(f2553,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f2553_D])).
fof(f2553_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2617,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2616])).
fof(f2616,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP11('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f2331,f831])).
fof(f2331,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP11('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f523,f769])).
fof(f2550,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f2550_D])).
fof(f2550_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2615,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2614])).
fof(f2614,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP11('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f2330,f831])).
fof(f2330,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP11('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f522,f769])).
fof(f2547,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f2547_D])).
fof(f2547_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2613,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2612])).
fof(f2612,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X78)) | sP11('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f2329,f831])).
fof(f2329,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | sP11('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f521,f769])).
fof(f2544,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f2544_D])).
fof(f2544_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2611,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2610])).
fof(f2610,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP10(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f2328,f831])).
fof(f2328,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP10(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f516,f769])).
fof(f2541,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2541_D])).
fof(f2541_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2609,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2608])).
fof(f2608,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f2327,f831])).
fof(f2327,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f506,f769])).
fof(f2538,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2538_D])).
fof(f2538_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2607,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2606])).
fof(f2606,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP9('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f2326,f831])).
fof(f2326,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | sP9('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f505,f769])).
fof(f2535,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2535_D])).
fof(f2535_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2605,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2604])).
fof(f2604,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X62)) | sP9('T_0',X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f2325,f831])).
fof(f2325,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~p(s(bool,X62)) | sP9('T_0',X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f504,f769])).
fof(f2532,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f2532_D])).
fof(f2532_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2603,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2602])).
fof(f2602,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | ~sP8('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f2324,f831])).
fof(f2324,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | p(s(bool,X58)) | p(s(bool,X59)) | ~sP8('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f499,f769])).
fof(f2529,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2529_D])).
fof(f2529_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2601,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2600])).
fof(f2600,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X55,X54,'T_0') | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f2323,f831])).
fof(f2323,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X55,X54,'T_0') | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f491,f769])).
fof(f2526,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f2526_D])).
fof(f2526_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2599,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2598])).
fof(f2598,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f2322,f831])).
fof(f2322,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f482,f769])).
fof(f2523,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f2523_D])).
fof(f2523_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2597,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2596])).
fof(f2596,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP3(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f2321,f831])).
fof(f2321,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP3(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f463,f769])).
fof(f2520,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2520_D])).
fof(f2520_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2595,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2594])).
fof(f2594,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP4(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f2320,f831])).
fof(f2320,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP4(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f459,f769])).
fof(f2517,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl22),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2593,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2592,f560])).
fof(f2592,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2591])).
fof(f2591,plain,(
  ( ! [X39,X37,X38,X40] : (sP5(X38,X39,'T_0') | ~sP4(X38,'T_0',X39) | ~sP3(X39,'T_0',X38) | ~sP2('T_0',X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f2319,f831])).
fof(f2319,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | sP5(X38,X39,'T_0') | ~sP4(X38,'T_0',X39) | ~sP3(X39,'T_0',X38) | ~sP2('T_0',X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f456,f769])).
fof(f2513,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl20),
  inference(cnf_transformation,[],[f2513_D])).
fof(f2513_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2590,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2589])).
fof(f2589,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f2318,f831])).
fof(f2318,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f423,f769])).
fof(f2509,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl16),
  inference(cnf_transformation,[],[f2509_D])).
fof(f2509_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2588,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2587])).
fof(f2587,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f2317,f831])).
fof(f2317,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f420,f769])).
fof(f2558,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2557])).
fof(f2557,plain,(
  ( ! [X94,X95,X93,X96] : (~sP8('T_0',X94,X95) | sP9('T_0',X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f2263,f831])).
fof(f2263,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | ~sP8('T_0',X94,X95) | sP9('T_0',X94,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f527,f769])).
fof(f2555,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2554])).
fof(f2554,plain,(
  ( ! [X92,X90,X91,X89] : (~sP10('T_0',X90,X91) | sP11('T_0',X90,X91) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f2262,f831])).
fof(f2262,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | ~sP10('T_0',X90,X91) | sP11('T_0',X90,X91) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f526,f769])).
fof(f2552,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2551])).
fof(f2551,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP11('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f2261,f831])).
fof(f2261,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP11('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f523,f769])).
fof(f2549,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2548])).
fof(f2548,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP11('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f2260,f831])).
fof(f2260,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP11('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f522,f769])).
fof(f2546,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2545])).
fof(f2545,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X78)) | sP11('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f2259,f831])).
fof(f2259,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | sP11('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f521,f769])).
fof(f2543,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f2542])).
fof(f2542,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP10(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f2258,f831])).
fof(f2258,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP10(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f516,f769])).
fof(f2540,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2539])).
fof(f2539,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f2257,f831])).
fof(f2257,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f506,f769])).
fof(f2537,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2536])).
fof(f2536,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP9('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f2256,f831])).
fof(f2256,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | sP9('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f505,f769])).
fof(f2534,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2533])).
fof(f2533,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X62)) | sP9('T_0',X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f2255,f831])).
fof(f2255,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~p(s(bool,X62)) | sP9('T_0',X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f504,f769])).
fof(f2531,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f2530])).
fof(f2530,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | ~sP8('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f2254,f831])).
fof(f2254,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | p(s(bool,X58)) | p(s(bool,X59)) | ~sP8('T_0',X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f499,f769])).
fof(f2528,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f2527])).
fof(f2527,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X55,X54,'T_0') | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f2253,f831])).
fof(f2253,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X55,X54,'T_0') | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f491,f769])).
fof(f2525,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f2524])).
fof(f2524,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f2252,f831])).
fof(f2252,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f482,f769])).
fof(f2522,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2521])).
fof(f2521,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP3(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f2251,f831])).
fof(f2251,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP3(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f463,f769])).
fof(f2519,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2518])).
fof(f2518,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP4(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f2250,f831])).
fof(f2250,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP4(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f459,f769])).
fof(f2516,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2515,f560])).
fof(f2515,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2514])).
fof(f2514,plain,(
  ( ! [X39,X37,X38,X40] : (sP5(X38,X39,'T_0') | ~sP4(X38,'T_0',X39) | ~sP3(X39,'T_0',X38) | ~sP2('T_0',X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f2249,f831])).
fof(f2249,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | sP5(X38,X39,'T_0') | ~sP4(X38,'T_0',X39) | ~sP3(X39,'T_0',X38) | ~sP2('T_0',X39,X38) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f456,f769])).
fof(f2512,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2511])).
fof(f2511,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f2248,f831])).
fof(f2248,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f423,f769])).
fof(f2508,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2507])).
fof(f2507,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f2247,f831])).
fof(f2247,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f420,f769])).
fof(f1982,plain,(
  ( ! [X116,X114,X117] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X116),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X117),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14)))))) ) | $spl12),
  inference(cnf_transformation,[],[f1982_D])).
fof(f1982_D,plain,(
  ( ! [X116,X114,X117] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X116),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X117),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1983,plain,(
  ( ! [X115] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X115),s('type.integer.int',sK14)))) ) | $spl14),
  inference(cnf_transformation,[],[f1983_D])).
fof(f1983_D,plain,(
  ( ! [X115] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X115),s('type.integer.int',sK14)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1979,plain,(
  ( ! [X116,X114,X117,X115] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X115),s('type.integer.int',sK14))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X116),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X117),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14)))))) ) | $spl4),
  inference(superposition,[],[f1606,f1677])).
fof(f1980,plain,(
  ( ! [X110] : (s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X110),s('type.integer.int',sK14)))) ) | $spl8),
  inference(cnf_transformation,[],[f1980_D])).
fof(f1980_D,plain,(
  ( ! [X110] : (s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X110),s('type.integer.int',sK14)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1981,plain,(
  ( ! [X111,X112,X113] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) | s('type.integer.int',X112) = s('type.integer.int',X113)) ) | $spl10),
  inference(cnf_transformation,[],[f1981_D])).
fof(f1981_D,plain,(
  ( ! [X111,X112,X113] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) | s('type.integer.int',X112) = s('type.integer.int',X113)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1978,plain,(
  ( ! [X111,X112,X110,X113] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X110),s('type.integer.int',sK14))) | s('type.integer.int',X112) = s('type.integer.int',X113)) ) | $spl4),
  inference(superposition,[],[f1605,f1677])).
fof(f1938,plain,(
  ( ! [X116,X114,X117,X115] : (s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X115),s('type.integer.int',sK14))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X116),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X117),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X114),s('type.integer.int',sK14)))))) ) | $spl4),
  inference(superposition,[],[f1606,f1677])).
fof(f1937,plain,(
  ( ! [X111,X112,X110,X113] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X112),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X113),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X111),s('type.integer.int',sK14))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X110),s('type.integer.int',sK14))) | s('type.integer.int',X112) = s('type.integer.int',X113)) ) | $spl4),
  inference(superposition,[],[f1605,f1677])).
fof(f1601,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1602,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))) | s('type.integer.int',sK12) = s('type.integer.int',X0)) ) | $spl6),
  inference(cnf_transformation,[],[f1602_D])).
fof(f1602_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))) | s('type.integer.int',sK12) = s('type.integer.int',X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1595,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',sK12) = s('type.integer.int',X0)) ) | $spl2),
  inference(superposition,[],[f437,f540])).
fof(f1586,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK14))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',sK12) = s('type.integer.int',X0)) ) | $spl2),
  inference(superposition,[],[f437,f540])).
fof(f538,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) | $spl0),
  inference(cnf_transformation,[],[f538_D])).
fof(f538_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int',sK14) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f540,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | $spl2),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f304,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | s('type.integer.int',sK13) = s('type.integer.int',sK14)),
  inference(cnf_transformation,[],[f201])).
fof(f539,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) | $spl1),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK14) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f541,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | $spl3),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f305,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK14))) | s('type.integer.int',sK13) != s('type.integer.int',sK14)),
  inference(cnf_transformation,[],[f201])).
% SZS output end Proof for thm.int_arith
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 488
% Passive clauses: 1608
% Generated clauses: 7750
% Final active clauses: 202
% Final passive clauses: 188
% Input formulas: 60
% Initial clauses: 223
% 
% Duplicate literals: 81
% Trivial inequalities: 28
% Fw subsumption resolutions: 317
% Fw demodulations: 863
% Bw demodulations: 713
% Condensations: 122
% Global subsumptions: 6
% 
% Simple tautologies: 108
% Equational tautologies: 125
% Forward subsumptions: 4278
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2051
% Factoring: 1
% Forward superposition: 1100
% Backward superposition: 1583
% Self superposition: 69
% Equality factoring: 26
% Equality resolution: 19
% 
% Split clauses: 86
% Split components: 177
% SAT solver clauses: 1893
% SAT solver unit clauses: 274
% SAT solver binary clauses: 481
% 
% Sat splits: 86
% Sat splitting refutations: 105
% 
% Memory used [KB]: 4349
% Time elapsed: 0.278 s
% ------------------------------
% ------------------------------
% Success in time 1.597 s
