Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (60330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 202
% Passive clauses: 2121
% Generated clauses: 23148
% Final active clauses: 180
% Final passive clauses: 1593
% Input formulas: 25
% Initial clauses: 58
% 
% Duplicate literals: 1149
% Trivial inequalities: 232
% Fw subsumption resolutions: 3567
% Fw demodulations: 3908
% Bw demodulations: 21
% 
% Simple tautologies: 376
% Equational tautologies: 25
% Forward subsumptions: 10557
% 
% Binary resolution: 648
% Forward superposition: 8095
% Backward superposition: 5100
% Self superposition: 110
% Equality factoring: 36
% Equality resolution: 41
% 
% Split clauses: 737
% Split components: 1485
% SAT solver clauses: 1530
% SAT solver unit clauses: 18
% SAT solver binary clauses: 120
% 
% Sat splits: 737
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1529
% TWLsolver calls for satisfiability: 328
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (60346)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 661
% Generated clauses: 13405
% Final active clauses: 110
% Final passive clauses: 425
% Input formulas: 25
% Initial clauses: 58
% 
% Duplicate literals: 823
% Trivial inequalities: 143
% Fw subsumption resolutions: 649
% Fw demodulations: 1016
% Bw demodulations: 56
% Condensations: 1777
% Global subsumptions: 5
% 
% Simple tautologies: 255
% Equational tautologies: 51
% Forward subsumptions: 6396
% Bw demodulations to eq. taut.: 53
% 
% Binary resolution: 191
% Forward superposition: 3873
% Backward superposition: 5305
% Self superposition: 194
% Equality factoring: 129
% Equality resolution: 30
% 
% Split clauses: 114
% Split components: 233
% SAT solver clauses: 1184
% SAT solver unit clauses: 62
% SAT solver binary clauses: 35
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 114
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1159
% TWLsolver calls for satisfiability: 311
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f3587,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f105,f182,f180,f104,f181,f179,f798,f799,f806,f803,f800,f801,f809,f802,f811,f804,f808,f834,f848,f846,f862,f863,f871,f864,f865,f874,f866,f876,f1051,f1052,f1053,f1054,f1055,f1056,f1062,f1063,f1064,f1076,f1065,f870,f869,f1082,f1083,f875,f873,f1492,f1494,f1496,f1509,f1511,f1514,f1323,f1322,f1517,f1320,f1319,f1521,f1369,f1524,f1399,f1398,f1396,f1515,f1394,f1379,f1519,f1380,f1278,f1516,f1277,f1520,f1390,f810,f1522,f1772,f1774,f1777,f1780,f1784,f1788,f1831,f1833,f1834,f1836,f1838,f1845,f1849,f3584,f3586])).
fof(f3586,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3585,f183])).
fof(f183,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(forward_demodulation,[],[f139,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X9,X10] : s(bool,'F_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X6,X9),s(X6,X10)))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.relation.EMPTY_REL_DEF')).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.relation.EMPTY_REL_2'(s(X0,X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1))),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X21,X22] : s(bool,'const.relation.EMPTY_REL_2'(s(X6,X21),s(X6,X22))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'),s(X6,X21))),s(X6,X22)))),
  file('thm.relation.REMPTY_SUBSET|split|1','ARITY_0_2_const.relation.EMPTY_REL')).
fof(f3585,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3577,f3095])).
fof(f3095,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3035,f116])).
fof(f116,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f3035,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3034,f108])).
fof(f108,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.REMPTY_SUBSET|split|1',aHL_FALSITY)).
fof(f3034,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f3033,f183])).
fof(f3033,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,X0))),s(sK0,X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2992,f106])).
fof(f106,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.REMPTY_SUBSET|split|1',aHL_TRUTH)).
fof(f2992,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X1)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,X0))),s(sK0,X1))))) ) | $spl0),
  inference(superposition,[],[f145,f1859])).
fof(f1859,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f179,f112])).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f145,plain,(
  ( ! [X6,X2,X0,X7,X3,X1] : (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK4(X3,X2,X1,X0)))),s(X1,sK5(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK4(X3,X2,X1,X0)))),s(X1,sK5(X3,X2,X1,X0)))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X6))),s(X1,X7)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(rectify,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (? [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) <=> ! [X4,X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X11] : ! [X19,X20] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X11,bool)),X19),s(fun(X6,fun(X11,bool)),X20)))) <=> ! [X9,X10] : (p(s(bool,happ(s(fun(X11,bool),happ(s(fun(X6,fun(X11,bool)),X19),s(X6,X9))),s(X11,X10)))) => p(s(bool,happ(s(fun(X11,bool),happ(s(fun(X6,fun(X11,bool)),X20),s(X6,X9))),s(X11,X10))))))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.relation.RSUBSET')).
fof(f3577,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1629,f144])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11,X6] : ! [X12,X13] : (s(fun(X6,X11),X12) = s(fun(X6,X11),X13) <=> ! [X9] : s(X11,happ(s(fun(X6,X11),X12),s(X6,X9))) = s(X11,happ(s(fun(X6,X11),X13),s(X6,X9))))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.bool.FUN_EQ_THM')).
fof(f1629,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))) | $spl3),
  inference(unit_resulting_resolution,[],[f182,f144])).
fof(f3584,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3583,f183])).
fof(f3583,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3580,f3095])).
fof(f3580,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool))))),s(sK0,sK3(happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3(sK1,'const.relation.EMPTY_REL_0',sK0,fun(sK0,bool)))),sK0,bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1629,f144])).
fof(f1849,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f105,f104,f106,f108,f109,f112,f111,f116,f115,f138,f183,f142,f144,f143,f147,f146,f145,f173,f185,f218,f226,f227,f315,f323,f376,f677,f678,f456,f457,f681,f779,f781,f785,f787,f912,f914,f916,f921,f923,f926,f928,f930,f1017,f1019,f1023,f1025,f1040,f1042,f873,f869,f1134,f1136,f1140,f1142,f1156,f1159,f1261,f1782,f1786,f1649,f1760,f1832,f1835,f1837,f1844,f1848])).
fof(f1848,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1733,f1760])).
fof(f1733,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | $spl20),
  inference(equality_factoring,[],[f869])).
fof(f1844,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1843,f1760])).
fof(f1843,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1))) ) | $spl20),
  inference(subsumption_resolution,[],[f1666,f111])).
fof(f1666,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl20),
  inference(superposition,[],[f111,f869])).
fof(f1837,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1796,f1760])).
fof(f1796,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)),X2,X0,bool)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) ) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1760,f926])).
fof(f1835,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1795,f1760])).
fof(f1795,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(X2,happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)),X0,bool)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) ) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1760,f912])).
fof(f1832,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1790,f1760])).
fof(f1790,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1760,f115])).
fof(f1760,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl20 | $spl24)),
  inference(duplicate_literal_removal,[],[f1656])).
fof(f1656,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl20 | $spl24)),
  inference(superposition,[],[f869,f873])).
fof(f1649,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(superposition,[],[f869,f218])).
fof(f1786,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X20,X21),s(X20,X22)))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1785,f185])).
fof(f1785,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X20,X21),s(X20,X22)))) ) | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1646,f873])).
fof(f1646,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X20,X21),s(X20,X22)))) ) | $spl20),
  inference(superposition,[],[f869,f138])).
fof(f1782,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'),s(X17,X18))),s(X17,X19)))) ) | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1781,f185])).
fof(f1781,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'),s(X17,X18))),s(X17,X19)))) ) | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1645,f873])).
fof(f1645,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),happ(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'),s(X17,X18))),s(X17,X19)))) ) | $spl20),
  inference(superposition,[],[f869,f183])).
fof(f1261,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X15))),s(X12,X16)))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X14),s(X11,X15))),s(X12,X16)))) | s(bool,X17) = s(bool,'const.relation.RSUBSET_2'(s(fun(X11,fun(X12,bool)),X13),s(fun(X11,fun(X12,bool)),X14))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1251,f106])).
fof(f1251,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X13),s(X11,X15))),s(X12,X16)))) | p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X11,fun(X12,bool)),X14),s(X11,X15))),s(X12,X16)))) | s(bool,X17) = s(bool,'const.relation.RSUBSET_2'(s(fun(X11,fun(X12,bool)),X13),s(fun(X11,fun(X12,bool)),X14))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f145,f218])).
fof(f1159,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f873,f218])).
fof(f1156,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X7,fun(X6,bool)),X9),s(fun(X7,fun(X6,bool)),X8)))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,sK4(X8,X9,X6,X7)))),s(X6,sK5(X8,X9,X6,X7)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1147,f106])).
fof(f1147,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X7,fun(X6,bool)),X9),s(fun(X7,fun(X6,bool)),X8)))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,sK4(X8,X9,X6,X7)))),s(X6,sK5(X8,X9,X6,X7)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f147,f218])).
fof(f1142,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X13,fun(bool,bool)),X11),s(X13,sK4(X11,X12,bool,X13)))),s(bool,X14)))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X13,fun(bool,bool)),X12),s(fun(X13,fun(bool,bool)),X11)))) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f147,f218])).
fof(f1140,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X5,fun(bool,bool)),X3),s(X5,sK4(X3,X4,bool,X5)))),s(bool,'T_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X5,fun(bool,bool)),X4),s(fun(X5,fun(bool,bool)),X3)))) | s(bool,X6) = s(bool,sK5(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f147,f218])).
fof(f1136,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X11),s(bool,X14))),s(X13,sK5(X11,X12,X13,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X13,bool)),X12),s(fun(bool,fun(X13,bool)),X11)))) | s(bool,'T_0') = s(bool,sK4(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f147,f218])).
fof(f1134,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X3),s(bool,'T_0'))),s(X5,sK5(X3,X4,X5,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X5,bool)),X4),s(fun(bool,fun(X5,bool)),X3)))) | s(bool,X6) = s(bool,sK4(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f147,f218])).
fof(f1042,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3)))) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK4(X3,X2,X1,X0)))),s(X1,sK5(X3,X2,X1,X0))))) )),
  inference(condensation,[],[f1041])).
fof(f1041,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,X20)) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X17,fun(X16,bool)),X18),s(fun(X17,fun(X16,bool)),X19)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X17,fun(X16,bool)),X18),s(X17,sK4(X19,X18,X16,X17)))),s(X16,sK5(X19,X18,X16,X17))))) )),
  inference(subsumption_resolution,[],[f1032,f111])).
fof(f1032,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,X20)) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X17,fun(X16,bool)),X18),s(fun(X17,fun(X16,bool)),X19)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X17,fun(X16,bool)),X18),s(X17,sK4(X19,X18,X16,X17)))),s(X16,sK5(X19,X18,X16,X17)))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f146,f218])).
fof(f1040,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(subsumption_resolution,[],[f1028,f108])).
fof(f1028,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(superposition,[],[f146,f183])).
fof(f1025,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X13,fun(bool,bool)),X12),s(X13,sK4(X11,X12,bool,X13)))),s(bool,X14)))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X13,fun(bool,bool)),X12),s(fun(X13,fun(bool,bool)),X11)))) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f146,f218])).
fof(f1023,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X5,fun(bool,bool)),X4),s(X5,sK4(X3,X4,bool,X5)))),s(bool,'T_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X5,fun(bool,bool)),X4),s(fun(X5,fun(bool,bool)),X3)))) | s(bool,X6) = s(bool,sK5(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f146,f218])).
fof(f1019,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(X13,bool),happ(s(fun(bool,fun(X13,bool)),X12),s(bool,X14))),s(X13,sK5(X11,X12,X13,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X13,bool)),X12),s(fun(bool,fun(X13,bool)),X11)))) | s(bool,'T_0') = s(bool,sK4(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f146,f218])).
fof(f1017,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X4),s(bool,'T_0'))),s(X5,sK5(X3,X4,X5,bool))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(bool,fun(X5,bool)),X4),s(fun(bool,fun(X5,bool)),X3)))) | s(bool,X6) = s(bool,sK4(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f146,f218])).
fof(f930,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) != s(bool,happ(s(fun(X14,bool),X16),s(X14,sK3(X15,X16,X14,bool)))) | s(fun(X14,bool),X15) = s(fun(X14,bool),X16) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK3(X15,X16,X14,bool)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f144,f218])).
fof(f928,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X8),s(X6,sK3(X7,X8,X6,bool)))) | s(fun(X6,bool),X7) = s(fun(X6,bool),X8) | s(bool,X9) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK3(X7,X8,X6,bool)))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f144,f218])).
fof(f926,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)),X2,X0,bool)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) )),
  inference(superposition,[],[f144,f183])).
fof(f923,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,happ(s(fun(bool,X13),X11),s(bool,X14))) != s(X13,happ(s(fun(bool,X13),X12),s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f144,f218])).
fof(f921,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK3(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f144,f218])).
fof(f916,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) != s(bool,happ(s(fun(X14,bool),X16),s(X14,sK3(X16,X15,X14,bool)))) | s(fun(X14,bool),X15) = s(fun(X14,bool),X16) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK3(X16,X15,X14,bool)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f144,f218])).
fof(f914,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X8),s(X6,sK3(X8,X7,X6,bool)))) | s(fun(X6,bool),X7) = s(fun(X6,bool),X8) | s(bool,X9) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK3(X8,X7,X6,bool)))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f144,f218])).
fof(f912,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK3(X2,happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)),X0,bool)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X1)))) )),
  inference(superposition,[],[f144,f183])).
fof(f787,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) != s(bool,f(s(X11,sK2(X12,X13,bool,X11)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK2(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f142,f218])).
fof(f785,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK2(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK2(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f142,f218])).
fof(f781,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,f(s(bool,X14))) != s(X13,g(s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK2(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f142,f218])).
fof(f779,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK2(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f142,f218])).
fof(f681,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f474,f106])).
fof(f474,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f115,f218])).
fof(f457,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f106,f218])).
fof(f456,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f111,f218])).
fof(f678,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) = s(bool,'const.relation.EMPTY_REL_2'(s(X19,X20),s(X19,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f388,f185])).
fof(f388,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X22) = s(bool,'const.relation.EMPTY_REL_2'(s(X19,X20),s(X19,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f218,f138])).
fof(f677,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X18) = s(bool,happ(s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),'const.relation.EMPTY_REL_0'),s(X15,X16))),s(X15,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f387,f185])).
fof(f387,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,happ(s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),'const.relation.EMPTY_REL_0'),s(X15,X16))),s(X15,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f218,f183])).
fof(f376,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f218,f218])).
fof(f323,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.relation.EMPTY_REL_0'),s(bool,X4))),s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f183,f109])).
fof(f315,plain,(
  ( ! [X4,X3] : (s(bool,X3) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.relation.EMPTY_REL_0'),s(bool,X3))),s(bool,X4))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f183,f109])).
fof(f227,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X8),s(bool,X9))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f138,f109])).
fof(f226,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X7),s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f138,f109])).
fof(f218,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f109,f109])).
fof(f185,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f108,f111])).
fof(f173,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f172])).
fof(f172,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f161])).
fof(f161,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X17,X10,X18] : ((s(bool,X9) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X10) = s(bool,X18))) => ((p(s(bool,X9)) => p(s(bool,X10))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.bool.IMP_CONG')).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK4(X3,X2,X1,X0)))),s(X1,sK5(X3,X2,X1,X0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f147,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,sK4(X3,X2,X1,X0)))),s(X1,sK5(X3,X2,X1,X0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f143,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.REMPTY_SUBSET|split|1',aHL_EXT)).
fof(f115,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f111,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f75])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.REMPTY_SUBSET|split|1',aHL_BOOL_CASES)).
fof(f1845,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f105,f104,f106,f108,f109,f112,f111,f116,f115,f138,f183,f142,f144,f143,f147,f146,f145,f173,f185,f218,f226,f227,f315,f323,f376,f677,f678,f456,f457,f681,f779,f781,f785,f787,f912,f914,f916,f921,f923,f926,f928,f930,f1017,f1019,f1023,f1025,f1040,f1042,f873,f869,f1134,f1136,f1140,f1142,f1156,f1159,f1261,f1782,f1786,f1649,f1760,f1832,f1835,f1837,f1844])).
fof(f1838,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f105,f104,f106,f108,f109,f112,f111,f116,f115,f138,f183,f142,f144,f143,f147,f146,f145,f173,f185,f218,f226,f227,f315,f323,f376,f677,f678,f456,f457,f681,f779,f781,f785,f787,f912,f914,f916,f921,f923,f926,f928,f930,f1017,f1019,f1023,f1025,f1040,f1042,f873,f869,f1134,f1136,f1140,f1142,f1156,f1159,f1261,f1782,f1786,f1649,f1760,f1832,f1835,f1837])).
fof(f1836,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f105,f104,f106,f108,f109,f112,f111,f116,f115,f138,f183,f142,f144,f143,f147,f146,f145,f173,f185,f218,f226,f227,f315,f323,f376,f677,f678,f456,f457,f681,f779,f781,f785,f787,f912,f914,f916,f921,f923,f926,f928,f930,f1017,f1019,f1023,f1025,f1040,f1042,f873,f869,f1134,f1136,f1140,f1142,f1156,f1159,f1261,f1782,f1786,f1649,f1760,f1832,f1835])).
fof(f1834,plain,(
  $false | ($spl9 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1792,f1760])).
fof(f1792,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl20 | $spl24)),
  inference(backward_demodulation,[],[f1760,f806])).
fof(f1833,plain,(
  $false | ($spl20 | $spl24)),
  inference(global_subsumption,[],[f105,f104,f106,f108,f109,f112,f111,f116,f115,f138,f183,f142,f144,f143,f147,f146,f145,f173,f185,f218,f226,f227,f315,f323,f376,f677,f678,f456,f457,f681,f779,f781,f785,f787,f912,f914,f916,f921,f923,f926,f928,f930,f1017,f1019,f1023,f1025,f1040,f1042,f873,f869,f1134,f1136,f1140,f1142,f1156,f1159,f1261,f1782,f1786,f1649,f1760,f1832])).
fof(f1831,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1789,f106])).
fof(f1789,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1760,f108])).
fof(f1788,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1787,f185])).
fof(f1787,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1786,f138])).
fof(f1784,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1783,f185])).
fof(f1783,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f1782,f183])).
fof(f1780,plain,(
  $false | ($spl9 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1779,f806])).
fof(f1779,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl20 | $spl24)),
  inference(condensation,[],[f1778])).
fof(f1778,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl9 | $spl20 | $spl24)),
  inference(forward_demodulation,[],[f1765,f873])).
fof(f1765,plain,(
  ( ! [X10] : (s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl9 | $spl20)),
  inference(duplicate_literal_removal,[],[f1640])).
fof(f1640,plain,(
  ( ! [X10] : (s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl9 | $spl20)),
  inference(superposition,[],[f869,f904])).
fof(f904,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(trivial_inequality_removal,[],[f898])).
fof(f898,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f806,f218])).
fof(f1777,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1776,f185])).
fof(f1776,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(condensation,[],[f1775])).
fof(f1775,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1770,f873])).
fof(f1770,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f1635])).
fof(f1635,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(superposition,[],[f869,f109])).
fof(f1774,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1773,f185])).
fof(f1773,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl24)),
  inference(forward_demodulation,[],[f1633,f873])).
fof(f1633,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f185,f869])).
fof(f1772,plain,(
  $false | ($spl9 | $spl20 | $spl24)),
  inference(subsumption_resolution,[],[f1771,f806])).
fof(f1771,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl20 | $spl24)),
  inference(forward_demodulation,[],[f1634,f873])).
fof(f1634,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl20)),
  inference(unit_resulting_resolution,[],[f806,f869])).
fof(f1522,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl44),
  inference(cnf_transformation,[],[f1522_D])).
fof(f1522_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f810,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl14),
  inference(cnf_transformation,[],[f810_D])).
fof(f810_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1390,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1266])).
fof(f1266,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f904,f109])).
fof(f1520,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X33,fun(X33,bool)),'const.relation.EMPTY_REL_0'),s(X33,X34))),s(X33,X35)))) ) | $spl40),
  inference(cnf_transformation,[],[f1520_D])).
fof(f1520_D,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X33,fun(X33,bool)),'const.relation.EMPTY_REL_0'),s(X33,X34))),s(X33,X35)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1277,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'const.relation.EMPTY_REL_0'),s(X16,X17))),s(X16,X18)))) ) | $spl9),
  inference(superposition,[],[f904,f183])).
fof(f1516,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X44,X45),s(X44,X46)))) ) | $spl34),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X44,X45),s(X44,X46)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1278,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X19,X20),s(X19,X21)))) ) | $spl9),
  inference(superposition,[],[f904,f138])).
fof(f1380,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1295])).
fof(f1295,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f109,f904])).
fof(f1519,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl39),
  inference(cnf_transformation,[],[f1519_D])).
fof(f1519_D,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1379,plain,(
  ( ! [X2] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1297])).
fof(f1297,plain,(
  ( ! [X2] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(superposition,[],[f112,f904])).
fof(f1394,plain,(
  ( ! [X3] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X3)) ) | $spl9),
  inference(subsumption_resolution,[],[f1298,f109])).
fof(f1298,plain,(
  ( ! [X3] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl9),
  inference(superposition,[],[f115,f904])).
fof(f1515,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl32),
  inference(cnf_transformation,[],[f1515_D])).
fof(f1515_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1396,plain,(
  ( ! [X6,X5] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5)) ) | $spl9),
  inference(subsumption_resolution,[],[f1300,f457])).
fof(f1300,plain,(
  ( ! [X6,X5] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,X5) != s(bool,X6) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5)) ) | $spl9),
  inference(superposition,[],[f173,f904])).
fof(f1398,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl9),
  inference(forward_demodulation,[],[f1374,f138])).
fof(f1374,plain,(
  ( ! [X15,X16] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X16),s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1305])).
fof(f1305,plain,(
  ( ! [X15,X16] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X16),s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl9),
  inference(superposition,[],[f226,f904])).
fof(f1399,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X19)) ) | $spl9),
  inference(forward_demodulation,[],[f1372,f138])).
fof(f1372,plain,(
  ( ! [X19,X20] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,f(s(bool,'F_0'))),s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1307])).
fof(f1307,plain,(
  ( ! [X19,X20] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,f(s(bool,'F_0'))),s(bool,X20))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl9),
  inference(superposition,[],[f227,f904])).
fof(f1524,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | $spl47),
  inference(cnf_transformation,[],[f1524_D])).
fof(f1524_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f1369,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X24)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1310])).
fof(f1310,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl9),
  inference(superposition,[],[f892,f904])).
fof(f892,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f806,f109])).
fof(f1521,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (~p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X29),s(X27,X31))),s(X28,X32)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X27,fun(X28,bool)),X29),s(fun(X27,fun(X28,bool)),X30))) | p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X30),s(X27,X31))),s(X28,X32))))) ) | $spl42),
  inference(cnf_transformation,[],[f1521_D])).
fof(f1521_D,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (~p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X29),s(X27,X31))),s(X28,X32)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X27,fun(X28,bool)),X29),s(fun(X27,fun(X28,bool)),X30))) | p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X30),s(X27,X31))),s(X28,X32))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1319,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (~p(s(bool,f(s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X29),s(X27,X31))),s(X28,X32)))) | p(s(bool,happ(s(fun(X28,bool),happ(s(fun(X27,fun(X28,bool)),X30),s(X27,X31))),s(X28,X32)))) | s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(X27,fun(X28,bool)),X29),s(fun(X27,fun(X28,bool)),X30)))) ) | $spl9),
  inference(superposition,[],[f145,f904])).
fof(f1320,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),happ(s(fun(X33,fun(X33,bool)),'const.relation.EMPTY_REL_0'),s(X33,X34))),s(X33,X35)))) ) | $spl9),
  inference(superposition,[],[f183,f904])).
fof(f1517,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X41,fun(X40,bool)),X43),s(fun(X41,fun(X40,bool)),X42)))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X41,fun(X40,bool)),X42),s(X41,sK4(X42,X43,X40,X41)))),s(X40,sK5(X42,X43,X40,X41))))) ) | $spl36),
  inference(cnf_transformation,[],[f1517_D])).
fof(f1517_D,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X41,fun(X40,bool)),X43),s(fun(X41,fun(X40,bool)),X42)))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X41,fun(X40,bool)),X42),s(X41,sK4(X42,X43,X40,X41)))),s(X40,sK5(X42,X43,X40,X41))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1322,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(X41,fun(X40,bool)),X43),s(fun(X41,fun(X40,bool)),X42)))) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),happ(s(fun(X41,fun(X40,bool)),X42),s(X41,sK4(X42,X43,X40,X41)))),s(X40,sK5(X42,X43,X40,X41))))) ) | $spl9),
  inference(superposition,[],[f147,f904])).
fof(f1323,plain,(
  ( ! [X45,X46,X44] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.relation.EMPTY_REL_2'(s(X44,X45),s(X44,X46)))) ) | $spl9),
  inference(superposition,[],[f138,f904])).
fof(f1514,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | $spl9),
  inference(subsumption_resolution,[],[f1347,f111])).
fof(f1347,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl9),
  inference(superposition,[],[f116,f904])).
fof(f1511,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1491,f108])).
fof(f1491,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f193,f1037,f173])).
fof(f1037,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),X0))))) ) | $spl2),
  inference(forward_demodulation,[],[f1036,f181])).
fof(f1036,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1035,f108])).
fof(f1035,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),X0))))) ) | $spl2),
  inference(forward_demodulation,[],[f1015,f183])).
fof(f1015,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(sK0,sK4(X0,sK1,sK0,sK0)))),s(sK0,sK5(X0,sK1,sK0,sK0))))) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),X0))))) ) | $spl2),
  inference(superposition,[],[f146,f181])).
fof(f193,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f184,f116])).
fof(f184,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f180,f181])).
fof(f1509,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1507,f185])).
fof(f1507,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f1495,f193])).
fof(f1495,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1037,f112])).
fof(f1496,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f186,f1037,f112])).
fof(f186,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f184,f111])).
fof(f1494,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f193,f1037,f115])).
fof(f1492,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f108,f193,f1037,f173])).
fof(f873,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f875,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl28),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1083,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(forward_demodulation,[],[f1068,f138])).
fof(f1068,plain,(
  ( ! [X7] : (s(bool,f(s(bool,'T_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,f(s(bool,'T_0'))),s(bool,X7))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f227,f847])).
fof(f847,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1082,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(forward_demodulation,[],[f1066,f138])).
fof(f1066,plain,(
  ( ! [X5] : (s(bool,f(s(bool,'T_0'))) = s(bool,'const.relation.EMPTY_REL_2'(s(bool,X5),s(bool,f(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f226,f847])).
fof(f869,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f870,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl22),
  inference(cnf_transformation,[],[f870_D])).
fof(f870_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1065,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl18),
  inference(superposition,[],[f218,f847])).
fof(f1076,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1064,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl18),
  inference(superposition,[],[f218,f847])).
fof(f1063,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f218,f847])).
fof(f1062,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl18),
  inference(superposition,[],[f218,f847])).
fof(f1056,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(superposition,[],[f109,f847])).
fof(f1055,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl18),
  inference(superposition,[],[f847,f109])).
fof(f1054,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl18),
  inference(superposition,[],[f847,f218])).
fof(f1053,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(superposition,[],[f847,f218])).
fof(f1052,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl18),
  inference(superposition,[],[f847,f218])).
fof(f1051,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f847,f218])).
fof(f876,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f876_D])).
fof(f876_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f866,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(superposition,[],[f848,f109])).
fof(f874,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(cnf_transformation,[],[f874_D])).
fof(f874_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f865,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl19),
  inference(superposition,[],[f848,f218])).
fof(f864,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl19),
  inference(superposition,[],[f848,f218])).
fof(f871,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f863,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl19),
  inference(superposition,[],[f848,f218])).
fof(f862,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl19),
  inference(superposition,[],[f848,f218])).
fof(f846,plain,(
  ( ! [X37,X38] : (s(fun(bool,bool),X37) = s(fun(bool,bool),X38)) ) | $spl16),
  inference(cnf_transformation,[],[f846_D])).
fof(f846_D,plain,(
  ( ! [X37,X38] : (s(fun(bool,bool),X37) = s(fun(bool,bool),X38)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f848,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(cnf_transformation,[],[f848_D])).
fof(f848_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f834,plain,(
  ( ! [X37,X38] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X37) = s(fun(bool,bool),X38)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f829])).
fof(f829,plain,(
  ( ! [X37,X38] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X37) = s(fun(bool,bool),X38) | s(fun(bool,bool),X37) = s(fun(bool,bool),X38)) ) | $spl6),
  inference(superposition,[],[f142,f804])).
fof(f808,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f808_D])).
fof(f808_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f804,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f811,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f802,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK2(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f790,f109])).
fof(f790,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f778,f193])).
fof(f778,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))))) != s(X2,g(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f142,f676])).
fof(f676,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f386,f185])).
fof(f386,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f218,f193])).
fof(f809,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(cnf_transformation,[],[f809_D])).
fof(f809_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f801,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f790,f218])).
fof(f800,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f790,f218])).
fof(f803,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(cnf_transformation,[],[f803_D])).
fof(f803_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f806,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(cnf_transformation,[],[f806_D])).
fof(f806_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f799,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK2(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f790,f218])).
fof(f798,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f790,f218])).
fof(f179,plain,(
  p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) | $spl0),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f181,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') = s(fun(sK0,fun(sK0,bool)),sK1) | $spl2),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') = s(fun(sK0,fun(sK0,bool)),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f104,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') = s(fun(sK0,fun(sK0,bool)),sK1) | p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  (p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') = s(fun(sK0,fun(sK0,bool)),sK1)) & (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f73])).
fof(f73,plain,(
  ? [X0,X1] : ((p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1)) & (~p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) | s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') != s(fun(X0,fun(X0,bool)),X1)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) <~> s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0') = s(fun(X0,fun(X0,bool)),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0') = s(fun(X6,fun(X6,bool)),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.RSUBSET_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0')))) <=> s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0') = s(fun(X6,fun(X6,bool)),X7))),
  file('thm.relation.REMPTY_SUBSET|split|1','thm.relation.REMPTY_SUBSET|split|1')).
fof(f180,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) | $spl1),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f182,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),sK1) | $spl3),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f105,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0') != s(fun(sK0,fun(sK0,bool)),sK1) | ~p(s(bool,'const.relation.RSUBSET_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(cnf_transformation,[],[f74])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 217
% Generated clauses: 3483
% Final active clauses: 48
% Final passive clauses: 75
% Input formulas: 25
% Initial clauses: 58
% 
% Duplicate literals: 252
% Trivial inequalities: 28
% Fw subsumption resolutions: 182
% Fw demodulations: 368
% Bw demodulations: 63
% Condensations: 215
% Global subsumptions: 5
% 
% Simple tautologies: 85
% Equational tautologies: 51
% Forward subsumptions: 2018
% Fw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 179
% Forward superposition: 1056
% Backward superposition: 883
% Self superposition: 52
% Equality factoring: 34
% Equality resolution: 3
% 
% Split clauses: 40
% Split components: 86
% SAT solver clauses: 324
% SAT solver unit clauses: 11
% SAT solver binary clauses: 81
% 
% Sat splits: 40
% Sat splitting refutations: 20
% 
% Memory used [KB]: 2430
% Time elapsed: 0.098 s
% ------------------------------
% ------------------------------
% Success in time 1.427 s
