Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (64586)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 1872
% Generated clauses: 21844
% Final active clauses: 120
% Final passive clauses: 1663
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 1512
% Trivial inequalities: 56
% Fw subsumption resolutions: 2078
% Fw demodulations: 3377
% Bw demodulations: 2
% 
% Simple tautologies: 203
% Equational tautologies: 19
% Forward subsumptions: 11461
% 
% Binary resolution: 11
% Forward superposition: 8431
% Backward superposition: 6008
% Self superposition: 169
% Equality factoring: 36
% Equality resolution: 5
% 
% Split clauses: 773
% Split components: 1551
% SAT solver clauses: 1256
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 773
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1256
% TWLsolver calls for satisfiability: 84
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f5173,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5153,f68,f92,f4835,f4937,f76,f100,f2174,f2175,f72,f98,f82,f106,f83,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) ) | $spl26),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|0','ARITY_1_3_const.ringNorm.r_interp_sp')).
fof(f106,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) ) | $spl24),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|0','ARITY_1_3_const.canonical.interp_sp')).
fof(f98,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) | $spl8),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_sp_def|split|0','thm.ringNorm.r_interp_sp_def')).
fof(f2175,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3)))) ) | $spl318),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3)))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f2174,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3)))) ) | ($spl8 | $spl24 | $spl26)),
  inference(forward_demodulation,[],[f2094,f107])).
fof(f2094,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) ) | ($spl8 | $spl24)),
  inference(superposition,[],[f106,f98])).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))) ) | $spl12),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X8,X9] : s(X6,X9) = s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9)))))),
  file('thm.ringNorm.interp_sp_def|split|0','thm.canonical.interp_sp_def|split|0')).
fof(f4937,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X7) = s(X4,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X4),X5),s('type.quote.varmap'(X4),X6),s('type.canonical.spolynom'(X4),'const.canonical.SPconst_1'(s(X4,X7)))))) ) | $spl802),
  inference(cnf_transformation,[],[f4937_D])).
fof(f4937_D,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X7) = s(X4,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X4),X5),s('type.quote.varmap'(X4),X6),s('type.canonical.spolynom'(X4),'const.canonical.SPconst_1'(s(X4,X7)))))) ) <=> ~$spl802),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl802])])).
fof(f4835,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X7) = s(X4,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X4),X5),s('type.quote.varmap'(X4),X6),s('type.canonical.spolynom'(X4),'const.canonical.SPconst_1'(s(X4,X7)))))) ) | ($spl12 | $spl318)),
  inference(superposition,[],[f2175,f100])).
fof(f92,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))))) | $spl1),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f68,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3)))))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  s(sK0,sK3) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3] : s(X0,X3) != s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9)))))),
  file('thm.ringNorm.interp_sp_def|split|0','thm.ringNorm.interp_sp_def|split|0')).
fof(f5153,plain,(
  $false | ($spl1 | $spl802)),
  inference(trivial_inequality_removal,[],[f5152])).
fof(f5152,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ($spl1 | $spl802)),
  inference(superposition,[],[f92,f4937])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 60
% Passive clauses: 414
% Generated clauses: 5052
% Final active clauses: 56
% Final passive clauses: 343
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 269
% Trivial inequalities: 13
% Fw subsumption resolutions: 142
% Fw demodulations: 361
% Bw demodulations: 11
% Condensations: 322
% 
% Simple tautologies: 59
% Equational tautologies: 16
% Forward subsumptions: 2811
% 
% Binary resolution: 7
% Forward superposition: 1877
% Backward superposition: 1472
% Self superposition: 102
% Equality factoring: 30
% Equality resolution: 1
% 
% Split clauses: 12
% Split components: 29
% SAT solver clauses: 639
% SAT solver unit clauses: 40
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 1
% 
% Sat splits: 12
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 638
% TWLsolver calls for satisfiability: 82
% 
% Memory used [KB]: 7164
% Time elapsed: 0.271 s
% ------------------------------
% ------------------------------
% Success in time 0.87 s
