Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.bool
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bool
% SZS output start Proof for thm.bool
fof(f100,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f96,f44,f71,f69,f80,f43,f70,f68,f49,f74])).
fof(f74,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f49,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bool.EQ_CLAUSES|split|1',aHL_FALSITY)).
fof(f68,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl0),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s(bool,'T_0') = s(bool,sK0) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f70,plain,(
  p(s(bool,sK0)) | $spl2),
  inference(cnf_transformation,[],[f70_D])).
fof(f70_D,plain,(
  p(s(bool,sK0)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f43,plain,(
  p(s(bool,sK0)) | s(bool,'T_0') = s(bool,sK0)),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  (s(bool,'T_0') = s(bool,sK0) | p(s(bool,sK0))) & (s(bool,'T_0') != s(bool,sK0) | ~p(s(bool,sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f34])).
fof(f34,plain,(
  ? [X0] : ((s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) & (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ? [X0] : (s(bool,'T_0') = s(bool,X0) <~> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s(bool,'T_0') = s(bool,X6) <=> p(s(bool,X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s(bool,'T_0') = s(bool,X6) <=> p(s(bool,X6)))),
  file('thm.bool.EQ_CLAUSES|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f80,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f79,f45])).
fof(f45,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bool.EQ_CLAUSES|split|1',aHL_TRUTH)).
fof(f79,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f68,f71])).
fof(f69,plain,(
  s(bool,'T_0') != s(bool,sK0) | $spl1),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s(bool,'T_0') != s(bool,sK0) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f71,plain,(
  ~p(s(bool,sK0)) | $spl3),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  ~p(s(bool,sK0)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f44,plain,(
  ~p(s(bool,sK0)) | s(bool,'T_0') != s(bool,sK0)),
  inference(cnf_transformation,[],[f35])).
fof(f96,plain,(
  $false | ($spl1 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f95,f69])).
fof(f95,plain,(
  s(bool,'T_0') = s(bool,sK0) | ($spl2 | $spl7)),
  inference(subsumption_resolution,[],[f90,f74])).
fof(f90,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK0) | $spl2),
  inference(superposition,[],[f70,f50])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bool.EQ_CLAUSES|split|1',aHL_BOOL_CASES)).
% SZS output end Proof for thm.bool
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 9
% Passive clauses: 22
% Generated clauses: 49
% Final active clauses: 7
% Final passive clauses: 12
% Input formulas: 12
% Initial clauses: 24
% 
% Duplicate literals: 1
% Fw subsumption resolutions: 4
% Bw demodulations: 1
% 
% Simple tautologies: 6
% Equational tautologies: 2
% Forward subsumptions: 5
% 
% Binary resolution: 2
% Backward superposition: 6
% Self superposition: 2
% Equality factoring: 2
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 12
% SAT solver unit clauses: 2
% SAT solver binary clauses: 8
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 7
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 9
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 639
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
% Success in time 0.017 s
