Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f13605,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f13594,f132,f193])).
fof(f193,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f132,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1',aHL_FALSITY)).
fof(f13594,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f13593,f193])).
fof(f13593,plain,(
  p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f13592,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X17] : ! [X18] : s(X17,X18) = s(X17,happ(s(fun(X17,X17),'_i_0_0'),s(X17,X18)))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1','_i0_ax')).
fof(f13592,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | $spl3),
  inference(subsumption_resolution,[],[f13524,f193])).
fof(f13524,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))))),
  inference(superposition,[],[f3738,f236])).
fof(f236,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun('type.one.one',bool),'_dst_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))),
  inference(resolution,[],[f125,f139])).
fof(f139,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f125,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun('type.one.one',bool),'_dst_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun('type.one.one',bool),'_dst_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun('type.one.one',bool),'_dst_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun('type.one.one',bool),'_dst_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1','thm.quantHeuristics.GUESS_RULES_BOOL|split|1')).
fof(f3738,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X28,X27),'_dst_x0_1'(s(X27,X29))),s(fun(X27,bool),X30)))) | p(s(bool,happ(s(fun(X27,bool),X30),s(X27,X29))))) )),
  inference(superposition,[],[f188,f171])).
fof(f171,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X17,X19] : ! [X20,X21] : s(X19,X20) = s(X19,happ(s(fun(X17,X19),'_dst_x0_1'(s(X19,X20))),s(X17,X21)))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1','_dst_x0_ax')).
fof(f188,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK11(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9,X6] : ! [X15,X16] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X9),X15),s(fun(X9,bool),X16)))) <=> ! [X14] : ~p(s(bool,happ(s(fun(X9,bool),X16),s(X9,happ(s(fun(X6,X9),X15),s(X6,X14)))))))),
  file('thm.quantHeuristics.GUESS_RULES_BOOL|split|1','thm.quantHeuristics.GUESS_REWRITES|split|3')).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 209
% Passive clauses: 3059
% Generated clauses: 13393
% Final active clauses: 194
% Final passive clauses: 2674
% Input formulas: 30
% Initial clauses: 64
% 
% Duplicate literals: 226
% Trivial inequalities: 395
% Fw subsumption resolutions: 1762
% Fw demodulations: 1773
% Bw demodulations: 13
% 
% Simple tautologies: 306
% Equational tautologies: 18
% Forward subsumptions: 5412
% 
% Binary resolution: 937
% Forward superposition: 4921
% Backward superposition: 3096
% Self superposition: 37
% Equality factoring: 32
% Equality resolution: 20
% 
% Split clauses: 233
% Split components: 466
% SAT solver clauses: 604
% SAT solver unit clauses: 17
% SAT solver binary clauses: 186
% SAT solver learnt clauses: 1
% 
% Sat splits: 233
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 598
% TWLsolver calls for satisfiability: 155
% 
% Memory used [KB]: 8059
% Time elapsed: 0.286 s
% ------------------------------
% ------------------------------
% Success in time 0.289 s
