Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.canonical
% (82296)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 142
% Passive clauses: 3111
% Generated clauses: 19621
% Final active clauses: 117
% Final passive clauses: 2534
% Input formulas: 69
% Initial clauses: 126
% 
% Duplicate literals: 508
% Trivial inequalities: 287
% Fw subsumption resolutions: 2188
% Fw demodulations: 2120
% Bw demodulations: 8
% 
% Simple tautologies: 273
% Equational tautologies: 47
% Forward subsumptions: 10233
% 
% Binary resolution: 341
% Factoring: 1
% Forward superposition: 7477
% Backward superposition: 6425
% Self superposition: 228
% Equality factoring: 96
% Equality resolution: 40
% 
% Split clauses: 682
% Split components: 1447
% SAT solver clauses: 1014
% SAT solver unit clauses: 16
% SAT solver binary clauses: 171
% 
% Sat splits: 682
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 986
% TWLsolver calls for satisfiability: 142
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.canonical
% (82318)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 123
% Passive clauses: 1411
% Generated clauses: 13468
% Final active clauses: 123
% Final passive clauses: 1286
% Input formulas: 69
% Initial clauses: 126
% 
% Duplicate literals: 344
% Trivial inequalities: 83
% Fw subsumption resolutions: 390
% Fw demodulations: 952
% Bw demodulations: 2
% Condensations: 1135
% Global subsumptions: 12
% 
% Simple tautologies: 136
% Equational tautologies: 42
% Forward subsumptions: 6819
% 
% Binary resolution: 142
% Factoring: 1
% Forward superposition: 4335
% Backward superposition: 4260
% Self superposition: 180
% Equality factoring: 72
% Equality resolution: 18
% 
% Split clauses: 111
% Split components: 222
% SAT solver clauses: 2220
% SAT solver unit clauses: 159
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 24
% 
% Sat splits: 110
% 
% TWLsolver clauses: 2172
% TWLsolver calls for satisfiability: 429
% 
% Memory used [KB]: 26993
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.canonical
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.canonical
% SZS output start Proof for thm.canonical
fof(f166787,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f246,f166782])).
fof(f166782,plain,(
  ( ! [X378] : (~p(s(bool,X378))) )),
  inference(forward_demodulation,[],[f166781,f340])).
fof(f340,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1))),X3),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1)),X4)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X2) = s(X1,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1))),X3),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1)),X4)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X2) = s(X1,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s(X1,X2),s(fun(X0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1))),X3),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),X1)),X4)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X13] : ! [X14,X15,X16] : s(X13,X14) = s(X13,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s(X13,X14),s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),X13))),X15),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),X13)),X16)))),
  file('thm.canonical.canonical_sum_distinct|split|1','thm.canonical.canonical_sum_case_def|split|0')).
fof(f166781,plain,(
  ( ! [X379,X380,X378] : (~p(s(bool,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s(bool,X378),s(fun(sK0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),bool))),X379),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),bool)),X380))))) )),
  inference(subsumption_resolution,[],[f166648,f166601])).
fof(f166601,plain,(
  ( ! [X4,X7,X5] : (s(X4,X5) = s(X4,happ(s(fun('type.canonical.canonical_sum'(sK0),X4),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),X4)),X7),s('type.list.list'('type.quote.index'),sK2))),s('type.canonical.canonical_sum'(sK0),sK1)))) )),
  inference(superposition,[],[f88857,f340])).
fof(f88857,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.canonical.canonical_sum'(sK0),X0),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),X0)),X1),s('type.list.list'('type.quote.index'),sK2))),s('type.canonical.canonical_sum'(sK0),sK1))) = s(X0,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s(X0,X2),s(fun(sK0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),X0))),X3),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),X0)),X1)))) )),
  inference(superposition,[],[f341,f245])).
fof(f245,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK1)))),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f180])).
fof(f180,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X7)))),
  file('thm.canonical.canonical_sum_distinct|split|1','thm.canonical.canonical_sum_distinct|split|1')).
fof(f341,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.canonical.canonical_sum'(X1),X0),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X1),X3))) = s(X0,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X1),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0))),X5),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.canonical.canonical_sum'(X1),X0),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X1),X3))) = s(X0,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X1),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0))),X5),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.canonical.canonical_sum'(X1),X0),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X1),X3))) = s(X0,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X1),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0))),X5),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X1),X0)),X6)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X6] : ! [X8,X7,X14,X15,X16] : s(X13,happ(s(fun('type.canonical.canonical_sum'(X6),X13),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),X13)),X16),s('type.list.list'('type.quote.index'),X8))),s('type.canonical.canonical_sum'(X6),X7))) = s(X13,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X7))),s(X13,X14),s(fun(X6,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),X13))),X15),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),X13)),X16)))),
  file('thm.canonical.canonical_sum_distinct|split|1','thm.canonical.canonical_sum_case_def|split|2')).
fof(f166648,plain,(
  ( ! [X379,X380,X378] : (s(bool,'F_0') != s(bool,happ(s(fun('type.canonical.canonical_sum'(sK0),bool),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),bool)),X380),s('type.list.list'('type.quote.index'),sK2))),s('type.canonical.canonical_sum'(sK0),sK1))) | ~p(s(bool,'const.canonical.canonical_sum_CASE_4'(s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s(bool,X378),s(fun(sK0,fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),bool))),X379),s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(sK0),bool)),X380))))) )),
  inference(superposition,[],[f261,f88857])).
fof(f261,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X20] : (s(bool,'F_0') = s(bool,X20) <=> ~p(s(bool,X20)))),
  file('thm.canonical.canonical_sum_distinct|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f246,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.canonical.canonical_sum_distinct|split|1','thm.bool.TRUTH')).
% SZS output end Proof for thm.canonical
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 192
% Passive clauses: 1868
% Generated clauses: 166392
% Final active clauses: 150
% Final passive clauses: 611
% Input formulas: 69
% Initial clauses: 126
% Discarded non-redundant clauses: 44802
% 
% Duplicate literals: 640
% Trivial inequalities: 134
% Fw subsumption resolutions: 1353
% Fw demodulations: 71634
% Bw demodulations: 248
% Condensations: 1509
% Global subsumptions: 7
% 
% Simple tautologies: 226
% Equational tautologies: 152
% Forward subsumptions: 20387
% Fw demodulations to eq. taut.: 25899
% Bw demodulations to eq. taut.: 16
% 
% Unit resulting resolution: 77861
% Factoring: 1
% Forward superposition: 6808
% Backward superposition: 5664
% Self superposition: 160
% Equality factoring: 63
% Equality resolution: 44
% 
% Split clauses: 47
% Split components: 94
% SAT solver clauses: 2007
% SAT solver unit clauses: 572
% SAT solver binary clauses: 596
% 
% Sat splits: 47
% Sat splitting refutations: 36
% 
% Memory used [KB]: 82130
% Time elapsed: 7.744 s
% ------------------------------
% ------------------------------
% Success in time 9.08 s
