Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.canonical
% (72867)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 4277
% Generated clauses: 14491
% Final active clauses: 112
% Final passive clauses: 3862
% Input formulas: 68
% Initial clauses: 123
% 
% Duplicate literals: 523
% Trivial inequalities: 228
% Fw subsumption resolutions: 1522
% Fw demodulations: 473
% Bw demodulations: 5
% 
% Simple tautologies: 300
% Equational tautologies: 86
% Forward subsumptions: 6178
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 5065
% Backward superposition: 5977
% Self superposition: 306
% Equality factoring: 190
% Equality resolution: 33
% 
% Split clauses: 463
% Split components: 972
% SAT solver clauses: 1069
% SAT solver unit clauses: 38
% SAT solver binary clauses: 388
% 
% Sat splits: 463
% 
% TWLsolver clauses: 1043
% TWLsolver calls for satisfiability: 142
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.canonical
% (73049)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 109
% Passive clauses: 1058
% Generated clauses: 12245
% Final active clauses: 107
% Final passive clauses: 911
% Input formulas: 68
% Initial clauses: 123
% 
% Duplicate literals: 443
% Trivial inequalities: 222
% Fw subsumption resolutions: 698
% Fw demodulations: 827
% Bw demodulations: 6
% Condensations: 1205
% Global subsumptions: 12
% 
% Simple tautologies: 181
% Equational tautologies: 48
% Forward subsumptions: 5514
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 32
% 
% Binary resolution: 100
% Factoring: 1
% Forward superposition: 3834
% Backward superposition: 3437
% Self superposition: 173
% Equality factoring: 80
% Equality resolution: 30
% 
% Split clauses: 50
% Split components: 101
% SAT solver clauses: 2076
% SAT solver unit clauses: 156
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 24
% 
% Sat splits: 50
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2028
% TWLsolver calls for satisfiability: 388
% 
% Memory used [KB]: 17142
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.canonical
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.canonical
% SZS output start Proof for thm.canonical
fof(f170449,plain,(
  $false),
  inference(subsumption_resolution,[],[f170448,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X19,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X19)))),
  file('thm.canonical.spolynom_distinct|split|1','thm.canonical.spolynom_distinct|split|0')).
fof(f170448,plain,(
  s('type.canonical.spolynom'('type.quote.index'),'const.canonical.SPvar_1'(s('type.quote.index',sK3))) = s('type.canonical.spolynom'('type.quote.index'),'const.canonical.SPconst_1'(s('type.quote.index',sK3)))),
  inference(forward_demodulation,[],[f169193,f169192])).
fof(f169192,plain,(
  ( ! [X142,X141] : (s('type.canonical.spolynom'(X141),'const.canonical.SPvar_1'(s('type.quote.index',sK3))) = s('type.canonical.spolynom'(X141),happ(s(fun('type.quote.index','type.canonical.spolynom'(X141)),X142),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f164532,f293])).
fof(f293,plain,(
  ( ! [X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X1))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.quote.index','type.canonical.spolynom'(X0)),'const.canonical.SPvar_0'),s('type.quote.index',X1)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X1))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.quote.index','type.canonical.spolynom'(X0)),'const.canonical.SPvar_0'),s('type.quote.index',X1)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X1))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.quote.index','type.canonical.spolynom'(X0)),'const.canonical.SPvar_0'),s('type.quote.index',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X12] : s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X12))) = s('type.canonical.spolynom'(X6),happ(s(fun('type.quote.index','type.canonical.spolynom'(X6)),'const.canonical.SPvar_0'),s('type.quote.index',X12)))),
  file('thm.canonical.spolynom_distinct|split|1','ARITY_0_1_const.canonical.SPvar')).
fof(f164532,plain,(
  ( ! [X2,X0,X3] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',sK3))) = s(X0,happ(s(fun('type.quote.index',X0),X3),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f26494,f26494])).
fof(f26494,plain,(
  ( ! [X6,X8,X5] : (s(X5,happ(s(fun('type.quote.index',X5),X6),s('type.quote.index',sK3))) = s(X5,happ(s(fun('type.canonical.spolynom'(sK0),X5),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X5)),X8),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK1)))) )),
  inference(superposition,[],[f10366,f335])).
fof(f335,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X5),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X6)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X5),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X6)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X5),s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),X1)),X6)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X14] : ! [X9,X15,X16,X17,X18] : s(X14,happ(s(fun('type.quote.index',X14),X15),s('type.quote.index',X9))) = s(X14,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X14)),X17),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X14)),X18)))),
  file('thm.canonical.spolynom_distinct|split|1','thm.canonical.spolynom_case_def|split|0')).
fof(f10366,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(sK0),X0),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK1))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(sK0),'const.canonical.SPvar_1'(s('type.quote.index',sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X4)))) )),
  inference(superposition,[],[f337,f241])).
fof(f241,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPvar_1'(s('type.quote.index',sK3))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPvar_1'(s('type.quote.index',sK3))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f181])).
fof(f181,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7)))),
  file('thm.canonical.spolynom_distinct|split|1','thm.canonical.spolynom_distinct|split|1')).
fof(f337,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X6] : ! [X8,X7,X15,X16,X17,X18] : s(X14,happ(s(fun('type.canonical.spolynom'(X6),X14),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X14)),X17),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.spolynom'(X6),X7))) = s(X14,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7))),s(fun('type.quote.index',X14),X15),s(fun(X6,X14),X16),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X14)),X17),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X14)),X18)))),
  file('thm.canonical.spolynom_distinct|split|1','thm.canonical.spolynom_case_def|split|2')).
fof(f169193,plain,(
  ( ! [X143] : (s('type.canonical.spolynom'('type.quote.index'),'const.canonical.SPconst_1'(s('type.quote.index',sK3))) = s('type.canonical.spolynom'('type.quote.index'),happ(s(fun('type.quote.index','type.canonical.spolynom'('type.quote.index')),X143),s('type.quote.index',sK3)))) )),
  inference(superposition,[],[f164532,f292])).
fof(f292,plain,(
  ( ! [X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1))) = s('type.canonical.spolynom'(X0),happ(s(fun(X0,'type.canonical.spolynom'(X0)),'const.canonical.SPconst_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1))) = s('type.canonical.spolynom'(X0),happ(s(fun(X0,'type.canonical.spolynom'(X0)),'const.canonical.SPconst_0'),s(X0,X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X1))) = s('type.canonical.spolynom'(X0),happ(s(fun(X0,'type.canonical.spolynom'(X0)),'const.canonical.SPconst_0'),s(X0,X1)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X12] : s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X12))) = s('type.canonical.spolynom'(X6),happ(s(fun(X6,'type.canonical.spolynom'(X6)),'const.canonical.SPconst_0'),s(X6,X12)))),
  file('thm.canonical.spolynom_distinct|split|1','ARITY_0_1_const.canonical.SPconst')).
% SZS output end Proof for thm.canonical
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 639
% Passive clauses: 11677
% Generated clauses: 175273
% Final active clauses: 235
% Final passive clauses: 1539
% Input formulas: 68
% Initial clauses: 123
% Discarded non-redundant clauses: 6320
% 
% Duplicate literals: 2611
% Trivial inequalities: 717
% Fw subsumption resolutions: 9725
% Fw demodulations: 46288
% Bw demodulations: 5647
% Condensations: 4721
% Global subsumptions: 242
% 
% Simple tautologies: 535
% Equational tautologies: 1435
% Forward subsumptions: 88589
% Fw demodulations to eq. taut.: 6153
% Bw demodulations to eq. taut.: 798
% 
% Unit resulting resolution: 23793
% Factoring: 21
% Forward superposition: 44946
% Backward superposition: 29180
% Self superposition: 461
% Equality factoring: 116
% Equality resolution: 72
% 
% Split clauses: 85
% Split components: 175
% SAT solver clauses: 12491
% SAT solver unit clauses: 566
% SAT solver binary clauses: 4029
% 
% Sat splits: 85
% Sat splitting refutations: 400
% 
% Memory used [KB]: 118462
% Time elapsed: 9.245 s
% ------------------------------
% ------------------------------
% Success in time 10.57 s
