Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (114322)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 220
% Passive clauses: 3390
% Generated clauses: 17699
% Final active clauses: 199
% Final passive clauses: 2967
% Input formulas: 41
% Initial clauses: 102
% 
% Duplicate literals: 634
% Trivial inequalities: 376
% Fw subsumption resolutions: 2523
% Fw demodulations: 2347
% Bw demodulations: 20
% 
% Simple tautologies: 394
% Equational tautologies: 30
% Forward subsumptions: 6960
% 
% Binary resolution: 1153
% Forward superposition: 6011
% Backward superposition: 7172
% Self superposition: 160
% Equality factoring: 43
% Equality resolution: 63
% 
% Split clauses: 685
% Split components: 1385
% SAT solver clauses: 1246
% SAT solver unit clauses: 15
% SAT solver binary clauses: 138
% 
% Sat splits: 685
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1243
% TWLsolver calls for satisfiability: 231
% 
% Memory used [KB]: 15479
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f1675,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1674,f163,f283,f167,f291,f168,f292,f223,f310,f1655,f1657,f164,f288,f328,f331,f161,f286,f171,f294,f1632,f1633,f1314,f1369,f224,f312,f1268,f1295,f804,f838,f236,f317,f1604,f1605,f162,f284,f227,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK6(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl30),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK6(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f227,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK6(X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK6(X3,X2,X1,X0))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f100])).
fof(f100,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X21] : ((p(s(bool,'const.bool.IN_2'(s(X6,X21),s(fun(X6,bool),X7)))) & ! [X22] : (p(s(bool,'const.bool.IN_2'(s(X6,X22),s(fun(X6,bool),X7)))) => s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X21))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X22))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X21)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME')).
fof(f284,plain,(
  ( ! [X4] : (s('type.realax.real',sK3) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) ) | $spl2),
  inference(cnf_transformation,[],[f284_D])).
fof(f284_D,plain,(
  ( ! [X4] : (s('type.realax.real',sK3) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f162,plain,(
  ( ! [X4] : (s('type.realax.real',sK3) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) & (! [X4] : s('type.realax.real',sK3) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f93])).
fof(f93,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2,X3] : (! [X4] : s('type.realax.real',X3) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.realax.real',X3) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.realax.real',X3) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : s('type.realax.real',X9) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : (! [X10] : s('type.realax.real',X9) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',X9)))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST')).
fof(f1605,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl334),
  inference(cnf_transformation,[],[f1605_D])).
fof(f1605_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f1604,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f1603,f284])).
fof(f1603,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f1599,f284])).
fof(f1599,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK3) != s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f313,f284])).
fof(f317,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f236,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X9,X10,X17] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X10),s(fun(X6,bool),X17)))))) <=> (s(X6,X9) = s(X6,X10) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X17))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.IN_INSERT')).
fof(f838,plain,(
  ( ! [X14,X12,X13,X11] : (s(X11,X12) != s(X11,X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X11,X12),s(fun(X11,bool),'const.pred_set.INSERT_2'(s(X11,X13),s(fun(X11,bool),X14)))))) ) | $spl116),
  inference(cnf_transformation,[],[f838_D])).
fof(f838_D,plain,(
  ( ! [X14,X12,X13,X11] : (s(X11,X12) != s(X11,X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X11,X12),s(fun(X11,bool),'const.pred_set.INSERT_2'(s(X11,X13),s(fun(X11,bool),X14)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f804,plain,(
  ( ! [X14,X12,X13,X11] : (s(X11,X12) != s(X11,X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X11,X12),s(fun(X11,bool),'const.pred_set.INSERT_2'(s(X11,X13),s(fun(X11,bool),X14)))))) ) | ($spl16 | $spl38)),
  inference(resolution,[],[f317,f294])).
fof(f1295,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) | $spl264),
  inference(cnf_transformation,[],[f1295_D])).
fof(f1295_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f1268,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) ) | $spl116),
  inference(equality_resolution,[],[f838])).
fof(f312,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl28),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f224,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),sK5(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X17] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X17) | ? [X9,X11] : (s(fun(X6,bool),X17) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X11))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X11))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.SET_CASES')).
fof(f1369,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl268),
  inference(cnf_transformation,[],[f1369_D])).
fof(f1369_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f1314,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl28 | $spl264)),
  inference(superposition,[],[f1295,f312])).
fof(f1633,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | $spl352),
  inference(cnf_transformation,[],[f1633_D])).
fof(f1633_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f1632,plain,(
  ( ! [X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | ($spl6 | $spl268 | $spl334)),
  inference(subsumption_resolution,[],[f1616,f288])).
fof(f1616,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))),s('type.realax.real',sK3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1)) ) | ($spl268 | $spl334)),
  inference(superposition,[],[f1605,f1369])).
fof(f294,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.bool.EQ_CLAUSES|split|0')).
fof(f286,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f286_D])).
fof(f286_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f161,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f107])).
fof(f331,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl60),
  inference(cnf_transformation,[],[f331_D])).
fof(f331_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f328,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | ($spl4 | $spl16)),
  inference(resolution,[],[f294,f286])).
fof(f288,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f164,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST',aHL_TRUTH)).
fof(f1657,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl366),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f1655,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl1 | $spl6 | $spl60 | $spl352)),
  inference(subsumption_resolution,[],[f1654,f283])).
fof(f1654,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl6 | $spl60 | $spl352)),
  inference(subsumption_resolution,[],[f1643,f288])).
fof(f1643,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl60 | $spl352)),
  inference(superposition,[],[f1633,f331])).
fof(f310,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl24),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f223,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X8] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.real_sigma.REAL_SUM_IMAGE_THM|split|0')).
fof(f292,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | $spl12),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f168,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X9] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.real.REAL_MUL_LZERO')).
fof(f291,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f167,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST','thm.pred_set.CARD_DEF|split|0')).
fof(f283,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3))) | $spl1),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f163,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',sK3)))),
  inference(cnf_transformation,[],[f107])).
fof(f1674,plain,(
  $false | ($spl1 | $spl10 | $spl12 | $spl24 | $spl366)),
  inference(subsumption_resolution,[],[f1673,f292])).
fof(f1673,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK3))) | ($spl1 | $spl10 | $spl24 | $spl366)),
  inference(forward_demodulation,[],[f1672,f310])).
fof(f1672,plain,(
  s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK3))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl10 | $spl366)),
  inference(forward_demodulation,[],[f1665,f291])).
fof(f1665,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),s('type.realax.real',sK3))) | ($spl1 | $spl366)),
  inference(backward_demodulation,[],[f1657,f283])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 58
% Passive clauses: 174
% Generated clauses: 1474
% Final active clauses: 51
% Final passive clauses: 115
% Input formulas: 41
% Initial clauses: 102
% 
% Duplicate literals: 92
% Trivial inequalities: 32
% Fw subsumption resolutions: 80
% Fw demodulations: 80
% Bw demodulations: 8
% Condensations: 128
% Global subsumptions: 12
% 
% Simple tautologies: 97
% Equational tautologies: 20
% Forward subsumptions: 453
% 
% Binary resolution: 53
% Forward superposition: 428
% Backward superposition: 209
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 422
% SAT solver unit clauses: 110
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 25
% 
% Sat splits: 4
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 373
% TWLsolver calls for satisfiability: 402
% 
% Memory used [KB]: 2430
% Time elapsed: 0.066 s
% ------------------------------
% ------------------------------
% Success in time 0.698 s
