Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.topology
% (73172)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1701
% Generated clauses: 14727
% Final active clauses: 68
% Final passive clauses: 1501
% Input formulas: 12
% Initial clauses: 28
% 
% Duplicate literals: 1194
% Trivial inequalities: 183
% Fw subsumption resolutions: 1845
% Fw demodulations: 743
% Bw demodulations: 18
% 
% Simple tautologies: 318
% Equational tautologies: 109
% Forward subsumptions: 6734
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 41
% Forward superposition: 5640
% Backward superposition: 4293
% Self superposition: 370
% Equality factoring: 305
% Equality resolution: 15
% 
% Split clauses: 712
% Split components: 1445
% SAT solver clauses: 2019
% SAT solver unit clauses: 323
% SAT solver binary clauses: 683
% 
% Sat splits: 712
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2001
% TWLsolver calls for satisfiability: 60
% 
% Memory used [KB]: 7291
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.topology
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.topology
% SZS output start Proof for thm.topology
fof(f3896,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3895,f51,f84,f70,f102,f2454,f2464,f107,f115,f53,f86,f54,f87,f712,f713,f55,f88,f59,f89,f3833,f3835,f116,f118,f49,f80,f119,f121,f48,f82,f432,f439,f437,f384,f385,f50,f78,f105,f114,f460,f1012,f1013,f62,f98,f72,f103])).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) ) | $spl42),
  inference(cnf_transformation,[],[f103_D])).
fof(f103_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f72,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X14,X15] : s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X14),s(fun(X6,bool),X15))) = s(bool,happ(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X14))),s(fun(X6,bool),X15)))),
  file('thm.topology.TOPOLOGY|split|2','ARITY_1_2_const.topology.open')).
fof(f98,plain,(
  ( ! [X6,X0,X7,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7)))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~sP0(X0,X1)) ) | $spl32),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  ( ! [X6,X0,X7,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7)))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~sP0(X0,X1)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f62,plain,(
  ( ! [X6,X0,X7,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7)))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X6,X7] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7))))))) & ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),sK5(X1,X0))))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),sK6(X1,X0))))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(X0,bool),sK6(X1,X0)))))))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),sK7(X1,X0)))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X6,X7] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X6)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X7)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X6),s(fun(X0,bool),X7))))))) & ! [X5] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X5),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X5))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))) & (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))) | sP0(X0,X1)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))) & ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ? [X3,X4] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) | ? [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) & ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1013,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),'const.topology.re_intersect_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7)))))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X7)))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X6)))) | ~sP0(X4,'const.topology.open_1'(s('type.topology.topology'(X4),X5)))) ) | $spl162),
  inference(cnf_transformation,[],[f1013_D])).
fof(f1013_D,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),'const.topology.re_intersect_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7)))))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X7)))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X6)))) | ~sP0(X4,'const.topology.open_1'(s('type.topology.topology'(X4),X5)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1012,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X6)))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X7)))) | p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),'const.topology.re_intersect_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7)))))) | ~sP0(X4,'const.topology.open_1'(s('type.topology.topology'(X4),X5)))) ) | ($spl32 | $spl42)),
  inference(forward_demodulation,[],[f1011,f103])).
fof(f1011,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),X7)))) | p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),'const.topology.re_intersect_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7)))))) | ~p(s(bool,happ(s(fun(fun(X4,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X4),X5))),s(fun(X4,bool),X6)))) | ~sP0(X4,'const.topology.open_1'(s('type.topology.topology'(X4),X5)))) ) | ($spl32 | $spl42)),
  inference(forward_demodulation,[],[f1002,f103])).
fof(f1002,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X4),X5),s(fun(X4,bool),'const.topology.re_intersect_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7)))))) | ~p(s(bool,happ(s(fun(fun(X4,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X4),X5))),s(fun(X4,bool),X7)))) | ~p(s(bool,happ(s(fun(fun(X4,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X4),X5))),s(fun(X4,bool),X6)))) | ~sP0(X4,'const.topology.open_1'(s('type.topology.topology'(X4),X5)))) ) | ($spl32 | $spl42)),
  inference(superposition,[],[f98,f103])).
fof(f460,plain,(
  $false | ($spl6 | $spl9 | $spl70)),
  inference(subsumption_resolution,[],[f443,f84])).
fof(f443,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl70)),
  inference(backward_demodulation,[],[f437,f86])).
fof(f114,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f105,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f87,f87])).
fof(f78,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) | $spl1),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f50,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4)))) & ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f32])).
fof(f32,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X3))))) & ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X3))))) => p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),X3))))) => p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X0),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),X9))))) => p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),'const.topology.re_intersect_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),X9))))) => p(s(bool,'const.topology.open_2'(s('type.topology.topology'(X6),X7),s(fun(X6,bool),'const.topology.re_intersect_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))))))),
  file('thm.topology.TOPOLOGY|split|2','thm.topology.TOPOLOGY|split|2')).
fof(f385,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl68),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f384,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f252,f84])).
fof(f252,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl1 | $spl46)),
  inference(superposition,[],[f78,f114])).
fof(f437,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) | $spl70),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f439,plain,(
  s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl72),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f432,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl10 | $spl68)),
  inference(duplicate_literal_removal,[],[f391])).
fof(f391,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),'const.topology.re_intersect_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl10 | $spl68)),
  inference(superposition,[],[f385,f87])).
fof(f82,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) | $spl4),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f48,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3))))),
  inference(cnf_transformation,[],[f38])).
fof(f121,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3))) | $spl52),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f119,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3))) | ($spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f110,f86])).
fof(f110,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3))) | ($spl4 | $spl10)),
  inference(superposition,[],[f82,f87])).
fof(f80,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4)))) | $spl2),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f49,plain,(
  p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4))))),
  inference(cnf_transformation,[],[f38])).
fof(f118,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4))) | $spl50),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f116,plain,(
  s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4))) | ($spl2 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f109,f86])).
fof(f109,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4))) | ($spl2 | $spl10)),
  inference(superposition,[],[f80,f87])).
fof(f3835,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | $spl305),
  inference(cnf_transformation,[],[f3835_D])).
fof(f3835_D,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f3833,plain,(
  ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl6 | $spl9 | $spl50 | $spl52 | $spl72 | $spl162)),
  inference(subsumption_resolution,[],[f3832,f84])).
fof(f3832,plain,(
  ~p(s(bool,'T_0')) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl6 | $spl9 | $spl50 | $spl52 | $spl72 | $spl162)),
  inference(forward_demodulation,[],[f3831,f121])).
fof(f3831,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl6 | $spl9 | $spl50 | $spl72 | $spl162)),
  inference(subsumption_resolution,[],[f3830,f84])).
fof(f3830,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl9 | $spl50 | $spl72 | $spl162)),
  inference(forward_demodulation,[],[f3829,f118])).
fof(f3829,plain,(
  ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl9 | $spl72 | $spl162)),
  inference(subsumption_resolution,[],[f3815,f86])).
fof(f3815,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.topology.open_2'(s('type.topology.topology'(sK1),sK2),s(fun(sK1,bool),sK3)))) | ~sP0(sK1,'const.topology.open_1'(s('type.topology.topology'(sK1),sK2))) | ($spl72 | $spl162)),
  inference(superposition,[],[f1013,f439])).
fof(f89,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) | $spl14),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f59,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1)))))) & (s(fun(fun(X0,bool),bool),X1) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> s(fun(fun(X0,bool),bool),X1) = s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X13] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X6,bool),bool),X13)))) <=> s(fun(fun(X6,bool),bool),X13) = s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),'const.topology.topology_1'(s(fun(fun(X6,bool),bool),X13))))))),
  file('thm.topology.TOPOLOGY|split|2','thm.topology.topology_tybij|split|1')).
fof(f88,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) ) | $spl12),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f55,plain,(
  ( ! [X0,X1] : (s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : s('type.topology.topology'(X0),X1) = s('type.topology.topology'(X0),'const.topology.topology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X10] : s('type.topology.topology'(X6),X10) = s('type.topology.topology'(X6),'const.topology.topology_1'(s(fun(fun(X6,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X6),X10)))))),
  file('thm.topology.TOPOLOGY|split|2','thm.topology.topology_tybij|split|0')).
fof(f713,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | $spl120),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f712,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | ($spl12 | $spl14)),
  inference(trivial_inequality_removal,[],[f711])).
fof(f711,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))) != s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X0),X1))))))) ) | ($spl12 | $spl14)),
  inference(superposition,[],[f89,f88])).
fof(f87,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f87_D])).
fof(f87_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f54,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.topology.TOPOLOGY|split|2',aHL_BOOL_CASES)).
fof(f86,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f53,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.topology.TOPOLOGY|split|2',aHL_FALSITY)).
fof(f115,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl48),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f107,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl10)),
  inference(superposition,[],[f86,f87])).
fof(f2464,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) | $spl244),
  inference(cnf_transformation,[],[f2464_D])).
fof(f2464_D,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2454,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.topology.istopology_1'(s(fun(fun(X2,bool),bool),'const.topology.open_1'(s('type.topology.topology'(X2),X3)))))) ) | ($spl48 | $spl120)),
  inference(resolution,[],[f713,f115])).
fof(f102,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) | $spl40),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f70,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) | sP0(X0,X1)) & (~sP0(X0,X1) | p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1))))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> sP0(X0,X1))),
  inference(definition_folding,[],[f34,f36])).
fof(f34,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X0,bool),bool),X1)))) <=> (p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) & ! [X3,X4] : ((p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X3)))) & p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),X4))))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.topology.re_intersect_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X4))))))) & ! [X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) => p(s(bool,happ(s(fun(fun(X0,bool),bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.topology.istopology_1'(s(fun(fun(X6,bool),bool),X7)))) <=> (p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0')))) & ! [X10,X11] : ((p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),X10)))) & p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),X11))))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.topology.re_intersect_2'(s(fun(X6,bool),X10),s(fun(X6,bool),X11))))))) & ! [X12] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(fun(X6,bool),bool),X12),s(fun(fun(X6,bool),bool),X7)))) => p(s(bool,happ(s(fun(fun(X6,bool),bool),X7),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X12)))))))))),
  file('thm.topology.TOPOLOGY|split|2','thm.topology.istopology')).
fof(f84,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f51,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.topology.TOPOLOGY|split|2',aHL_TRUTH)).
fof(f3895,plain,(
  $false | ($spl6 | $spl40 | $spl244 | $spl305)),
  inference(subsumption_resolution,[],[f3894,f84])).
fof(f3894,plain,(
  ~p(s(bool,'T_0')) | ($spl40 | $spl244 | $spl305)),
  inference(forward_demodulation,[],[f3858,f2464])).
fof(f3858,plain,(
  ~p(s(bool,'const.topology.istopology_1'(s(fun(fun(sK1,bool),bool),'const.topology.open_1'(s('type.topology.topology'(sK1),sK2)))))) | ($spl40 | $spl305)),
  inference(resolution,[],[f3835,f102])).
% SZS output end Proof for thm.topology
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 78
% Passive clauses: 186
% Generated clauses: 3798
% Final active clauses: 66
% Final passive clauses: 70
% Input formulas: 12
% Initial clauses: 28
% 
% Duplicate literals: 364
% Trivial inequalities: 18
% Fw subsumption resolutions: 332
% Fw demodulations: 475
% Bw demodulations: 44
% Condensations: 299
% 
% Simple tautologies: 78
% Equational tautologies: 18
% Forward subsumptions: 1744
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 45
% Forward superposition: 1405
% Backward superposition: 517
% Self superposition: 47
% Equality factoring: 32
% Equality resolution: 2
% 
% Split clauses: 24
% Split components: 70
% SAT solver clauses: 298
% SAT solver unit clauses: 34
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 1
% 
% Sat splits: 24
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 297
% TWLsolver calls for satisfiability: 389
% 
% Memory used [KB]: 2814
% Time elapsed: 0.096 s
% ------------------------------
% ------------------------------
% Success in time 0.709 s
