Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (102219)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 662
% Passive clauses: 10153
% Generated clauses: 18753
% Final active clauses: 657
% Final passive clauses: 9071
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 283
% Trivial inequalities: 60
% Fw subsumption resolutions: 1570
% Fw demodulations: 209
% Bw demodulations: 5
% 
% Simple tautologies: 221
% Equational tautologies: 56
% Forward subsumptions: 5884
% 
% Binary resolution: 11559
% Factoring: 1130
% Forward superposition: 1532
% Backward superposition: 2007
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 327
% Split components: 661
% SAT solver clauses: 674
% SAT solver unit clauses: 20
% SAT solver binary clauses: 186
% 
% Sat splits: 327
% 
% TWLsolver clauses: 674
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (102231)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 452
% Passive clauses: 3569
% Generated clauses: 16569
% Final active clauses: 357
% Final passive clauses: 1260
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 219
% Trivial inequalities: 82
% Fw subsumption resolutions: 1689
% Fw demodulations: 714
% Bw demodulations: 1744
% Condensations: 357
% Global subsumptions: 18
% 
% Simple tautologies: 269
% Equational tautologies: 17
% Forward subsumptions: 6133
% Bw demodulations to eq. taut.: 147
% 
% Binary resolution: 3118
% Factoring: 100
% Forward superposition: 1681
% Backward superposition: 1793
% Self superposition: 25
% Equality factoring: 28
% Equality resolution: 8
% 
% Split clauses: 134
% Split components: 273
% SAT solver clauses: 3545
% SAT solver unit clauses: 195
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 11
% 
% Sat splits: 134
% Sat splitting refutations: 21
% 
% TWLsolver clauses: 3472
% TWLsolver calls for satisfiability: 1659
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f45751,plain,(
  $false),
  inference(subsumption_resolution,[],[f45750,f523])).
fof(f523,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f486,f295])).
fof(f295,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f486,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f283,f280])).
fof(f280,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0',aHL_TRUTH)).
fof(f283,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f25])).
fof(f25,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f45750,plain,(
  s(bool,'T_0') = s(bool,'F_0')),
  inference(condensation,[],[f45749])).
fof(f45749,plain,(
  ( ! [X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X117)) )),
  inference(forward_demodulation,[],[f45748,f43440])).
fof(f43440,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f43439,f42834])).
fof(f42834,plain,(
  ( ! [X1198,X1200,X1199] : (s(bool,'T_0') != s(bool,happ(s(fun(X1198,bool),X1199),s(X1198,X1200))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X1198,bool),X1199)))) )),
  inference(subsumption_resolution,[],[f42121,f524])).
fof(f524,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f279,f295])).
fof(f279,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f131])).
fof(f131,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0')).
fof(f42121,plain,(
  ( ! [X1198,X1200,X1199] : (s(bool,'T_0') != s(bool,happ(s(fun(X1198,bool),X1199),s(X1198,X1200))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X1198,bool),X1199)))) )),
  inference(superposition,[],[f524,f8004])).
fof(f8004,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X24) = s(bool,happ(s(fun(X22,bool),X23),s(X22,X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) )),
  inference(subsumption_resolution,[],[f8003,f486])).
fof(f8003,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'F_0')) | s(bool,X24) = s(bool,happ(s(fun(X22,bool),X23),s(X22,X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) )),
  inference(forward_demodulation,[],[f7998,f544])).
fof(f544,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f279,f300])).
fof(f300,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f7998,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | s(bool,X24) = s(bool,happ(s(fun(X22,bool),X23),s(X22,X25))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) )),
  inference(superposition,[],[f2259,f1425])).
fof(f1425,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1422,f486])).
fof(f1422,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f741,f544])).
fof(f741,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f733,f486])).
fof(f733,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f345,f287])).
fof(f287,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0',aHL_BOOL_CASES)).
fof(f345,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f2259,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,'|dollar|exists_1'(s(fun(X113,bool),X114)))) | s(bool,X116) = s(bool,happ(s(fun(X113,bool),X114),s(X113,X115))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f1876,f280])).
fof(f1876,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X113,bool),X114)))) | s(bool,X116) = s(bool,happ(s(fun(X113,bool),X114),s(X113,X115))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f487,f682])).
fof(f682,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f287,f287])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f342,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.EXISTS_DEF')).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,'const.min.@_1'(s(fun(X6,bool),X9)))))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.SELECT_AX')).
fof(f43439,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(forward_demodulation,[],[f42608,f43309])).
fof(f43309,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) )),
  inference(duplicate_literal_removal,[],[f43251])).
fof(f43251,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0))) | s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) )),
  inference(superposition,[],[f5613,f4437])).
fof(f4437,plain,(
  ( ! [X8] : (s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK20(X8,sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) )),
  inference(subsumption_resolution,[],[f4436,f486])).
fof(f4436,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK20(X8,sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) )),
  inference(forward_demodulation,[],[f4427,f544])).
fof(f4427,plain,(
  ( ! [X8] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | s(sK11,X8) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK20(X8,sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X8)))) )),
  inference(superposition,[],[f4400,f1425])).
fof(f4400,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) | s(sK11,X0) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK20(X0,sK13,sK12,sK10,sK11))))) )),
  inference(subsumption_resolution,[],[f4393,f280])).
fof(f4393,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,X0)))) | s(sK11,X0) = s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,sK20(X0,sK13,sK12,sK10,sK11))))) )),
  inference(superposition,[],[f387,f527])).
fof(f527,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f278,f296])).
fof(f296,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f278,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13))))),
  inference(cnf_transformation,[],[f183])).
fof(f387,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK20(X6,X3,X2,X1,X0)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK21(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK21(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X7,X6] : ! [X20,X21] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X6,X7),X20),s(fun(X7,bool),X21)))) <=> ! [X18] : (~p(s(bool,happ(s(fun(X7,bool),X21),s(X7,X18)))) => ? [X19] : s(X7,X18) = s(X7,happ(s(fun(X6,X7),X20),s(X6,X19)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.quantHeuristics.GUESS_REWRITES|split|5')).
fof(f5613,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X0)))))) )),
  inference(subsumption_resolution,[],[f5605,f486])).
fof(f5605,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,happ(s(fun(sK10,sK11),sK12),s(sK10,X0)))))) )),
  inference(superposition,[],[f1435,f544])).
fof(f1435,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(subsumption_resolution,[],[f1434,f486])).
fof(f1434,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(superposition,[],[f392,f287])).
fof(f392,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK22(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK23(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7,X6] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) <=> ! [X18] : (~p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X18)))) => ? [X19] : ~p(s(bool,happ(s(fun(X7,bool),X9),s(X7,happ(s(fun(X6,X7),X8),s(X6,X19))))))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f42608,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11))))) )),
  inference(condensation,[],[f42607])).
fof(f42607,plain,(
  ( ! [X494,X496,X495,X497] : (s(bool,happ(s(fun(X495,bool),X496),s(X495,X497))) = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,X494) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X495,bool),X496)))) )),
  inference(duplicate_literal_removal,[],[f42085])).
fof(f42085,plain,(
  ( ! [X494,X496,X495,X497] : (s(bool,happ(s(fun(X495,bool),X496),s(X495,X497))) = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,X494) | s(bool,'T_0') = s(bool,X494) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X495,bool),X496)))) )),
  inference(superposition,[],[f16003,f8004])).
fof(f16003,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f15984,f486])).
fof(f15984,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2262,f544])).
fof(f2262,plain,(
  ( ! [X151,X149,X147,X150,X148] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X150,X147),X149),s(fun(X147,bool),X148)))) | s(bool,X151) = s(bool,happ(s(fun(X147,bool),X148),s(X147,sK23(X148,X149,X150,X147)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(subsumption_resolution,[],[f1883,f280])).
fof(f1883,plain,(
  ( ! [X151,X149,X147,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X150,X147),X149),s(fun(X147,bool),X148)))) | s(bool,X151) = s(bool,happ(s(fun(X147,bool),X148),s(X147,sK23(X148,X149,X150,X147)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f391,f682])).
fof(f391,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK23(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f45748,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(subsumption_resolution,[],[f45747,f287])).
fof(f45747,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(forward_demodulation,[],[f41047,f43309])).
fof(f41047,plain,(
  ( ! [X118,X119,X117] : (s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(duplicate_literal_removal,[],[f40981])).
fof(f40981,plain,(
  ( ! [X118,X119,X117] : (s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) != s(bool,happ(s(fun(sK11,bool),sK13),s(sK11,sK23(sK13,sK12,sK10,sK11)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f7803,f16003])).
fof(f7803,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f7802,f486])).
fof(f7802,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'F_0')) | s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(forward_demodulation,[],[f7778,f544])).
fof(f7778,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) != s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK10,sK11),sK12),s(fun(sK11,bool),sK13)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f5245,f1425])).
fof(f5245,plain,(
  ( ! [X85,X83,X84] : (s(bool,X85) != s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84))) | p(s(bool,X85)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) )),
  inference(forward_demodulation,[],[f5244,f337])).
fof(f5244,plain,(
  ( ! [X85,X83,X84] : (s(bool,X85) != s(bool,happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84))))) | p(s(bool,X85)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) )),
  inference(subsumption_resolution,[],[f5193,f280])).
fof(f5193,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,'T_0')) | s(bool,X85) != s(bool,happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84))))) | p(s(bool,X85)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) )),
  inference(superposition,[],[f477,f1236])).
fof(f1236,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f337,f287])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f465])).
fof(f465,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10,X16,X14,X17] : ((s(bool,X10) = s(bool,X16) & (p(s(bool,X16)) => s(bool,X14) = s(bool,X17))) => ((p(s(bool,X10)) => p(s(bool,X14))) <=> (p(s(bool,X16)) => p(s(bool,X17)))))),
  file('thm.quantHeuristics.GUESSES_WEAKEN_THM|split|0','thm.bool.IMP_CONG')).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1061
% Passive clauses: 3917
% Generated clauses: 46075
% Final active clauses: 840
% Final passive clauses: 1150
% Input formulas: 52
% Initial clauses: 188
% 
% Duplicate literals: 3373
% Trivial inequalities: 124
% Fw subsumption resolutions: 2888
% Fw demodulations: 4377
% Bw demodulations: 1457
% Condensations: 4777
% Global subsumptions: 35
% 
% Simple tautologies: 226
% Equational tautologies: 862
% Forward subsumptions: 24027
% Bw demodulations to eq. taut.: 203
% 
% Unit resulting resolution: 8824
% Factoring: 2
% Forward superposition: 8733
% Backward superposition: 9324
% Self superposition: 186
% Equality factoring: 159
% Equality resolution: 6
% 
% Split clauses: 411
% Split components: 848
% SAT solver clauses: 5074
% SAT solver unit clauses: 634
% SAT solver binary clauses: 905
% 
% Sat splits: 411
% Sat splitting refutations: 298
% 
% Memory used [KB]: 28016
% Time elapsed: 4.187 s
% ------------------------------
% ------------------------------
% Success in time 5.503 s
