Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f7463,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7381,f71,f109,f148,f314,f4258,f5338,f70,f107,f353,f351,f830,f828,f805,f807,f688,f757])).
fof(f757,plain,(
  $false | $spl12),
  inference(trivial_inequality_removal,[],[f754])).
fof(f754,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl12),
  inference(superposition,[],[f549,f688])).
fof(f549,plain,(
  s(bool,'T_0') != s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))),
  inference(superposition,[],[f448,f98])).
fof(f98,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X8] : ! [X26,X27] : s(X7,'const.quotient.respects_2'(s(fun(X8,fun(X8,X7)),X26),s(X8,X27))) = s(X7,happ(s(fun(X8,X7),'const.quotient.respects_1'(s(fun(X8,fun(X8,X7)),X26))),s(X8,X27)))),
  file('thm.quotient.quotient_89_prove_ltb','ARITY_1_2_const.quotient.respects')).
fof(f448,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))),s(sK2,sK4)))),
  inference(superposition,[],[f140,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : ! [X20,X10] : s(bool,happ(s(fun(X8,bool),X20),s(X8,X10))) = s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),X20)))),
  file('thm.quotient.quotient_89_prove_ltb','thm.pred_set.SPECIFICATION')).
fof(f140,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10)))))),
  inference(resolution,[],[f69,f73])).
fof(f73,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X19] : (s(bool,'T_0') = s(bool,X19) <=> p(s(bool,X19)))),
  file('thm.quotient.quotient_89_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f69,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))))))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5)))) & p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3)))) & ! [X11,X12] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,X11))),s(sK2,X12)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK9),s(sK1,happ(s(fun(sK2,sK1),sK7),s(sK2,X11))))),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,X12))))))) & ~p(s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f45])).
fof(f45,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) & p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) & ! [X11,X12] : (~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) & ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) & (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) & (! [X11,X12] : (~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) & ~p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => (! [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) => p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => (! [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) => p(s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X16),s(fun(X8,X6),X14),s(fun(X6,X8),X11)))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X10))),s(X8,X9)))) => (! [X17,X18] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X17))),s(X8,X18)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X15),s(X7,happ(s(fun(X8,X7),X13),s(X8,X17))))),s(X7,happ(s(fun(X8,X7),X12),s(X8,X18))))))) => p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.quotient.respects_1'(s(fun(X8,fun(X8,bool)),X16)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X16),s(fun(X8,X6),X14),s(fun(X6,X8),X11)))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X10))),s(X8,X9)))) => (! [X17,X18] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X16),s(X8,X17))),s(X8,X18)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X15),s(X7,happ(s(fun(X8,X7),X13),s(X8,X17))))),s(X7,happ(s(fun(X8,X7),X12),s(X8,X18))))))) => p(s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),'const.quotient.respects_1'(s(fun(X8,fun(X8,bool)),X16)))))))))),
  file('thm.quotient.quotient_89_prove_ltb','thm.quotient.quotient_89_prove_ltb')).
fof(f688,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl12),
  inference(cnf_transformation,[],[f688_D])).
fof(f688_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f807,plain,(
  ~p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | $spl17),
  inference(cnf_transformation,[],[f807_D])).
fof(f807_D,plain,(
  ~p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f805,plain,(
  ( ! [X13] : (~p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl5),
  inference(subsumption_resolution,[],[f798,f72])).
fof(f72,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.quotient_89_prove_ltb',aHL_BOOL_CASES)).
fof(f798,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | ~p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl5),
  inference(superposition,[],[f77,f592])).
fof(f592,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(subsumption_resolution,[],[f572,f112])).
fof(f112,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl5),
  inference(resolution,[],[f73,f109])).
fof(f572,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f542,f118])).
fof(f118,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f72,f72])).
fof(f542,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))),
  inference(superposition,[],[f98,f442])).
fof(f442,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10))),s(sK2,sK4)))),
  inference(superposition,[],[f96,f139])).
fof(f139,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK2,sK4),s(fun(sK2,bool),'const.quotient.respects_1'(s(fun(sK2,fun(sK2,bool)),sK10)))))),
  inference(resolution,[],[f69,f78])).
fof(f78,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X19] : (s(bool,'F_0') = s(bool,X19) <=> ~p(s(bool,X19)))),
  file('thm.quotient.quotient_89_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f77,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f51])).
fof(f828,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))) | $spl18),
  inference(cnf_transformation,[],[f828_D])).
fof(f828_D,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f830,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))) | $spl17),
  inference(resolution,[],[f807,f78])).
fof(f351,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) | $spl10),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f353,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))) | $spl5),
  inference(subsumption_resolution,[],[f348,f109])).
fof(f348,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3)))),
  inference(superposition,[],[f67,f72])).
fof(f67,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4))),s(sK2,sK3))))),
  inference(cnf_transformation,[],[f48])).
fof(f107,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f70,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.quotient_89_prove_ltb',aHL_TRUTH)).
fof(f5338,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | $spl114),
  inference(cnf_transformation,[],[f5338_D])).
fof(f5338_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4258,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl5 | $spl10 | $spl18)),
  inference(subsumption_resolution,[],[f4257,f109])).
fof(f4257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl10 | $spl18)),
  inference(forward_demodulation,[],[f4256,f828])).
fof(f4256,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f4216,f107])).
fof(f4216,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.quotient.respects_2'(s(fun(sK2,fun(sK2,bool)),sK10),s(sK2,sK4)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,X0),X1),s(fun(X0,sK2),X2)))) ) | $spl10),
  inference(superposition,[],[f1219,f351])).
fof(f1219,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,X4),X5),s(fun(X4,X0),X6)))) )),
  inference(resolution,[],[f1218,f78])).
fof(f1218,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | p(s(bool,'const.quotient.respects_2'(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))))) )),
  inference(forward_demodulation,[],[f100,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ! [X21,X10] : s(bool,'const.quotient.respects_2'(s(fun(X8,fun(X8,bool)),X21),s(X8,X10))) = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X21),s(X8,X10))),s(X8,X10)))),
  file('thm.quotient.quotient_89_prove_ltb','thm.quotient.RESPECTS')).
fof(f100,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X8] : ! [X21,X22,X23] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X21),s(fun(X8,X7),X22),s(fun(X7,X8),X23)))) => ! [X24,X25] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X21),s(X8,X24))),s(X8,X25)))) <=> (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X21),s(X8,X24))),s(X8,X24)))) & p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X21),s(X8,X25))),s(X8,X25)))) & s(X7,happ(s(fun(X8,X7),X22),s(X8,X24))) = s(X7,happ(s(fun(X8,X7),X22),s(X8,X25))))))),
  file('thm.quotient.quotient_89_prove_ltb','thm.quotient.QUOTIENT_REL')).
fof(f314,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))) | $spl8),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f148,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5)))),
  inference(resolution,[],[f66,f74])).
fof(f74,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f66,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK10),s(fun(sK2,sK0),sK8),s(fun(sK0,sK2),sK5))))),
  inference(cnf_transformation,[],[f48])).
fof(f109,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f71,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.quotient_89_prove_ltb',aHL_FALSITY)).
fof(f7381,plain,(
  $false | ($spl5 | $spl8 | $spl114)),
  inference(subsumption_resolution,[],[f7292,f112])).
fof(f7292,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl114)),
  inference(superposition,[],[f5338,f314])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 107
% Passive clauses: 979
% Generated clauses: 7302
% Final active clauses: 91
% Final passive clauses: 738
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 383
% Trivial inequalities: 104
% Fw subsumption resolutions: 800
% Fw demodulations: 1427
% Bw demodulations: 14
% 
% Simple tautologies: 193
% Equational tautologies: 20
% Forward subsumptions: 3201
% 
% Binary resolution: 170
% Forward superposition: 2622
% Backward superposition: 1585
% Self superposition: 55
% Equality factoring: 34
% Equality resolution: 15
% 
% Split clauses: 119
% Split components: 244
% SAT solver clauses: 330
% SAT solver unit clauses: 14
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 1
% 
% Sat splits: 119
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 323
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 4733
% Time elapsed: 0.121 s
% ------------------------------
% ------------------------------
% Success in time 0.118 s
