Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.list
% (79265)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 773
% Passive clauses: 8605
% Generated clauses: 20809
% Final active clauses: 713
% Final passive clauses: 7283
% Input formulas: 45
% Initial clauses: 187
% 
% Duplicate literals: 577
% Trivial inequalities: 121
% Fw subsumption resolutions: 2155
% Fw demodulations: 514
% Bw demodulations: 58
% 
% Simple tautologies: 244
% Equational tautologies: 75
% Forward subsumptions: 7848
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 10231
% Factoring: 1392
% Forward superposition: 3183
% Backward superposition: 1962
% Self superposition: 40
% Equality factoring: 28
% Equality resolution: 60
% 
% Split clauses: 634
% Split components: 1276
% SAT solver clauses: 1160
% SAT solver unit clauses: 18
% SAT solver binary clauses: 175
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 634
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1157
% TWLsolver calls for satisfiability: 473
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (79275)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 407
% Passive clauses: 2229
% Generated clauses: 10046
% Final active clauses: 398
% Final passive clauses: 1794
% Input formulas: 45
% Initial clauses: 187
% 
% Duplicate literals: 237
% Trivial inequalities: 19
% Fw subsumption resolutions: 787
% Fw demodulations: 222
% Bw demodulations: 30
% Condensations: 263
% Global subsumptions: 14
% 
% Simple tautologies: 227
% Equational tautologies: 29
% Forward subsumptions: 3373
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3435
% Factoring: 352
% Forward superposition: 1269
% Backward superposition: 864
% Self superposition: 36
% Equality factoring: 26
% Equality resolution: 24
% 
% Split clauses: 52
% Split components: 110
% SAT solver clauses: 2828
% SAT solver unit clauses: 184
% SAT solver binary clauses: 90
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 11
% 
% Sat splits: 52
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2771
% TWLsolver calls for satisfiability: 1830
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f38367,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f245,f444,f442,f440,f244,f443,f243,f441,f439,f242,f448,f446,f453,f1186,f1190,f1193,f1196,f1199,f1202,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1267,f1187,f1269,f1191,f1271,f1194,f1273,f1197,f1275,f1200,f1277,f1203,f1279,f1206,f1281,f1209,f1283,f1212,f1285,f1215,f1287,f1218,f1289,f1221,f1291,f1224,f1293,f1227,f1295,f1230,f1297,f1233,f1299,f1236,f1301,f1188,f1239,f1308,f1309,f1814,f1815,f1817,f1819,f1818,f1894,f1895,f1896,f1898,f1897,f1933,f1934,f1910,f1939,f1909,f1908,f1940,f1936,f1907,f1906,f1943,f1941,f2626,f2627,f8466,f8468,f8470,f8472,f5926,f8498,f8497,f8499,f8500,f8504,f8505,f8602,f8467,f8603,f8469,f8604,f8471,f8605,f1816,f8473,f9349,f9352,f9353,f9354,f9475,f9472,f9357,f9358,f9477,f10507,f10512,f10515,f10518,f10521,f10523,f10489,f10488,f10487,f10486,f10522,f10485,f10526,f10484,f10527,f10483,f10482,f10531,f10481,f10480,f10479,f10528,f10478,f10508,f10477,f10509,f6853,f6313,f2639,f10532,f5773,f5620,f10533,f2637,f10833,f10834,f10859,f10835,f10839,f10860,f1937,f11017,f9474,f11077,f11079,f9473,f17927,f17928,f17931,f17932,f17933,f17934,f17935,f17936,f17937,f17938,f17948,f17949,f17950,f17951,f17952,f17981,f17953,f17983,f17956,f17984,f17957,f17958,f17986,f17959,f17985,f17989,f17960,f17987,f17961,f17988,f17990,f17962,f17991,f17963,f17982,f17992,f17973,f17994,f17974,f17996,f17995,f17975,f17998,f17997,f17976,f17999,f17993,f18000,f11080,f14799,f10530,f18002,f32832,f35078,f36360,f38172,f38173])).
fof(f38173,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f37929])).
fof(f37929,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) | ($spl0 | $spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f442,f36612,f37375,f294])).
fof(f294,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | ! [X3,X4] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | ((s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4))) & (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),X4) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | ! [X3,X4] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | ((s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4))) & ((s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),X4)) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | ! [X3,X4] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) <=> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) => ! [X3,X4] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) => (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) <=> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1,X2] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) => ! [X3,X4] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) => (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) <=> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X7,X9] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9))) => ! [X8,X10] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X10))) => (s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X10))) <=> (s('type.list.list'(X6),X7) = s('type.list.list'(X6),X9) & s('type.list.list'(X6),X8) = s('type.list.list'(X6),X10)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.list.APPEND_LENGTH_EQ')).
fof(f37375,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f36945,f2715])).
fof(f2715,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),'const.list.APPEND_2'(s('type.list.list'(X8),X10),s('type.list.list'(X8),X9))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),'const.list.APPEND_2'(s('type.list.list'(X8),X11),s('type.list.list'(X8),X10))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X9))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X11)))) )),
  inference(superposition,[],[f1603,f1594])).
fof(f1594,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X4),s('type.list.list'(X3),X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X5))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X4)))))) )),
  inference(superposition,[],[f293,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X11,X12] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X11)))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.arithmetic.ADD_SYM')).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X8] : s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.list.LENGTH_APPEND')).
fof(f1603,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) | s('type.num.num',X3) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))) )),
  inference(superposition,[],[f347,f293])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12,X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))) <=> s('type.num.num',X11) = s('type.num.num',X12))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.arithmetic.EQ_MONO_ADD_EQ')).
fof(f36945,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK14),s('type.list.list'(sK12),sK15))))) | ($spl0 | $spl4)),
  inference(superposition,[],[f1832,f36612])).
fof(f1832,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X4),s('type.list.list'(X3),X5))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X5),s('type.list.list'(X3),X4)))))) )),
  inference(superposition,[],[f1594,f293])).
fof(f36612,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f443,f439])).
fof(f38172,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f37906])).
fof(f37906,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) | ($spl0 | $spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f442,f36612,f37375,f294])).
fof(f36360,plain,(
  $false | ($spl0 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f36058,f35893])).
fof(f35893,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),X0))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X0)))))) ) | ($spl0 | $spl5 | $spl8)),
  inference(unit_resulting_resolution,[],[f35450,f2716])).
fof(f2716,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X12),'const.list.APPEND_2'(s('type.list.list'(X12),X13),s('type.list.list'(X12),X14))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X12),'const.list.APPEND_2'(s('type.list.list'(X12),X15),s('type.list.list'(X12),X14))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X12),X13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X12),X15)))) )),
  inference(superposition,[],[f1603,f293])).
fof(f35450,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | ($spl0 | $spl5 | $spl8)),
  inference(unit_resulting_resolution,[],[f444,f439,f448,f295])).
fof(f295,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),X4)) )),
  inference(cnf_transformation,[],[f181])).
fof(f36058,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))))) | ($spl0 | $spl8)),
  inference(superposition,[],[f35683,f439])).
fof(f35683,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK14))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16)))))) ) | $spl8),
  inference(forward_demodulation,[],[f35457,f293])).
fof(f35457,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X0))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))) ) | $spl8),
  inference(superposition,[],[f293,f448])).
fof(f35078,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f34773,f33698])).
fof(f33698,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK14))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16)))))) ) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f33656,f293])).
fof(f33656,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X0))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl0 | $spl5 | $spl6)),
  inference(unit_resulting_resolution,[],[f33282,f1863])).
fof(f1863,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X21),X23))),s('type.num.num',X24))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X21),'const.list.APPEND_2'(s('type.list.list'(X21),X23),s('type.list.list'(X21),X22))))) | s('type.num.num',X24) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X21),X22)))) )),
  inference(superposition,[],[f821,f1594])).
fof(f821,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(superposition,[],[f347,f303])).
fof(f33282,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))) | ($spl0 | $spl5 | $spl6)),
  inference(unit_resulting_resolution,[],[f444,f446,f439,f295])).
fof(f34773,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK16))))) | ($spl0 | $spl6)),
  inference(superposition,[],[f33089,f439])).
fof(f33089,plain,(
  ( ! [X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),X1))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X1)))))) ) | $spl6),
  inference(forward_demodulation,[],[f32863,f293])).
fof(f32863,plain,(
  ( ! [X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X1)))))) ) | $spl6),
  inference(superposition,[],[f293,f446])).
fof(f32832,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f32829])).
fof(f32829,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f4455,f454])).
fof(f454,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK16))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f439,f441])).
fof(f4455,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK14))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16)))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f4445,f293])).
fof(f4445,plain,(
  ( ! [X0] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),X0),s('type.list.list'(sK12),sK16))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),X0))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3220,f1594])).
fof(f3220,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))),s('type.num.num',X0)))) ) | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f2880,f821])).
fof(f2880,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))) | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f444,f454,f1647])).
fof(f1647,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),X3)) )),
  inference(equality_resolution,[],[f295])).
fof(f18002,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,g(s(bool,sK19(X4,X5,bool,bool))))) ) | $spl164),
  inference(cnf_transformation,[],[f18002_D])).
fof(f18002_D,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,g(s(bool,sK19(X4,X5,bool,bool))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f10530,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl111),
  inference(cnf_transformation,[],[f10530_D])).
fof(f10530_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f14799,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,X6) = s(bool,g(s(bool,sK19(X4,X5,bool,bool)))) | s(bool,'T_0') = s(bool,X6)) ) | $spl60),
  inference(duplicate_literal_removal,[],[f14766])).
fof(f14766,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,X6) = s(bool,g(s(bool,sK19(X4,X5,bool,bool)))) | s(bool,'T_0') = s(bool,X6) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5)) ) | $spl60),
  inference(superposition,[],[f1575,f1937])).
fof(f1575,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK19(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f346,f561])).
fof(f561,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f254,f254])).
fof(f254,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.list_|49||48|8_prove_ltb',aHL_BOOL_CASES)).
fof(f346,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK19(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK19(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.list.list_|49||48|8_prove_ltb',aHL_EXT)).
fof(f11080,plain,(
  ( ! [X263,X264] : (s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) | $spl124),
  inference(cnf_transformation,[],[f11080_D])).
fof(f11080_D,plain,(
  ( ! [X263,X264] : (s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f18000,plain,(
  ( ! [X263,X264] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) | ($spl60 | $spl92)),
  inference(forward_demodulation,[],[f11015,f9474])).
fof(f11015,plain,(
  ( ! [X263,X264] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) | $spl60),
  inference(duplicate_literal_removal,[],[f11012])).
fof(f11012,plain,(
  ( ! [X263,X264] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X263) = s(fun(bool,bool),X264) | s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) | $spl60),
  inference(superposition,[],[f346,f1937])).
fof(f17993,plain,(
  ( ! [X101,X102] : (s(fun(fun(bool,bool),bool),X101) = s(fun(fun(bool,bool),bool),X102) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X101,X102,bool,fun(bool,bool)))))) ) | $spl150),
  inference(cnf_transformation,[],[f17993_D])).
fof(f17993_D,plain,(
  ( ! [X101,X102] : (s(fun(fun(bool,bool),bool),X101) = s(fun(fun(bool,bool),bool),X102) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X101,X102,bool,fun(bool,bool)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f17999,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X114)))) ) | $spl162),
  inference(cnf_transformation,[],[f17999_D])).
fof(f17999_D,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X114)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f17976,plain,(
  ( ! [X114,X112,X113] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X114))) | s(fun(fun(bool,bool),bool),X112) = s(fun(fun(bool,bool),bool),X113) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X112,X113,bool,fun(bool,bool)))))) ) | $spl124),
  inference(superposition,[],[f1579,f11080])).
fof(f1579,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK19(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK19(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f346,f254])).
fof(f17997,plain,(
  ( ! [X109,X110] : (s(fun(fun(bool,bool),bool),X109) = s(fun(fun(bool,bool),bool),X110) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK19(X109,X110,bool,fun(bool,bool)))))) ) | $spl158),
  inference(cnf_transformation,[],[f17997_D])).
fof(f17997_D,plain,(
  ( ! [X109,X110] : (s(fun(fun(bool,bool),bool),X109) = s(fun(fun(bool,bool),bool),X110) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK19(X109,X110,bool,fun(bool,bool)))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f17998,plain,(
  ( ! [X111] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X111)))) ) | $spl160),
  inference(cnf_transformation,[],[f17998_D])).
fof(f17998_D,plain,(
  ( ! [X111] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X111)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f17975,plain,(
  ( ! [X111,X109,X110] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X111))) | s(fun(fun(bool,bool),bool),X109) = s(fun(fun(bool,bool),bool),X110) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK19(X109,X110,bool,fun(bool,bool)))))) ) | $spl124),
  inference(superposition,[],[f2641,f11080])).
fof(f2641,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK19(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK19(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f2640])).
fof(f2640,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2633,f254])).
fof(f2633,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1579,f561])).
fof(f17995,plain,(
  ( ! [X105,X108,X106] : (s(fun(fun(bool,bool),bool),X105) = s(fun(fun(bool,bool),bool),X106) | s(bool,'T_0') = s(bool,X108) | s(bool,X108) = s(bool,g(s(fun(bool,bool),sK19(X105,X106,bool,fun(bool,bool)))))) ) | $spl154),
  inference(cnf_transformation,[],[f17995_D])).
fof(f17995_D,plain,(
  ( ! [X105,X108,X106] : (s(fun(fun(bool,bool),bool),X105) = s(fun(fun(bool,bool),bool),X106) | s(bool,'T_0') = s(bool,X108) | s(bool,X108) = s(bool,g(s(fun(bool,bool),sK19(X105,X106,bool,fun(bool,bool)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f17996,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X107)))) ) | $spl156),
  inference(cnf_transformation,[],[f17996_D])).
fof(f17996_D,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X107)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f17974,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X107))) | s(fun(fun(bool,bool),bool),X105) = s(fun(fun(bool,bool),bool),X106) | s(bool,X108) = s(bool,g(s(fun(bool,bool),sK19(X105,X106,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X108)) ) | $spl124),
  inference(superposition,[],[f1575,f11080])).
fof(f17994,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,f(s(fun(bool,bool),X103))) | s(bool,'T_0') = s(bool,X104)) ) | $spl152),
  inference(cnf_transformation,[],[f17994_D])).
fof(f17994_D,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,f(s(fun(bool,bool),X103))) | s(bool,'T_0') = s(bool,X104)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f17973,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X104) != s(bool,f(s(fun(bool,bool),X103))) | s(fun(fun(bool,bool),bool),X101) = s(fun(fun(bool,bool),bool),X102) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X101,X102,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X104)) ) | $spl124),
  inference(superposition,[],[f1577,f11080])).
fof(f1577,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK19(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK19(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f346,f561])).
fof(f17992,plain,(
  ( ! [X54,X55,X51] : (p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) = s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK18(X51,X54,bool))) ) | $spl148),
  inference(cnf_transformation,[],[f17992_D])).
fof(f17992_D,plain,(
  ( ! [X54,X55,X51] : (p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) = s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK18(X51,X54,bool))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f17982,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2))))) ) | $spl128),
  inference(cnf_transformation,[],[f17982_D])).
fof(f17982_D,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f17963,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X53)))) | p(s(bool,X54)) | s(bool,'T_0') = s(bool,sK18(X51,X54,bool)) | s(bool,X55) = s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X55)) ) | $spl124),
  inference(superposition,[],[f2602,f11080])).
fof(f2602,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2597,f246])).
fof(f246,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.list_|49||48|8_prove_ltb',aHL_TRUTH)).
fof(f2597,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1559,f561])).
fof(f1559,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | p(s(bool,X17)) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK18(X16,X17,bool))) )),
  inference(superposition,[],[f302,f254])).
fof(f302,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X19,X20] : (! [X17] : (p(s(bool,X19)) | p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X17))))) <=> (p(s(bool,X19)) | ! [X17] : p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X17))))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f17991,plain,(
  ( ! [X50,X46,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK17(X46,X49,bool))) ) | $spl146),
  inference(cnf_transformation,[],[f17991_D])).
fof(f17991_D,plain,(
  ( ! [X50,X46,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK17(X46,X49,bool))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f17962,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,happ(s(fun(bool,bool),X47),s(bool,X48)))) | p(s(bool,X49)) | s(bool,'T_0') = s(bool,sK17(X46,X49,bool)) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X50)) ) | $spl124),
  inference(superposition,[],[f2567,f11080])).
fof(f2567,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2562,f246])).
fof(f2562,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1542,f561])).
fof(f1542,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | p(s(bool,X17)) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK17(X16,X17,bool))) )),
  inference(superposition,[],[f300,f254])).
fof(f300,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f186])).
fof(f17990,plain,(
  ( ! [X37,X41,X40] : (p(s(bool,X40)) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X41)))) | s(bool,'T_0') = s(bool,sK18(X37,X40,bool))) ) | $spl144),
  inference(cnf_transformation,[],[f17990_D])).
fof(f17990_D,plain,(
  ( ! [X37,X41,X40] : (p(s(bool,X40)) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X41)))) | s(bool,'T_0') = s(bool,sK18(X37,X40,bool))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f17988,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0'))))) ) | $spl140),
  inference(cnf_transformation,[],[f17988_D])).
fof(f17988_D,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0'))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f17961,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'F_0')))) | p(s(bool,X44)) | p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X45)))) | s(bool,'T_0') = s(bool,sK18(X42,X44,bool))) ) | $spl124),
  inference(superposition,[],[f1559,f11080])).
fof(f17987,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl138),
  inference(cnf_transformation,[],[f17987_D])).
fof(f17987_D,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | s(bool,'T_0') = s(bool,X25)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f17960,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X39)))) | p(s(bool,X40)) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X41)))) | s(bool,'T_0') = s(bool,sK18(X37,X40,bool)) | s(bool,'T_0') = s(bool,X39)) ) | $spl124),
  inference(superposition,[],[f1557,f11080])).
fof(f1557,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X10)))) | p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | s(bool,'T_0') = s(bool,sK18(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f302,f561])).
fof(f17989,plain,(
  ( ! [X35,X36,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,X36) | s(bool,X36) = s(bool,sK18(X32,X34,bool)) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X35))))) ) | $spl142),
  inference(cnf_transformation,[],[f17989_D])).
fof(f17989_D,plain,(
  ( ! [X35,X36,X34,X32] : (p(s(bool,X34)) | s(bool,'T_0') = s(bool,X36) | s(bool,X36) = s(bool,sK18(X32,X34,bool)) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X35))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f17985,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))))) ) | $spl134),
  inference(cnf_transformation,[],[f17985_D])).
fof(f17985_D,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0'))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f17959,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | p(s(bool,X34)) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X35)))) | s(bool,X36) = s(bool,sK18(X32,X34,bool)) | s(bool,'T_0') = s(bool,X36)) ) | $spl124),
  inference(superposition,[],[f1555,f11080])).
fof(f1555,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,X3) = s(bool,sK18(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f302,f561])).
fof(f17986,plain,(
  ( ! [X26,X23,X27] : (p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X27)))) | s(bool,'T_0') = s(bool,sK17(X23,X26,bool))) ) | $spl136),
  inference(cnf_transformation,[],[f17986_D])).
fof(f17986_D,plain,(
  ( ! [X26,X23,X27] : (p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X27)))) | s(bool,'T_0') = s(bool,sK17(X23,X26,bool))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f17958,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) | p(s(bool,X30)) | p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X31)))) | s(bool,'T_0') = s(bool,sK17(X28,X30,bool))) ) | $spl124),
  inference(superposition,[],[f1542,f11080])).
fof(f17957,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X27)))) | s(bool,'T_0') = s(bool,sK17(X23,X26,bool)) | s(bool,'T_0') = s(bool,X25)) ) | $spl124),
  inference(superposition,[],[f1540,f11080])).
fof(f1540,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X10)))) | p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | s(bool,'T_0') = s(bool,sK17(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f300,f561])).
fof(f17984,plain,(
  ( ! [X21,X22,X20,X18] : (p(s(bool,X20)) | s(bool,'T_0') = s(bool,X22) | s(bool,X22) = s(bool,sK17(X18,X20,bool)) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X21))))) ) | $spl132),
  inference(cnf_transformation,[],[f17984_D])).
fof(f17984_D,plain,(
  ( ! [X21,X22,X20,X18] : (p(s(bool,X20)) | s(bool,'T_0') = s(bool,X22) | s(bool,X22) = s(bool,sK17(X18,X20,bool)) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X21))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f17956,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) | p(s(bool,X20)) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X21)))) | s(bool,X22) = s(bool,sK17(X18,X20,bool)) | s(bool,'T_0') = s(bool,X22)) ) | $spl124),
  inference(superposition,[],[f1538,f11080])).
fof(f1538,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,X3) = s(bool,sK17(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f300,f561])).
fof(f17983,plain,(
  ( ! [X8,X5,X9] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X5),s(bool,sK17(X5,X8,bool))))) ) | $spl130),
  inference(cnf_transformation,[],[f17983_D])).
fof(f17983_D,plain,(
  ( ! [X8,X5,X9] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X5),s(bool,sK17(X5,X8,bool))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f17953,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X7)))) | p(s(bool,X8)) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X5),s(bool,sK17(X5,X8,bool)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl124),
  inference(superposition,[],[f1548,f11080])).
fof(f1548,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1543,f246])).
fof(f1543,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f300,f561])).
fof(f17981,plain,(
  ( ! [X4,X0,X3] : (p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X0,X3,bool))))) ) | $spl126),
  inference(cnf_transformation,[],[f17981_D])).
fof(f17981_D,plain,(
  ( ! [X4,X0,X3] : (p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X0,X3,bool))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f17952,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X0,X3,bool)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl124),
  inference(superposition,[],[f1565,f11080])).
fof(f1565,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1560,f246])).
fof(f1560,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f302,f561])).
fof(f17951,plain,(
  ( ! [X114,X112,X113] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X114))) | s(fun(fun(bool,bool),bool),X112) = s(fun(fun(bool,bool),bool),X113) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X112,X113,bool,fun(bool,bool)))))) ) | $spl124),
  inference(superposition,[],[f1579,f11080])).
fof(f17950,plain,(
  ( ! [X111,X109,X110] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X111))) | s(fun(fun(bool,bool),bool),X109) = s(fun(fun(bool,bool),bool),X110) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK19(X109,X110,bool,fun(bool,bool)))))) ) | $spl124),
  inference(superposition,[],[f2641,f11080])).
fof(f17949,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X107))) | s(fun(fun(bool,bool),bool),X105) = s(fun(fun(bool,bool),bool),X106) | s(bool,X108) = s(bool,g(s(fun(bool,bool),sK19(X105,X106,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X108)) ) | $spl124),
  inference(superposition,[],[f1575,f11080])).
fof(f17948,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X104) != s(bool,f(s(fun(bool,bool),X103))) | s(fun(fun(bool,bool),bool),X101) = s(fun(fun(bool,bool),bool),X102) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK19(X101,X102,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X104)) ) | $spl124),
  inference(superposition,[],[f1577,f11080])).
fof(f17938,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X53)))) | p(s(bool,X54)) | s(bool,'T_0') = s(bool,sK18(X51,X54,bool)) | s(bool,X55) = s(bool,happ(s(fun(bool,bool),X51),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X55)) ) | $spl124),
  inference(superposition,[],[f2602,f11080])).
fof(f17937,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,happ(s(fun(bool,bool),X47),s(bool,X48)))) | p(s(bool,X49)) | s(bool,'T_0') = s(bool,sK17(X46,X49,bool)) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X50)) ) | $spl124),
  inference(superposition,[],[f2567,f11080])).
fof(f17936,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'F_0')))) | p(s(bool,X44)) | p(s(bool,happ(s(fun(bool,bool),X42),s(bool,X45)))) | s(bool,'T_0') = s(bool,sK18(X42,X44,bool))) ) | $spl124),
  inference(superposition,[],[f1559,f11080])).
fof(f17935,plain,(
  ( ! [X39,X37,X41,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X39)))) | p(s(bool,X40)) | p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X41)))) | s(bool,'T_0') = s(bool,sK18(X37,X40,bool)) | s(bool,'T_0') = s(bool,X39)) ) | $spl124),
  inference(superposition,[],[f1557,f11080])).
fof(f17934,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'T_0')))) | p(s(bool,X34)) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X35)))) | s(bool,X36) = s(bool,sK18(X32,X34,bool)) | s(bool,'T_0') = s(bool,X36)) ) | $spl124),
  inference(superposition,[],[f1555,f11080])).
fof(f17933,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) | p(s(bool,X30)) | p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X31)))) | s(bool,'T_0') = s(bool,sK17(X28,X30,bool))) ) | $spl124),
  inference(superposition,[],[f1542,f11080])).
fof(f17932,plain,(
  ( ! [X26,X24,X23,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | p(s(bool,X26)) | p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X27)))) | s(bool,'T_0') = s(bool,sK17(X23,X26,bool)) | s(bool,'T_0') = s(bool,X25)) ) | $spl124),
  inference(superposition,[],[f1540,f11080])).
fof(f17931,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,'T_0')))) | p(s(bool,X20)) | p(s(bool,happ(s(fun(bool,bool),X18),s(bool,X21)))) | s(bool,X22) = s(bool,sK17(X18,X20,bool)) | s(bool,'T_0') = s(bool,X22)) ) | $spl124),
  inference(superposition,[],[f1538,f11080])).
fof(f17928,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X7)))) | p(s(bool,X8)) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X5),s(bool,sK17(X5,X8,bool)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl124),
  inference(superposition,[],[f1548,f11080])).
fof(f17927,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X0),s(bool,sK18(X0,X3,bool)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl124),
  inference(superposition,[],[f1565,f11080])).
fof(f9473,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl93),
  inference(cnf_transformation,[],[f9473_D])).
fof(f9473_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f11079,plain,(
  ( ! [X263,X264] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X263) = s(fun(bool,bool),X264)) ) | ($spl60 | $spl73)),
  inference(forward_demodulation,[],[f11015,f10641])).
fof(f10641,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f246,f475,f2639,f1241])).
fof(f1241,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f979,f246])).
fof(f979,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f308,f561])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f475,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f449,f264])).
fof(f264,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f449,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f249,f246])).
fof(f249,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f11077,plain,(
  ( ! [X259,X260] : (s(fun(bool,bool),X259) = s(fun(bool,bool),X260) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X259,X260,bool,bool))))) ) | $spl122),
  inference(cnf_transformation,[],[f11077_D])).
fof(f11077_D,plain,(
  ( ! [X259,X260] : (s(fun(bool,bool),X259) = s(fun(bool,bool),X260) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X259,X260,bool,bool))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f9474,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl92),
  inference(cnf_transformation,[],[f9474_D])).
fof(f9474_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11017,plain,(
  ( ! [X259,X260] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X259) = s(fun(bool,bool),X260) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X259,X260,bool,bool))))) ) | $spl60),
  inference(duplicate_literal_removal,[],[f11010])).
fof(f11010,plain,(
  ( ! [X259,X260] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X259) = s(fun(bool,bool),X260) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X259,X260,bool,bool)))) | s(fun(bool,bool),X259) = s(fun(bool,bool),X260)) ) | $spl60),
  inference(superposition,[],[f2641,f1937])).
fof(f1937,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13)) ) | $spl60),
  inference(cnf_transformation,[],[f1937_D])).
fof(f1937_D,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10860,plain,(
  ( ! [X37,X41,X38,X42] : (s(bool,X38) != s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X38) | s(bool,X41) = s(bool,g(s(bool,X37))) | s(bool,'T_0') = s(bool,X37)) ) | $spl120),
  inference(cnf_transformation,[],[f10860_D])).
fof(f10860_D,plain,(
  ( ! [X37,X41,X38,X42] : (s(bool,X38) != s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X38) | s(bool,X41) = s(bool,g(s(bool,X37))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f10839,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,X62) != s(bool,f(s(bool,X61))) | s(fun(bool,bool),X63) = s(fun(bool,bool),X64) | s(bool,'T_0') = s(bool,sK19(X63,X64,bool,bool)) | s(bool,'T_0') = s(bool,X61) | s(bool,X65) = s(bool,g(s(bool,X61))) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f1572,f849])).
fof(f849,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f561,f561])).
fof(f1572,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK19(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f346,f561])).
fof(f10835,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X38) != s(bool,f(s(bool,X37))) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,'T_0') = s(bool,sK19(X39,X40,bool,bool)) | s(bool,'T_0') = s(bool,X37) | s(bool,X41) = s(bool,g(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f1572,f849])).
fof(f10859,plain,(
  ( ! [X30,X26,X29,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X25)) ) | $spl118),
  inference(cnf_transformation,[],[f10859_D])).
fof(f10859_D,plain,(
  ( ! [X30,X26,X29,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X25)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10834,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,X32) != s(bool,f(s(bool,X31))) | s(fun(bool,bool),X33) = s(fun(bool,bool),X34) | s(bool,'T_0') = s(bool,sK19(X33,X34,bool,bool)) | s(bool,'T_0') = s(bool,X31) | s(bool,X35) = s(bool,g(s(bool,X31))) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1572,f849])).
fof(f10833,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK19(X27,X28,bool,bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1572,f849])).
fof(f2637,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl70),
  inference(cnf_transformation,[],[f2637_D])).
fof(f2637_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f10533,plain,(
  ( ! [X669,X671] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(bool,'T_0') = s(bool,X671) | s(bool,X669) = s(bool,X671)) ) | $spl116),
  inference(cnf_transformation,[],[f10533_D])).
fof(f10533_D,plain,(
  ( ! [X669,X671] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(bool,'T_0') = s(bool,X671) | s(bool,X669) = s(bool,X671)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5620,plain,(
  ( ! [X668,X670,X667,X669,X671] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(fun(bool,bool),X667) = s(fun(bool,bool),X668) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X667,X668,bool,bool)))) | s(bool,X669) = s(bool,X670) | s(bool,'T_0') = s(bool,X670) | s(bool,X671) = s(bool,sK19(X667,X668,bool,bool)) | s(bool,'T_0') = s(bool,X671)) )),
  inference(superposition,[],[f1579,f849])).
fof(f5773,plain,(
  ( ! [X668,X670,X667,X669,X671] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(fun(bool,bool),X667) = s(fun(bool,bool),X668) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X667,X668,bool,bool)))) | s(bool,X670) = s(bool,sK19(X667,X668,bool,bool)) | s(bool,'T_0') = s(bool,X670) | s(bool,X669) = s(bool,X671) | s(bool,'T_0') = s(bool,X671)) )),
  inference(superposition,[],[f1579,f849])).
fof(f10532,plain,(
  ( ! [X652,X648,X650,X649,X651] : (s(fun(bool,bool),X648) = s(fun(bool,bool),X649) | s(bool,X650) = s(bool,sK19(X648,X649,bool,bool)) | s(bool,X650) = s(bool,X651) | s(bool,'T_0') = s(bool,X652) | s(bool,X651) = s(bool,X652) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X648,X649,bool,bool))))) ) | $spl114),
  inference(cnf_transformation,[],[f10532_D])).
fof(f10532_D,plain,(
  ( ! [X652,X648,X650,X649,X651] : (s(fun(bool,bool),X648) = s(fun(bool,bool),X649) | s(bool,X650) = s(bool,sK19(X648,X649,bool,bool)) | s(bool,X650) = s(bool,X651) | s(bool,'T_0') = s(bool,X652) | s(bool,X651) = s(bool,X652) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X648,X649,bool,bool))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2639,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl73),
  inference(cnf_transformation,[],[f2639_D])).
fof(f2639_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f6313,plain,(
  ( ! [X652,X648,X650,X649,X651] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X648) = s(fun(bool,bool),X649) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X648,X649,bool,bool)))) | s(bool,X650) = s(bool,X651) | s(bool,X651) = s(bool,sK19(X648,X649,bool,bool)) | s(bool,X650) = s(bool,X652) | s(bool,'T_0') = s(bool,X652)) )),
  inference(superposition,[],[f1579,f849])).
fof(f6853,plain,(
  ( ! [X652,X648,X650,X649,X651] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X648) = s(fun(bool,bool),X649) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X648,X649,bool,bool)))) | s(bool,X650) = s(bool,X651) | s(bool,X651) = s(bool,X652) | s(bool,'T_0') = s(bool,X652) | s(bool,X650) = s(bool,sK19(X648,X649,bool,bool))) )),
  inference(superposition,[],[f1579,f849])).
fof(f10509,plain,(
  ( ! [X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool))) ) | $spl100),
  inference(cnf_transformation,[],[f10509_D])).
fof(f10509_D,plain,(
  ( ! [X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f10477,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f1570,f254])).
fof(f1570,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK19(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f346,f561])).
fof(f10508,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl98),
  inference(cnf_transformation,[],[f10508_D])).
fof(f10508_D,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f10478,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1570,f561])).
fof(f10528,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X55)) ) | $spl108),
  inference(cnf_transformation,[],[f10528_D])).
fof(f10528_D,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X55)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10479,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,X10) = s(bool,sK19(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f1570,f561])).
fof(f10480,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X11) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,X14) = s(bool,sK19(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1570,f561])).
fof(f10481,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,X17) = s(bool,sK19(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1570,f561])).
fof(f10531,plain,(
  ( ! [X30,X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,X25) = s(bool,X30)) ) | $spl112),
  inference(cnf_transformation,[],[f10531_D])).
fof(f10531_D,plain,(
  ( ! [X30,X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,X25) = s(bool,X30)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10482,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,X19) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,X22) = s(bool,sK19(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10483,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,X28) = s(bool,sK19(X26,X27,bool,bool)) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X29) | s(bool,X25) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10527,plain,(
  ( ! [X59,X60] : (s(bool,X59) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,X59) = s(bool,X60)) ) | $spl106),
  inference(cnf_transformation,[],[f10527_D])).
fof(f10527_D,plain,(
  ( ! [X59,X60] : (s(bool,X59) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,X59) = s(bool,X60)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f10484,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,X31) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,X34) = s(bool,sK19(X32,X33,bool,bool)) | s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10526,plain,(
  ( ! [X61,X66,X65] : (s(bool,X61) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) = s(bool,X66) | s(bool,X61) = s(bool,X65)) ) | $spl104),
  inference(cnf_transformation,[],[f10526_D])).
fof(f10526_D,plain,(
  ( ! [X61,X66,X65] : (s(bool,X61) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,X65) = s(bool,X66) | s(bool,X61) = s(bool,X65)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f10485,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X37) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,X40) = s(bool,sK19(X38,X39,bool,bool)) | s(bool,'T_0') = s(bool,X40) | s(bool,X37) = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10522,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,X47) = s(bool,g(s(bool,'T_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f10522_D])).
fof(f10522_D,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,X47) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f10486,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,X45) = s(bool,sK19(X43,X44,bool,bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,g(s(bool,'T_0'))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10487,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,X51) = s(bool,sK19(X49,X50,bool,bool)) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f1570,f849])).
fof(f10488,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,X55) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X56) = s(fun(bool,bool),X57) | s(bool,X58) = s(bool,sK19(X56,X57,bool,bool)) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,g(s(bool,'T_0'))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f1570,f849])).
fof(f10489,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,X61) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X62) = s(fun(bool,bool),X63) | s(bool,X64) = s(bool,sK19(X62,X63,bool,bool)) | s(bool,'T_0') = s(bool,X64) | s(bool,X61) = s(bool,X65) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f1570,f849])).
fof(f10523,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,X51) = s(bool,sK19(X49,X50,bool,bool)) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,X53) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X52) = s(bool,g(s(bool,'T_0')))) ) | $spl73),
  inference(subsumption_resolution,[],[f10487,f2669])).
fof(f2669,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f449,f2639,f607])).
fof(f607,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f604,f449])).
fof(f604,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f305,f254])).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f131])).
fof(f10521,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,X45) = s(bool,sK19(X43,X44,bool,bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,g(s(bool,'T_0'))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | $spl73),
  inference(subsumption_resolution,[],[f10486,f2669])).
fof(f10518,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK19(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl73),
  inference(duplicate_literal_removal,[],[f10517])).
fof(f10517,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK19(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) | $spl73),
  inference(condensation,[],[f10516])).
fof(f10516,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,X28) = s(bool,sK19(X26,X27,bool,bool)) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X29) | s(bool,X25) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl73),
  inference(subsumption_resolution,[],[f10483,f4127])).
fof(f4127,plain,(
  ( ! [X30,X31] : (s(bool,X30) = s(bool,f(s(bool,'T_0'))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl73),
  inference(subsumption_resolution,[],[f4111,f246])).
fof(f4111,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,f(s(bool,'T_0'))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl73),
  inference(superposition,[],[f1241,f2669])).
fof(f10515,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK19(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl73),
  inference(duplicate_literal_removal,[],[f10514])).
fof(f10514,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK19(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl73),
  inference(condensation,[],[f10513])).
fof(f10513,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,X22) = s(bool,sK19(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X24)) ) | $spl73),
  inference(subsumption_resolution,[],[f10482,f4127])).
fof(f10512,plain,(
  ( ! [X17,X15,X18,X16] : (s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,X17) = s(bool,sK19(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) ) | $spl73),
  inference(subsumption_resolution,[],[f10481,f2669])).
fof(f10507,plain,(
  ( ! [X6,X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl73),
  inference(subsumption_resolution,[],[f10478,f2669])).
fof(f9477,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,X38) | s(bool,X36) = s(bool,X38) | s(bool,X37) = s(bool,sK19(X34,X35,bool,bool)) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X34,X35,bool,bool))))) ) | $spl96),
  inference(cnf_transformation,[],[f9477_D])).
fof(f9477_D,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,X38) | s(bool,X36) = s(bool,X38) | s(bool,X37) = s(bool,sK19(X34,X35,bool,bool)) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X34,X35,bool,bool))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9358,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X39,X40,bool,bool)))) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,sK19(X39,X40,bool,bool))) )),
  inference(superposition,[],[f2641,f849])).
fof(f9357,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X34,X35,bool,bool)))) | s(bool,X36) = s(bool,X37) | s(bool,X37) = s(bool,sK19(X34,X35,bool,bool)) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f2641,f849])).
fof(f9472,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK19(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X2,X3,bool,bool))))) ) | $spl90),
  inference(cnf_transformation,[],[f9472_D])).
fof(f9472_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK19(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X2,X3,bool,bool))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f9475,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,X17)) ) | $spl94),
  inference(cnf_transformation,[],[f9475_D])).
fof(f9475_D,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,X17)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f9354,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,g(s(bool,X21))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X19,X20,bool,bool)))) | s(bool,X22) = s(bool,sK19(X19,X20,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f2641,f849])).
fof(f9353,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X14,X15,bool,bool)))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,sK19(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f2641,f849])).
fof(f9352,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK19(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2641,f561])).
fof(f9349,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,f(s(bool,sK19(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK19(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2641,f561])).
fof(f8473,plain,(
  ( ! [X218,X216,X215,X213,X217] : (sP1(X213,X215) | s(bool,'T_0') = s(bool,X218) | s(bool,X217) = s(bool,X218) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X216) | s(bool,X215) = s(bool,X216)) ) | $spl80),
  inference(cnf_transformation,[],[f8473_D])).
fof(f8473_D,plain,(
  ( ! [X218,X216,X215,X213,X217] : (sP1(X213,X215) | s(bool,'T_0') = s(bool,X218) | s(bool,X217) = s(bool,X218) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X216) | s(bool,X215) = s(bool,X216)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1816,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(cnf_transformation,[],[f1816_D])).
fof(f1816_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8605,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,X214)) | sP1(X213,X215) | s(bool,X215) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X213) = s(bool,X217) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(subsumption_resolution,[],[f6354,f264])).
fof(f6354,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,X214)) | sP1(X213,X215) | s(bool,X215) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X213) = s(bool,X217) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f1245,f849])).
fof(f1245,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f983,f246])).
fof(f983,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f333,f561])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f8471,plain,(
  ( ! [X212,X210,X211,X209,X207] : (sP1(X209,X207) | s(bool,'T_0') = s(bool,X212) | s(bool,X211) = s(bool,X212) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,X209) = s(bool,X210)) ) | $spl78),
  inference(cnf_transformation,[],[f8471_D])).
fof(f8471_D,plain,(
  ( ! [X212,X210,X211,X209,X207] : (sP1(X209,X207) | s(bool,'T_0') = s(bool,X212) | s(bool,X211) = s(bool,X212) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X210) | s(bool,X209) = s(bool,X210)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8604,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X209,X207) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f6353,f264])).
fof(f6353,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X209,X207) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f1244,f849])).
fof(f1244,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f982,f246])).
fof(f982,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f330,f561])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f8469,plain,(
  ( ! [X194,X192,X195] : (sP1(X192,'T_0') | s(bool,'T_0') = s(bool,X195) | s(bool,X194) = s(bool,X195) | s(bool,X192) = s(bool,X194)) ) | $spl76),
  inference(cnf_transformation,[],[f8469_D])).
fof(f8469_D,plain,(
  ( ! [X194,X192,X195] : (sP1(X192,'T_0') | s(bool,'T_0') = s(bool,X195) | s(bool,X194) = s(bool,X195) | s(bool,X192) = s(bool,X194)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8603,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X193)) | sP1(X192,'T_0') | s(bool,X192) = s(bool,X194) | s(bool,X194) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) ) | $spl14),
  inference(subsumption_resolution,[],[f6350,f264])).
fof(f6350,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X193)) | sP1(X192,'T_0') | s(bool,X192) = s(bool,X194) | s(bool,X194) = s(bool,X195) | s(bool,'T_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X193)) ) | $spl14),
  inference(superposition,[],[f1191,f849])).
fof(f8467,plain,(
  ( ! [X191,X190,X188] : (sP1('T_0',X188) | s(bool,'T_0') = s(bool,X191) | s(bool,X190) = s(bool,X191) | s(bool,X188) = s(bool,X190)) ) | $spl74),
  inference(cnf_transformation,[],[f8467_D])).
fof(f8467_D,plain,(
  ( ! [X191,X190,X188] : (sP1('T_0',X188) | s(bool,'T_0') = s(bool,X191) | s(bool,X190) = s(bool,X191) | s(bool,X188) = s(bool,X190)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8602,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X189)) | sP1('T_0',X188) | s(bool,X188) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | $spl10),
  inference(subsumption_resolution,[],[f6349,f264])).
fof(f6349,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X189)) | sP1('T_0',X188) | s(bool,X188) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191) | s(bool,'T_0') = s(bool,X189)) ) | $spl10),
  inference(superposition,[],[f1187,f849])).
fof(f8505,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | s(bool,'T_0') = s(bool,X187) | s(bool,X186) = s(bool,X187) | s(bool,X183) = s(bool,X186) | s(bool,'T_0') = s(bool,X185) | s(bool,X184) = s(bool,X185)) ) | $spl88),
  inference(cnf_transformation,[],[f8505_D])).
fof(f8505_D,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | s(bool,'T_0') = s(bool,X187) | s(bool,X186) = s(bool,X187) | s(bool,X183) = s(bool,X186) | s(bool,'T_0') = s(bool,X185) | s(bool,X184) = s(bool,X185)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8504,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | p(s(bool,X182)) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,X183) = s(bool,X186) | s(bool,X186) = s(bool,X187) | s(bool,'T_0') = s(bool,X187)) )),
  inference(subsumption_resolution,[],[f5961,f264])).
fof(f5961,plain,(
  ( ! [X187,X185,X182,X186,X184,X183] : (s(bool,X183) != s(bool,X184) | p(s(bool,X182)) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,X183) = s(bool,X186) | s(bool,'T_0') = s(bool,X182) | s(bool,X186) = s(bool,X187) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f937,f849])).
fof(f937,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f264,f561])).
fof(f8500,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | $spl86),
  inference(cnf_transformation,[],[f8500_D])).
fof(f8500_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8499,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5928,f264])).
fof(f5928,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f264,f849])).
fof(f8497,plain,(
  ( ! [X668,X670,X667,X671] : (s(fun(bool,bool),X667) = s(fun(bool,bool),X668) | s(bool,'T_0') = s(bool,X671) | s(bool,X670) = s(bool,X671) | s(bool,X670) = s(bool,sK19(X667,X668,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X667,X668,bool,bool))))) ) | $spl82),
  inference(cnf_transformation,[],[f8497_D])).
fof(f8497_D,plain,(
  ( ! [X668,X670,X667,X671] : (s(fun(bool,bool),X667) = s(fun(bool,bool),X668) | s(bool,'T_0') = s(bool,X671) | s(bool,X670) = s(bool,X671) | s(bool,X670) = s(bool,sK19(X667,X668,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X667,X668,bool,bool))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8498,plain,(
  ( ! [X669] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(bool,'T_0') = s(bool,X669)) ) | $spl84),
  inference(cnf_transformation,[],[f8498_D])).
fof(f8498_D,plain,(
  ( ! [X669] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(bool,'T_0') = s(bool,X669)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5926,plain,(
  ( ! [X668,X670,X667,X669,X671] : (s(bool,'F_0') != s(bool,f(s(bool,X669))) | s(fun(bool,bool),X667) = s(fun(bool,bool),X668) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X667,X668,bool,bool)))) | s(bool,X670) = s(bool,sK19(X667,X668,bool,bool)) | s(bool,'T_0') = s(bool,X669) | s(bool,X670) = s(bool,X671) | s(bool,'T_0') = s(bool,X671)) )),
  inference(superposition,[],[f1579,f849])).
fof(f8472,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,X214)) | sP1(X213,X215) | s(bool,X215) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X213) = s(bool,X217) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(subsumption_resolution,[],[f5814,f264])).
fof(f5814,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (p(s(bool,X214)) | sP1(X213,X215) | s(bool,X215) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X214) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f1245,f849])).
fof(f8470,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X209,X207) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f5813,f264])).
fof(f5813,plain,(
  ( ! [X212,X210,X208,X211,X209,X207] : (p(s(bool,X208)) | sP1(X209,X207) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X207) = s(bool,X211) | s(bool,'T_0') = s(bool,X208) | s(bool,X211) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f1244,f849])).
fof(f8468,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X193)) | sP1(X192,'T_0') | s(bool,X192) = s(bool,X194) | s(bool,X194) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) ) | $spl14),
  inference(subsumption_resolution,[],[f5810,f264])).
fof(f5810,plain,(
  ( ! [X194,X192,X195,X193] : (p(s(bool,X193)) | sP1(X192,'T_0') | s(bool,X192) = s(bool,X194) | s(bool,'T_0') = s(bool,X193) | s(bool,X194) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) ) | $spl14),
  inference(superposition,[],[f1191,f849])).
fof(f8466,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X189)) | sP1('T_0',X188) | s(bool,X188) = s(bool,X190) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | $spl10),
  inference(subsumption_resolution,[],[f5809,f264])).
fof(f5809,plain,(
  ( ! [X191,X189,X190,X188] : (p(s(bool,X189)) | sP1('T_0',X188) | s(bool,X188) = s(bool,X190) | s(bool,'T_0') = s(bool,X189) | s(bool,X190) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | $spl10),
  inference(superposition,[],[f1187,f849])).
fof(f2627,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK19(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1579,f561])).
fof(f2626,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK19(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK19(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1579,f561])).
fof(f1941,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl66),
  inference(cnf_transformation,[],[f1941_D])).
fof(f1941_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1943,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl69),
  inference(cnf_transformation,[],[f1943_D])).
fof(f1943_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f1906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1574,f561])).
fof(f1574,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK19(X16,X17,X18,bool))) )),
  inference(superposition,[],[f346,f254])).
fof(f1907,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK19(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1574,f561])).
fof(f1936,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl58),
  inference(cnf_transformation,[],[f1936_D])).
fof(f1936_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1940,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl64),
  inference(cnf_transformation,[],[f1940_D])).
fof(f1940_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1908,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK19(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1574,f561])).
fof(f1909,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK19(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1574,f561])).
fof(f1939,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl63),
  inference(cnf_transformation,[],[f1939_D])).
fof(f1939_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f1910,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK19(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1574,f254])).
fof(f1934,plain,(
  $false | $spl50),
  inference(subsumption_resolution,[],[f1914,f246])).
fof(f1914,plain,(
  ~p(s(bool,'T_0')) | $spl50),
  inference(backward_demodulation,[],[f1911,f449])).
fof(f1911,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl50),
  inference(global_subsumption,[],[f245,f244,f243,f242,f246,f254,f265,f264,f293,f296,f295,f294,f303,f308,f305,f315,f314,f450,f333,f330,f329,f328,f346,f348,f347,f361,f360,f359,f358,f357,f356,f355,f365,f364,f363,f362,f369,f368,f367,f366,f373,f372,f371,f370,f377,f376,f375,f374,f431,f381,f380,f379,f378,f388,f387,f386,f385,f394,f391,f390,f389,f401,f400,f399,f430,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f418,f417,f416,f429,f414,f413,f412,f422,f421,f420,f419,f426,f425,f424,f423,f455,f456,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f482,f483,f485,f491,f496,f497,f500,f503,f506,f513,f514,f516,f523,f527,f528,f531,f537,f561,f607,f616,f615,f626,f624,f638,f644,f654,f657,f662,f666,f674,f680,f703,f710,f716,f739,f746,f752,f784,f798,f819,f821,f849,f937,f1186,f1190,f1193,f1196,f1199,f1202,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1240,f1241,f1242,f1243,f1244,f1245,f1246,f1247,f1248,f1249,f1250,f1251,f1252,f1253,f1254,f1255,f1256,f1257,f1258,f1259,f1260,f1261,f1267,f1269,f1271,f1273,f1275,f1277,f1279,f1281,f1283,f1285,f1287,f1289,f1291,f1293,f1295,f1297,f1299,f1301,f302,f300,f449,f269,f268,f1528,f1529,f1538,f1540,f1542,f1548,f1555,f1557,f1559,f1565,f1570,f1572,f1574,f1575,f1577,f1579,f1587,f1590,f1593,f1594,f1598,f1601,f1603,f1609,f1612,f1622,f1643,f1647,f1660,f1665,f1682,f1684,f1685,f1688,f1689,f1691,f1692,f1697,f1699,f1712,f1713,f1714,f1716,f1745,f1749,f1752,f1754,f1756,f1774,f1814,f1817,f1819,f1820,f1826,f1827,f1830,f1831,f1832,f1836,f1837,f1838,f1840,f1842,f1846,f1849,f1851,f1853,f1863,f1868,f1871,f1873,f1875,f1894,f1896,f1898,f1899,f1816,f1906])).
fof(f1899,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f1891,f449])).
fof(f1891,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'F_0')) | sP1(X19,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f1245,f254])).
fof(f1875,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X79),s('type.num.num',X80))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X76),'const.list.APPEND_2'(s('type.list.list'(X76),X78),s('type.list.list'(X76),X77))))),s('type.num.num',X79))) | s('type.num.num',X80) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X76),X77))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X76),X78)))))) )),
  inference(superposition,[],[f821,f1594])).
fof(f1873,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X70),s('type.num.num',X69))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X69),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X66),'const.list.APPEND_2'(s('type.list.list'(X66),X68),s('type.list.list'(X66),X67))))))) | s('type.num.num',X70) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X66),X67))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X66),X68)))))) )),
  inference(superposition,[],[f821,f1594])).
fof(f1871,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s('type.num.num',X59) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X56),'const.list.APPEND_2'(s('type.list.list'(X56),X58),s('type.list.list'(X56),X57))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X60))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X56),X57))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X56),X58))))),s('type.num.num',X60)))) )),
  inference(superposition,[],[f348,f1594])).
fof(f1868,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X44))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X41),'const.list.APPEND_2'(s('type.list.list'(X41),X43),s('type.list.list'(X41),X42))))),s('type.num.num',X44))) | s('type.num.num',X45) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X41),X42))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X41),X43)))))) )),
  inference(superposition,[],[f347,f1594])).
fof(f1853,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X93),s('type.num.num',X94))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X90),X92))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X90),X91))))),s('type.num.num',X93))) | s('type.num.num',X94) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X90),'const.list.APPEND_2'(s('type.list.list'(X90),X91),s('type.list.list'(X90),X92)))))) )),
  inference(superposition,[],[f821,f1594])).
fof(f1851,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X84),s('type.num.num',X83))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X83),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X80),X82))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X80),X81))))))) | s('type.num.num',X84) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X80),'const.list.APPEND_2'(s('type.list.list'(X80),X81),s('type.list.list'(X80),X82)))))) )),
  inference(superposition,[],[f821,f1594])).
fof(f1849,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.num.num',X73) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X70),X72))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X70),X71))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X73),s('type.num.num',X74))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X70),'const.list.APPEND_2'(s('type.list.list'(X70),X71),s('type.list.list'(X70),X72))))),s('type.num.num',X74)))) )),
  inference(superposition,[],[f348,f1594])).
fof(f1846,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X58))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X55),X57))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X55),X56))))),s('type.num.num',X58))) | s('type.num.num',X59) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X55),'const.list.APPEND_2'(s('type.list.list'(X55),X56),s('type.list.list'(X55),X57)))))) )),
  inference(superposition,[],[f347,f1594])).
fof(f1842,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X35),X38))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X35),X37))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X35),X36))))) | s('type.list.list'(X35),X39) != s('type.list.list'(X35),X40) | s('type.list.list'(X35),X38) != s('type.list.list'(X35),'const.list.APPEND_2'(s('type.list.list'(X35),X36),s('type.list.list'(X35),X37))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X35),X39))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X35),X40))) | s('type.list.list'(X35),'const.list.APPEND_2'(s('type.list.list'(X35),X38),s('type.list.list'(X35),X39))) = s('type.list.list'(X35),'const.list.APPEND_2'(s('type.list.list'(X35),'const.list.APPEND_2'(s('type.list.list'(X35),X36),s('type.list.list'(X35),X37))),s('type.list.list'(X35),X40)))) )),
  inference(superposition,[],[f296,f1594])).
fof(f1840,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X26))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X25))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X24))))) | s('type.list.list'(X23),'const.list.APPEND_2'(s('type.list.list'(X23),X26),s('type.list.list'(X23),X27))) != s('type.list.list'(X23),'const.list.APPEND_2'(s('type.list.list'(X23),'const.list.APPEND_2'(s('type.list.list'(X23),X24),s('type.list.list'(X23),X25))),s('type.list.list'(X23),X28))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X27))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X28))) | s('type.list.list'(X23),X27) = s('type.list.list'(X23),X28)) )),
  inference(superposition,[],[f295,f1594])).
fof(f1838,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X14))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X13))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X12))))) | s('type.list.list'(X11),'const.list.APPEND_2'(s('type.list.list'(X11),X14),s('type.list.list'(X11),X15))) != s('type.list.list'(X11),'const.list.APPEND_2'(s('type.list.list'(X11),'const.list.APPEND_2'(s('type.list.list'(X11),X12),s('type.list.list'(X11),X13))),s('type.list.list'(X11),X16))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X11),X16))) | s('type.list.list'(X11),X14) = s('type.list.list'(X11),'const.list.APPEND_2'(s('type.list.list'(X11),X12),s('type.list.list'(X11),X13)))) )),
  inference(superposition,[],[f294,f1594])).
fof(f1837,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),'const.list.APPEND_2'(s('type.list.list'(X7),'const.list.APPEND_2'(s('type.list.list'(X7),X8),s('type.list.list'(X7),X9))),s('type.list.list'(X7),X10))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X9))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X10)))))) )),
  inference(superposition,[],[f293,f1594])).
fof(f1836,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X6),s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X4),s('type.list.list'(X3),X5))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X5))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X4)))))))) )),
  inference(superposition,[],[f293,f1594])).
fof(f1831,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),X5),s('type.list.list'(X4),X6))),s('type.list.list'(X4),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X5))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X6)))))))) )),
  inference(superposition,[],[f1594,f293])).
fof(f1830,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))))) )),
  inference(superposition,[],[f1594,f1594])).
fof(f1827,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),X7),s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),X5),s('type.list.list'(X4),X6))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X5))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X6))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X4),X7)))))) )),
  inference(superposition,[],[f1594,f293])).
fof(f1826,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) )),
  inference(superposition,[],[f1594,f1594])).
fof(f1820,plain,(
  ( ! [X21,X19,X20] : (sP1(X20,X19) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f1811,f449])).
fof(f1811,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'F_0')) | sP1(X20,X19) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f1244,f254])).
fof(f1774,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1240,f254])).
fof(f1756,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num',X69))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X66),s('type.num.num',X65))),s('type.num.num',X67))),s('type.num.num',X68))) | s('type.num.num',X69) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X66))),s('type.num.num',X67)))) )),
  inference(superposition,[],[f821,f1528])).
fof(f1754,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X58))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X58),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X55))),s('type.num.num',X57))))) | s('type.num.num',X59) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56))),s('type.num.num',X57)))) )),
  inference(superposition,[],[f821,f1528])).
fof(f1752,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s('type.num.num',X48) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X45))),s('type.num.num',X47))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',X49))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))),s('type.num.num',X47))),s('type.num.num',X49)))) )),
  inference(superposition,[],[f348,f1528])).
fof(f1749,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X33))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X30))),s('type.num.num',X32))),s('type.num.num',X33))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X31))),s('type.num.num',X32)))) )),
  inference(superposition,[],[f347,f1528])).
fof(f1745,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))),s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num',X17)) )),
  inference(superposition,[],[f821,f1528])).
fof(f1716,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f1528,f303])).
fof(f1714,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))))),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X10))),s('type.num.num',X12))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f1528,f1528])).
fof(f1713,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))))),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f1528,f293])).
fof(f1712,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))),s('type.num.num',X3)))) )),
  inference(unit_resulting_resolution,[],[f1528,f348])).
fof(f1699,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))),s('type.num.num',X9))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))))) | s('type.num.num',X9) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))) )),
  inference(superposition,[],[f821,f293])).
fof(f1697,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(superposition,[],[f821,f303])).
fof(f1692,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),X14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),X15))))),s('type.num.num',X16))) | s('type.num.num',X17) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),'const.list.APPEND_2'(s('type.list.list'(X13),X14),s('type.list.list'(X13),X15)))))) )),
  inference(superposition,[],[f821,f293])).
fof(f1691,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),'const.list.APPEND_2'(s('type.list.list'(X8),X9),s('type.list.list'(X8),X10))))),s('type.num.num',X11))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X9))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X10)))))) )),
  inference(superposition,[],[f821,f293])).
fof(f1689,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f821,f303])).
fof(f1688,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))))) | s('type.num.num',X9) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))) )),
  inference(superposition,[],[f821,f293])).
fof(f1685,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X16))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),X14))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),X15))))))) | s('type.num.num',X17) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X13),'const.list.APPEND_2'(s('type.list.list'(X13),X14),s('type.list.list'(X13),X15)))))) )),
  inference(superposition,[],[f821,f293])).
fof(f1684,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X11))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),'const.list.APPEND_2'(s('type.list.list'(X8),X9),s('type.list.list'(X8),X10))))))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X9))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X10)))))) )),
  inference(superposition,[],[f821,f293])).
fof(f1682,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f821,f303])).
fof(f1665,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) )),
  inference(trivial_inequality_removal,[],[f1664])).
fof(f1664,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),X3) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2)))) )),
  inference(equality_resolution,[],[f296])).
fof(f1660,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),X4) != s('type.list.list'(X0),X5) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X5))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X5))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X4)))) )),
  inference(superposition,[],[f296,f293])).
fof(f1643,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X5))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X5))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X4))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) | s('type.list.list'(X0),X4) = s('type.list.list'(X0),X5)) )),
  inference(superposition,[],[f295,f293])).
fof(f1622,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X5))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X5))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X4))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) | s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) )),
  inference(superposition,[],[f294,f293])).
fof(f1612,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num',X40) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X37),'const.list.APPEND_2'(s('type.list.list'(X37),X38),s('type.list.list'(X37),X39))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X37),X38))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X37),X39))))),s('type.num.num',X41)))) )),
  inference(superposition,[],[f348,f293])).
fof(f1609,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X25))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X22),'const.list.APPEND_2'(s('type.list.list'(X22),X23),s('type.list.list'(X22),X24))))),s('type.num.num',X25))) | s('type.num.num',X26) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X22),X23))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X22),X24)))))) )),
  inference(superposition,[],[f347,f293])).
fof(f1601,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s('type.num.num',X26) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X24))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),X25))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X23),'const.list.APPEND_2'(s('type.list.list'(X23),X24),s('type.list.list'(X23),X25))))),s('type.num.num',X27)))) )),
  inference(superposition,[],[f348,f293])).
fof(f1598,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X11))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X9))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),X10))))),s('type.num.num',X11))) | s('type.num.num',X12) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X8),'const.list.APPEND_2'(s('type.list.list'(X8),X9),s('type.list.list'(X8),X10)))))) )),
  inference(superposition,[],[f347,f293])).
fof(f1593,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))))) )),
  inference(superposition,[],[f293,f293])).
fof(f1590,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) )),
  inference(superposition,[],[f293,f293])).
fof(f1587,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))),s('type.num.num',X3)))) )),
  inference(unit_resulting_resolution,[],[f293,f348])).
fof(f1529,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f348,f303])).
fof(f1528,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f303,f348])).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f269,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f1261,plain,(
  ( ! [X101,X99,X100,X98] : (~sP8(X98,X99,X100) | sP9(X98,X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f999,f246])).
fof(f999,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | ~sP8(X98,X99,X100) | sP9(X98,X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f430,f561])).
fof(f1260,plain,(
  ( ! [X94,X97,X95,X96] : (~sP10(X94,X95,X96) | sP11(X94,X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f998,f246])).
fof(f998,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | ~sP10(X94,X95,X96) | sP11(X94,X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f429,f561])).
fof(f1259,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X91)) | p(s(bool,X92)) | ~sP11(X90,X92,X91) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f997,f246])).
fof(f997,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X91)) | p(s(bool,X92)) | ~sP11(X90,X92,X91) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f426,f561])).
fof(f1258,plain,(
  ( ! [X88,X87,X89,X86] : (sP11(X86,X87,X88) | ~p(s(bool,X87)) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f996,f246])).
fof(f996,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X87)) | sP11(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f425,f561])).
fof(f1257,plain,(
  ( ! [X85,X83,X84,X82] : (sP11(X82,X84,X83) | p(s(bool,X83)) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f995,f246])).
fof(f995,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X83)) | sP11(X82,X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f424,f561])).
fof(f1256,plain,(
  ( ! [X80,X78,X81,X79] : (sP11(X80,X79,X78) | p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f994,f246])).
fof(f994,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | sP11(X80,X79,X78) | p(s(bool,X80)) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f423,f561])).
fof(f1255,plain,(
  ( ! [X76,X74,X77,X75] : (~sP10(X76,X75,X74) | ~p(s(bool,X76)) | p(s(bool,X75)) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f993,f246])).
fof(f993,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,X74) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f419,f561])).
fof(f1254,plain,(
  ( ! [X70,X72,X71,X73] : (~sP9(X70,X72,X71) | p(s(bool,X72)) | p(s(bool,X71)) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f992,f246])).
fof(f992,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9(X70,X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f409,f561])).
fof(f1253,plain,(
  ( ! [X68,X66,X69,X67] : (sP9(X66,X67,X68) | ~p(s(bool,X67)) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f991,f246])).
fof(f991,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X67)) | sP9(X66,X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f408,f561])).
fof(f1252,plain,(
  ( ! [X64,X62,X65,X63] : (sP9(X62,X64,X63) | ~p(s(bool,X63)) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f990,f246])).
fof(f990,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP9(X62,X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f407,f561])).
fof(f1251,plain,(
  ( ! [X61,X59,X60,X58] : (~sP8(X58,X59,X60) | p(s(bool,X60)) | p(s(bool,X59)) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f989,f246])).
fof(f989,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP8(X58,X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f402,f561])).
fof(f1250,plain,(
  ( ! [X57,X54,X56,X55] : (~sP7(X56,X55,X54) | p(s(bool,X56)) | ~p(s(bool,X55)) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f988,f246])).
fof(f988,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,X54) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f394,f561])).
fof(f1249,plain,(
  ( ! [X52,X50,X53,X51] : (sP7(X50,X51,X52) | ~p(s(bool,X52)) | ~p(s(bool,X51)) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f987,f246])).
fof(f987,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | sP7(X50,X51,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f391,f561])).
fof(f1248,plain,(
  ( ! [X47,X48,X46,X49] : (~sP6(X47,X46,X48) | p(s(bool,X48)) | ~p(s(bool,X47)) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f986,f246])).
fof(f986,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,X46,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f385,f561])).
fof(f1247,plain,(
  ( ! [X45,X43,X44,X42] : (~sP3(X43,X42,X44) | p(s(bool,X44)) | ~p(s(bool,X43)) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f985,f246])).
fof(f985,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f366,f561])).
fof(f1246,plain,(
  ( ! [X39,X41,X38,X40] : (~sP4(X39,X38,X40) | p(s(bool,X40)) | ~p(s(bool,X39)) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f984,f246])).
fof(f984,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,X38,X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f362,f561])).
fof(f1243,plain,(
  ( ! [X30,X31,X29] : (sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f981,f246])).
fof(f981,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f315,f561])).
fof(f1242,plain,(
  ( ! [X28,X26,X27] : (sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f980,f246])).
fof(f980,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f314,f561])).
fof(f1240,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f975,f246])).
fof(f975,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f268,f561])).
fof(f819,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f347,f303])).
fof(f798,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f463,f460,f458,f449,f360])).
fof(f784,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f465,f460,f457,f246,f359])).
fof(f752,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f746,f414])).
fof(f746,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f449,f246,f419])).
fof(f739,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f474,f246,f449,f418])).
fof(f716,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f710,f397])).
fof(f710,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f449,f246,f402])).
fof(f703,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f471,f449,f449,f401])).
fof(f680,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f674,f378])).
fof(f674,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f246,f246,f385])).
fof(f666,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f466,f246,f246,f381])).
fof(f662,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f657,f355])).
fof(f657,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f449,f449,f370])).
fof(f654,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f644,f356])).
fof(f644,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f246,f246,f366])).
fof(f638,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f246,f246,f362])).
fof(f624,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f475,f376])).
fof(f626,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f246,f475,f376])).
fof(f615,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f449,f475,f374])).
fof(f616,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f475,f374])).
fof(f537,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f246,f425])).
fof(f531,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f449,f246,f424])).
fof(f528,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f449,f413])).
fof(f527,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f449,f449,f412])).
fof(f523,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f246,f408])).
fof(f516,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f246,f407])).
fof(f514,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f449,f396])).
fof(f513,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f449,f395])).
fof(f506,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f449,f449,f390])).
fof(f503,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f449,f449,f389])).
fof(f500,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f449,f380])).
fof(f497,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f246,f449,f379])).
fof(f496,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f449,f450])).
fof(f491,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f449,f246,f333])).
fof(f485,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f246,f330])).
fof(f483,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f246,f329])).
fof(f482,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f449,f449,f328])).
fof(f474,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f449,f422])).
fof(f473,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f421])).
fof(f472,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f449,f420])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f449,f405])).
fof(f470,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f404])).
fof(f469,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f403])).
fof(f468,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f449,f388])).
fof(f467,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f449,f387])).
fof(f466,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f386])).
fof(f465,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f246,f373])).
fof(f464,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f372])).
fof(f463,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f371])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f449,f369])).
fof(f461,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f449,f368])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f367])).
fof(f459,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f449,f365])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f449,f364])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f363])).
fof(f456,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f315])).
fof(f455,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f314])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f139,f156,f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X23,X24,X13] : ((p(s(bool,X13)) <=> (p(s(bool,X24)) => p(s(bool,X23)))) <=> ((p(s(bool,X13)) | p(s(bool,X24))) & (p(s(bool,X13)) | ~p(s(bool,X23))) & (~p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X13)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f112,f153,f152])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X23,X24,X13] : ((p(s(bool,X13)) <=> (p(s(bool,X24)) | p(s(bool,X23)))) <=> ((p(s(bool,X13)) | ~p(s(bool,X24))) & (p(s(bool,X13)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X13)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f231])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f110,f150,f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X23,X24,X13] : ((p(s(bool,X13)) <=> (p(s(bool,X24)) & p(s(bool,X23)))) <=> ((p(s(bool,X13)) | ~p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X13))) & (p(s(bool,X23)) | ~p(s(bool,X13)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.sat.sat_23_prove_ltb')).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f149])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f223])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f108,f147,f146,f145,f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f145,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f146,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f147,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X23,X24,X13] : ((p(s(bool,X13)) <=> s(bool,X23) = s(bool,X24)) <=> ((p(s(bool,X13)) | p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X13)) | ~p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X13))) & (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X13)))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f144])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f214])).
fof(f214,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f145])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f215])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f212])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f212])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f204])).
fof(f328,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f197])).
fof(f329,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f450,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f313,f449])).
fof(f313,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f315,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f296,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) | s('type.list.list'(X0),X3) != s('type.list.list'(X0),X4) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4))) | s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f1933,plain,(
  $false | $spl50),
  inference(subsumption_resolution,[],[f1932,f1816])).
fof(f1932,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl50),
  inference(subsumption_resolution,[],[f1912,f1911])).
fof(f1912,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl50),
  inference(backward_demodulation,[],[f1911,f268])).
fof(f1897,plain,(
  ( ! [X14,X13,X11] : (sP1(X11,X13) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X14)) ) | $spl56),
  inference(cnf_transformation,[],[f1897_D])).
fof(f1897_D,plain,(
  ( ! [X14,X13,X11] : (sP1(X11,X13) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X14)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1898,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | sP1(X15,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1890,f264])).
fof(f1890,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | sP1(X15,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1245,f561])).
fof(f1896,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X12)) | sP1(X11,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1889,f264])).
fof(f1889,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X12)) | sP1(X11,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1245,f561])).
fof(f1895,plain,(
  ( ! [X2,X1] : (sP1('F_0',X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl54),
  inference(cnf_transformation,[],[f1895_D])).
fof(f1895_D,plain,(
  ( ! [X2,X1] : (sP1('F_0',X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1894,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | sP1('F_0',X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1886,f264])).
fof(f1886,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | sP1('F_0',X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1245,f254])).
fof(f1818,plain,(
  ( ! [X14,X13,X11] : (sP1(X13,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X14)) ) | $spl52),
  inference(cnf_transformation,[],[f1818_D])).
fof(f1818_D,plain,(
  ( ! [X14,X13,X11] : (sP1(X13,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X14)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1819,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | sP1(X17,X15) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1810,f264])).
fof(f1810,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X16)) | sP1(X17,X15) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1244,f561])).
fof(f1817,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X12)) | sP1(X13,X11) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1809,f264])).
fof(f1809,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,X12)) | sP1(X13,X11) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1244,f561])).
fof(f1815,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0') | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f1815_D])).
fof(f1815_D,plain,(
  ( ! [X2,X1] : (sP1(X1,'F_0') | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1814,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | sP1(X1,'F_0') | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1806,f264])).
fof(f1806,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | sP1(X1,'F_0') | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1244,f254])).
fof(f1309,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1304,f246])).
fof(f1304,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f1188,f449])).
fof(f1308,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f245,f244,f243,f242,f246,f449,f254,f265,f264,f269,f268,f293,f296,f295,f294,f302,f300,f303,f308,f305,f315,f314,f450,f333,f330,f329,f328,f346,f348,f347,f361,f360,f359,f358,f357,f356,f355,f365,f364,f363,f362,f369,f368,f367,f366,f373,f372,f371,f370,f377,f376,f375,f374,f431,f381,f380,f379,f378,f388,f387,f386,f385,f394,f391,f390,f389,f401,f400,f399,f430,f397,f396,f395,f405,f404,f403,f402,f409,f408,f407,f406,f418,f417,f416,f429,f414,f413,f412,f422,f421,f420,f419,f426,f425,f424,f423,f455,f456,f457,f458,f459,f460,f461,f462,f463,f464,f465,f466,f467,f468,f469,f470,f471,f472,f473,f474,f475,f482,f483,f485,f491,f496,f497,f500,f503,f506,f513,f514,f516,f523,f527,f528,f531,f537,f561,f607,f616,f615,f626,f624,f638,f644,f654,f657,f662,f666,f674,f680,f703,f710,f716,f739,f746,f752,f784,f798,f819,f821,f849,f937,f1186,f1190,f1193,f1196,f1199,f1202,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1240,f1241,f1242,f1243,f1244,f1245,f1246,f1247,f1248,f1249,f1250,f1251,f1252,f1253,f1254,f1255,f1256,f1257,f1258,f1259,f1260,f1261,f1267,f1269,f1271,f1273,f1275,f1277,f1279,f1281,f1283,f1285,f1287,f1289,f1291,f1293,f1295,f1297,f1299,f1301,f1188,f1307])).
fof(f1307,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f1302,f1188])).
fof(f1302,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl12),
  inference(backward_demodulation,[],[f1188,f268])).
fof(f1239,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f1239_D])).
fof(f1239_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1188,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f1188_D])).
fof(f1188_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1301,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1300])).
fof(f1300,plain,(
  ( ! [X97,X95,X98,X96] : (~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f1029,f607])).
fof(f1029,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f430,f561])).
fof(f1236,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f1236_D])).
fof(f1236_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1299,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1298])).
fof(f1298,plain,(
  ( ! [X94,X92,X93,X91] : (~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1028,f607])).
fof(f1028,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f429,f561])).
fof(f1233,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f1233_D])).
fof(f1233_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1297,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1296])).
fof(f1296,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1027,f607])).
fof(f1027,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f426,f561])).
fof(f1230,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1295,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1294])).
fof(f1294,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1026,f607])).
fof(f1026,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f425,f561])).
fof(f1227,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f1227_D])).
fof(f1227_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1293,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1292])).
fof(f1292,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1025,f607])).
fof(f1025,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f424,f561])).
fof(f1224,plain,(
  ( ! [X2,X1] : (sP11(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  ( ! [X2,X1] : (sP11(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1291,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1290])).
fof(f1290,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1024,f607])).
fof(f1024,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f423,f561])).
fof(f1221,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1221_D])).
fof(f1221_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1289,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1288])).
fof(f1288,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1023,f607])).
fof(f1023,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f419,f561])).
fof(f1218,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f1218_D])).
fof(f1218_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1287,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1286])).
fof(f1286,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1022,f607])).
fof(f1022,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f409,f561])).
fof(f1215,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1285,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1284])).
fof(f1284,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1021,f607])).
fof(f1021,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f408,f561])).
fof(f1212,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1283,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1282])).
fof(f1282,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1020,f607])).
fof(f1020,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f407,f561])).
fof(f1209,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f1209_D])).
fof(f1209_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1281,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1280])).
fof(f1280,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1019,f607])).
fof(f1019,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f402,f561])).
fof(f1206,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1206_D])).
fof(f1206_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1279,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1278])).
fof(f1278,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1018,f607])).
fof(f1018,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f394,f561])).
fof(f1203,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f1203_D])).
fof(f1203_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1277,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1276])).
fof(f1276,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1017,f607])).
fof(f1017,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f391,f561])).
fof(f1200,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1200_D])).
fof(f1200_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1275,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1274])).
fof(f1274,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1016,f607])).
fof(f1016,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f385,f561])).
fof(f1197,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f1197_D])).
fof(f1197_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1273,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1272])).
fof(f1272,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1015,f607])).
fof(f1015,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f366,f561])).
fof(f1194,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1194_D])).
fof(f1194_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1271,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1270])).
fof(f1270,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1014,f607])).
fof(f1014,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f362,f561])).
fof(f1191,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f1191_D])).
fof(f1191_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1269,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1268])).
fof(f1268,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1013,f607])).
fof(f1013,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f333,f561])).
fof(f1187,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl10),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1267,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1266])).
fof(f1266,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1012,f607])).
fof(f1012,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f330,f561])).
fof(f1238,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1237])).
fof(f1237,plain,(
  ( ! [X97,X95,X98,X96] : (~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f965,f607])).
fof(f965,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP8('T_0',X96,X97) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f430,f561])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1234])).
fof(f1234,plain,(
  ( ! [X94,X92,X93,X91] : (~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f964,f607])).
fof(f964,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP10('T_0',X92,X93) | sP11('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f429,f561])).
fof(f1232,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1231])).
fof(f1231,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f963,f607])).
fof(f963,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | p(s(bool,X89)) | ~sP11('T_0',X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f426,f561])).
fof(f1229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1228])).
fof(f1228,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f962,f607])).
fof(f962,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f425,f561])).
fof(f1226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1225])).
fof(f1225,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f961,f607])).
fof(f961,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f424,f561])).
fof(f1223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1222])).
fof(f1222,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f960,f607])).
fof(f960,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11(X77,X76,'T_0') | p(s(bool,X77)) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f423,f561])).
fof(f1220,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1219])).
fof(f1219,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f959,f607])).
fof(f959,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f419,f561])).
fof(f1217,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1216])).
fof(f1216,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f958,f607])).
fof(f958,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f409,f561])).
fof(f1214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1213])).
fof(f1213,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f957,f607])).
fof(f957,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f408,f561])).
fof(f1211,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1210])).
fof(f1210,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f956,f607])).
fof(f956,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f407,f561])).
fof(f1208,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1207])).
fof(f1207,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f955,f607])).
fof(f955,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f402,f561])).
fof(f1205,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1204])).
fof(f1204,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f954,f607])).
fof(f954,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f394,f561])).
fof(f1202,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f1201])).
fof(f1201,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f953,f607])).
fof(f953,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | ~p(s(bool,X49)) | sP7('T_0',X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f391,f561])).
fof(f1199,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1198])).
fof(f1198,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f952,f607])).
fof(f952,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f385,f561])).
fof(f1196,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1195])).
fof(f1195,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f951,f607])).
fof(f951,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f366,f561])).
fof(f1193,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1192])).
fof(f1192,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f950,f607])).
fof(f950,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f362,f561])).
fof(f1190,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1189])).
fof(f1189,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f949,f607])).
fof(f949,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f333,f561])).
fof(f1186,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1185])).
fof(f1185,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f948,f607])).
fof(f948,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f330,f561])).
fof(f453,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f452])).
fof(f452,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f451,f441])).
fof(f451,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK14))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f440,f443])).
fof(f446,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | $spl6),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f448,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))) | $spl8),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f242,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15)))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK16)))) & (s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) | (s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) & s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),sK16))) & (s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) | s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK15) | s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),sK16))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4)))) & (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) | s('type.list.list'(X0),X1) != s('type.list.list'(X0),X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4)))) & ((s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))) & (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) | (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X3) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),X4)))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4)))) & (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) <~> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3,X4] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4)))) => (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) <=> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X4)))) => (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))) <=> (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X3) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),X4))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X10)))) => (s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X10))) <=> (s('type.list.list'(X6),X7) = s('type.list.list'(X6),X9) & s('type.list.list'(X6),X8) = s('type.list.list'(X6),X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X10)))) => (s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X7),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X10))) <=> (s('type.list.list'(X6),X7) = s('type.list.list'(X6),X9) & s('type.list.list'(X6),X8) = s('type.list.list'(X6),X10))))),
  file('thm.list.list_|49||48|8_prove_ltb','thm.list.list_|49||48|8_prove_ltb')).
fof(f439,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) | $spl0),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f441,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | $spl2),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f243,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK15) | s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16)))),
  inference(cnf_transformation,[],[f160])).
fof(f443,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),sK16) | $spl4),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),sK16) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f244,plain,(
  s('type.list.list'(sK12),sK14) = s('type.list.list'(sK12),sK16) | s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16)))),
  inference(cnf_transformation,[],[f160])).
fof(f440,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) | $spl1),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f442,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK15) | $spl3),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f444,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),sK16) | $spl5),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),sK16) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f245,plain,(
  s('type.list.list'(sK12),sK14) != s('type.list.list'(sK12),sK16) | s('type.list.list'(sK12),sK13) != s('type.list.list'(sK12),sK15) | s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK13),s('type.list.list'(sK12),sK14))) != s('type.list.list'(sK12),'const.list.APPEND_2'(s('type.list.list'(sK12),sK15),s('type.list.list'(sK12),sK16)))),
  inference(cnf_transformation,[],[f160])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 586
% Passive clauses: 6448
% Generated clauses: 35684
% Final active clauses: 478
% Final passive clauses: 2828
% Input formulas: 45
% Initial clauses: 187
% Discarded non-redundant clauses: 5388
% 
% Duplicate literals: 1898
% Trivial inequalities: 60
% Fw subsumption resolutions: 641
% Fw demodulations: 3387
% Bw demodulations: 90
% Condensations: 2576
% Global subsumptions: 6
% 
% Simple tautologies: 119
% Equational tautologies: 409
% Forward subsumptions: 19834
% Fw demodulations to eq. taut.: 129
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 3874
% Factoring: 5
% Forward superposition: 11864
% Backward superposition: 10378
% Self superposition: 235
% Equality factoring: 110
% Equality resolution: 54
% 
% Split clauses: 129
% Split components: 277
% SAT solver clauses: 6724
% SAT solver unit clauses: 864
% SAT solver binary clauses: 4105
% 
% Sat splits: 129
% Sat splitting refutations: 10
% 
% Memory used [KB]: 71000
% Time elapsed: 5.488 s
% ------------------------------
% ------------------------------
% Success in time 6.82 s
