Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.while
% (90409)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 817
% Passive clauses: 8471
% Generated clauses: 19842
% Final active clauses: 812
% Final passive clauses: 7411
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 358
% Trivial inequalities: 21
% Fw subsumption resolutions: 1502
% Fw demodulations: 1065
% Bw demodulations: 2
% 
% Simple tautologies: 280
% Equational tautologies: 67
% Forward subsumptions: 7707
% 
% Binary resolution: 9636
% Factoring: 1106
% Forward superposition: 2407
% Backward superposition: 3238
% Self superposition: 72
% Equality factoring: 26
% Equality resolution: 10
% 
% Split clauses: 337
% Split components: 684
% SAT solver clauses: 499
% SAT solver unit clauses: 7
% SAT solver binary clauses: 90
% 
% Sat splits: 337
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 499
% TWLsolver calls for satisfiability: 338
% 
% Memory used [KB]: 14967
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.while
% (90423)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 384
% Passive clauses: 2836
% Generated clauses: 10051
% Final active clauses: 384
% Final passive clauses: 2442
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 117
% Trivial inequalities: 8
% Fw subsumption resolutions: 476
% Fw demodulations: 500
% Bw demodulations: 10
% Condensations: 142
% Global subsumptions: 18
% 
% Simple tautologies: 252
% Equational tautologies: 8
% Forward subsumptions: 2682
% 
% Binary resolution: 2933
% Factoring: 325
% Forward superposition: 1275
% Backward superposition: 1102
% Self superposition: 52
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 3286
% SAT solver unit clauses: 261
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 1
% 
% TWLsolver clauses: 3214
% TWLsolver calls for satisfiability: 2038
% 
% Memory used [KB]: 18166
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.while
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.while
% SZS output start Proof for thm.while
fof(f33975,plain,(
  $false),
  inference(subsumption_resolution,[],[f33854,f15719])).
fof(f15719,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f596,f2493,f457])).
fof(f457,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2493,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2492,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.ZERO_0'))) & ! [X6] : s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X2),s('type.num.num',X6))),s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num',X6))))) & ! [X5] : s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X3),s('type.num.num',X5))),s(X0,happ(s(fun('type.num.num',X0),sK18(X3,X2,X1,X0)),s('type.num.num',X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.ZERO_0'))) & ! [X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X2),s('type.num.num',X6))),s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X6))))) & ! [X5] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X3),s('type.num.num',X5))),s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X5))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.ZERO_0'))) & ! [X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X2),s('type.num.num',X6))),s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X6))))) & ! [X5] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.num.num',fun(X0,X0)),X3),s('type.num.num',X5))),s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X5))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X25] : ! [X29,X30,X31] : ? [X32] : (s(X25,X29) = s(X25,happ(s(fun('type.num.num',X25),X32),s('type.num.num','const.arithmetic.ZERO_0'))) & ! [X9] : s(X25,happ(s(fun('type.num.num',X25),X32),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X9))))) = s(X25,happ(s(fun(X25,X25),happ(s(fun('type.num.num',fun(X25,X25)),X30),s('type.num.num',X9))),s(X25,happ(s(fun('type.num.num',X25),X32),s('type.num.num',X9))))) & ! [X9] : s(X25,happ(s(fun('type.num.num',X25),X32),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9))))) = s(X25,happ(s(fun(X25,X25),happ(s(fun('type.num.num',fun(X25,X25)),X31),s('type.num.num',X9))),s(X25,happ(s(fun('type.num.num',X25),X32),s('type.num.num',X9))))))),
  file('thm.while.OLEAST_EQNS|split|0','thm.numeral.bit_initiality')).
fof(f2492,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X3))),s(X0,happ(s(fun('type.num.num',X0),sK18(X1,X2,X3,X0)),s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(forward_demodulation,[],[f2456,f559])).
fof(f559,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12,X13,X14] : ! [X15,X16,X17] : s(X12,happ(s(fun(X13,X12),'_dst_x1x0_2'(s(fun(X13,fun(X14,X12)),X15),s(X14,X16))),s(X13,X17))) = s(X12,happ(s(fun(X14,X12),happ(s(fun(X13,fun(X14,X12)),X15),s(X13,X17))),s(X14,X16)))),
  file('thm.while.OLEAST_EQNS|split|0','_dst_x1x0_ax')).
fof(f2456,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),sK18(X1,X2,X3,X0)),s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X3))))) )),
  inference(unit_resulting_resolution,[],[f409,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X18,X19,X20] : (p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X18,fun(X18,bool)),'=_0'),s(X18,X19))),s(X18,X20)))) <=> s(X18,X19) = s(X18,X20))),
  file('thm.while.OLEAST_EQNS|split|0',arity_0_HL_EQ)).
fof(f596,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f337,f333])).
fof(f333,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f21])).
fof(f21,axiom,(
  p(s(bool,'T_0'))),
  file('thm.while.OLEAST_EQNS|split|0','thm.bool.TRUTH')).
fof(f337,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f48])).
fof(f48,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.while.OLEAST_EQNS|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f33854,plain,(
  ~sP1(happ(s(fun('type.option.option'('type.num.num'),bool),'_dst_x1x0_2'(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13)))),'F_0')),
  inference(backward_demodulation,[],[f33730,f17012])).
fof(f17012,plain,(
  ~sP1(happ(s(fun('type.option.option'('type.num.num'),bool),'_dst_x1x0_2'(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))),s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK13)))))),'F_0')),
  inference(unit_resulting_resolution,[],[f596,f2384,f455])).
fof(f455,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f284])).
fof(f2384,plain,(
  ~p(s(bool,happ(s(fun('type.option.option'('type.num.num'),bool),'_dst_x1x0_2'(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))),s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK13))))))))),
  inference(forward_demodulation,[],[f2359,f559])).
fof(f2359,plain,(
  ~p(s(bool,happ(s(fun('type.option.option'('type.num.num'),bool),happ(s(fun('type.option.option'('type.num.num'),fun('type.option.option'('type.num.num'),bool)),'=_0'),s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK13))))))),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))))))),
  inference(unit_resulting_resolution,[],[f332,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f287])).
fof(f332,plain,(
  s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK13)))))),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',sK13))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f171])).
fof(f171,plain,(
  ? [X0] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) != s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X6))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X6))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X6)))))),
  file('thm.while.OLEAST_EQNS|split|0','thm.while.OLEAST_EQNS|split|0')).
fof(f33730,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f33666,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X25] : ! [X34,X35] : s(X25,X34) = s(X25,'const.bool.COND_3'(s(bool,'T_0'),s(X25,X34),s(X25,X35)))),
  file('thm.while.OLEAST_EQNS|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f33666,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,'T_0'),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))) = s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f32821,f2791])).
fof(f2791,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK19('_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num',X0))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(superposition,[],[f414,f342])).
fof(f342,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X6)))))),
  file('thm.while.OLEAST_EQNS|split|0','thm.while.LEAST_EQ|split|0')).
fof(f414,plain,(
  ( ! [X0] : (s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK19(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : ((~p(s(bool,sK19(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK20(X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK19(X0)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK19(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f268])).
fof(f268,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ? [X8] : ((p(s(bool,X8)) <=> ? [X9] : p(s(bool,happ(s(fun('type.num.num',bool),X7),s('type.num.num',X9))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X7))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X8),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X7))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  file('thm.while.OLEAST_EQNS|split|0','thm.while.OLEAST_def')).
fof(f32821,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK19('_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f333,f15012,f429])).
fof(f429,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X34,X35] : ((p(s(bool,X34)) => p(s(bool,X35))) => ((p(s(bool,X35)) => p(s(bool,X34))) => s(bool,X34) = s(bool,X35)))),
  file('thm.while.OLEAST_EQNS|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f15012,plain,(
  ( ! [X0] : (p(s(bool,sK19('_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2487,f413])).
fof(f413,plain,(
  ( ! [X2,X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK19(X0)))) )),
  inference(cnf_transformation,[],[f269])).
fof(f2487,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f2486,f342])).
fof(f2486,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f2459,f559])).
fof(f2459,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num',X0))))))),s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f342,f467])).
% SZS output end Proof for thm.while
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 832
% Passive clauses: 6919
% Generated clauses: 34550
% Final active clauses: 764
% Final passive clauses: 4594
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 254
% Trivial inequalities: 27
% Fw subsumption resolutions: 1588
% Fw demodulations: 5135
% Bw demodulations: 1536
% Condensations: 681
% Global subsumptions: 1034
% 
% Simple tautologies: 92
% Equational tautologies: 55
% Forward subsumptions: 16605
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 9962
% Factoring: 1
% Forward superposition: 7722
% Backward superposition: 5152
% Self superposition: 82
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 76
% Split components: 157
% SAT solver clauses: 9980
% SAT solver unit clauses: 4778
% SAT solver binary clauses: 2867
% 
% Sat splits: 76
% Sat splitting refutations: 1873
% 
% Memory used [KB]: 97994
% Time elapsed: 2.922 s
% ------------------------------
% ------------------------------
% Success in time 4.251 s
