Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f6796,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f6783])).
fof(f6783,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f56,f6766])).
fof(f6766,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X4),X5),s('type.list.list'('type.quote.index'),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f2943,f63])).
fof(f63,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X11,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X6),X11),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar2_def|split|2','thm.canonical.canonical_sum_scalar2_def|split|2')).
fof(f2943,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f2934,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X15),s('type.list.list'('type.quote.index'),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X6),X15))),s('type.list.list'('type.quote.index'),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_scalar2_def|split|2','ARITY_1_3_const.ringNorm.r_canonical_sum_scalar2')).
fof(f2934,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f69,f60])).
fof(f60,plain,(
  ( ! [X0,X1] : (s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_scalar2_def|split|2','thm.ringNorm.r_canonical_sum_scalar2_def')).
fof(f69,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.list.list'('type.quote.index'),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.list.list'('type.quote.index'),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_scalar2_def|split|2','ARITY_1_3_const.canonical.canonical_sum_scalar2')).
fof(f56,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f51])).
fof(f51,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_scalar2_def|split|2','thm.ringNorm.canonical_sum_scalar2_def|split|2')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 46
% Passive clauses: 717
% Generated clauses: 6507
% Final active clauses: 44
% Final passive clauses: 623
% Input formulas: 18
% Initial clauses: 19
% 
% Duplicate literals: 574
% Trivial inequalities: 87
% Fw subsumption resolutions: 856
% Fw demodulations: 3
% 
% Simple tautologies: 176
% Equational tautologies: 109
% Forward subsumptions: 2909
% 
% Binary resolution: 25
% Forward superposition: 2282
% Backward superposition: 1928
% Self superposition: 370
% Equality factoring: 326
% Equality resolution: 8
% 
% Split clauses: 340
% Split components: 682
% SAT solver clauses: 1110
% SAT solver unit clauses: 232
% SAT solver binary clauses: 380
% 
% Sat splits: 340
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1108
% TWLsolver calls for satisfiability: 32
% 
% Memory used [KB]: 3709
% Time elapsed: 0.398 s
% ------------------------------
% ------------------------------
% Success in time 0.404 s
