Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% (51313)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1279
% Passive clauses: 8228
% Generated clauses: 45054
% Final active clauses: 1158
% Final passive clauses: 6230
% Input formulas: 261
% Initial clauses: 506
% 
% Duplicate literals: 274
% Trivial inequalities: 70
% Fw subsumption resolutions: 2212
% Fw demodulations: 6053
% Bw demodulations: 119
% 
% Simple tautologies: 268
% Equational tautologies: 31
% Forward subsumptions: 27770
% Fw demodulations to eq. taut.: 22
% 
% Binary resolution: 7813
% Factoring: 1418
% Forward superposition: 14475
% Backward superposition: 11837
% Self superposition: 47
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 511
% Split components: 1024
% SAT solver clauses: 773
% SAT solver unit clauses: 23
% SAT solver binary clauses: 110
% 
% Sat splits: 511
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 771
% TWLsolver calls for satisfiability: 477
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (51338)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 626
% Passive clauses: 2834
% Generated clauses: 24582
% Final active clauses: 600
% Final passive clauses: 1973
% Input formulas: 261
% Initial clauses: 506
% 
% Duplicate literals: 96
% Trivial inequalities: 14
% Fw subsumption resolutions: 661
% Fw demodulations: 3599
% Bw demodulations: 258
% Condensations: 89
% Global subsumptions: 37
% 
% Simple tautologies: 238
% Equational tautologies: 12
% Forward subsumptions: 13636
% Fw demodulations to eq. taut.: 22
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2636
% Factoring: 192
% Forward superposition: 7384
% Backward superposition: 6136
% Self superposition: 86
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 21
% Split components: 42
% SAT solver clauses: 3949
% SAT solver unit clauses: 710
% SAT solver binary clauses: 145
% SAT solver learnt clauses: 74
% SAT solver learnt literals: 14
% 
% Sat splits: 21
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3801
% TWLsolver calls for satisfiability: 3222
% 
% Memory used [KB]: 27121
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f64410,plain,(
  $false),
  inference(subsumption_resolution,[],[f64408,f24319])).
fof(f24319,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat',X0))),s('type.rat.rat',X0))))) )),
  inference(forward_demodulation,[],[f24037,f854])).
fof(f854,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_CLASS_1'(s(fun('type.frac.frac',bool),'const.rat.rep_rat_CLASS_1'(s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_CLASS_1'(s(fun('type.frac.frac',bool),'const.rat.rep_rat_CLASS_1'(s('type.rat.rat',X0)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X41] : s('type.rat.rat',X41) = s('type.rat.rat','const.rat.abs_rat_CLASS_1'(s(fun('type.frac.frac',bool),'const.rat.rep_rat_CLASS_1'(s('type.rat.rat',X41)))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.rat.rat_bijections|split|0')).
fof(f24037,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.abs_rat_CLASS_1'(s(fun('type.frac.frac',bool),'const.rat.rep_rat_CLASS_1'(s('type.rat.rat',X0))))))),s('type.rat.rat',X0))))) )),
  inference(unit_resulting_resolution,[],[f854,f1188])).
fof(f1188,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f150])).
fof(f150,axiom,(
  ! [X27,X15,X16] : (p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X27,fun(X27,bool)),'=_0'),s(X27,X15))),s(X27,X16)))) <=> s(X27,X15) = s(X27,X16))),
  file('thm.rat.rat_of_num_def_compute|split|2',arity_0_HL_EQ)).
fof(f64408,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK15))))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK15))))))))),
  inference(backward_demodulation,[],[f64407,f23509])).
fof(f23509,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK15))))))),s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))),s('type.rat.rat','const.rat.rat_|49|_0'))))))),
  inference(unit_resulting_resolution,[],[f1355,f1187])).
fof(f1187,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f753])).
fof(f1355,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK15))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(backward_demodulation,[],[f856,f1346])).
fof(f1346,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(backward_demodulation,[],[f828,f801])).
fof(f801,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(cnf_transformation,[],[f628])).
fof(f628,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK15))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f564])).
fof(f564,plain,(
  ? [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.rat.rat_of_num_def_compute|split|2')).
fof(f828,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : s('type.num.num',X12) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.NUMERAL_DEF')).
fof(f856,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f205])).
fof(f205,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.numeral.numeral_suc|split|1')).
fof(f64407,plain,(
  ( ! [X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(forward_demodulation,[],[f64125,f856])).
fof(f64125,plain,(
  ( ! [X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(superposition,[],[f896,f22518])).
fof(f22518,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f22511,f13558])).
fof(f13558,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f1042,f9942])).
fof(f9942,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f9832,f1383])).
fof(f1383,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f895,f1335])).
fof(f1335,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f814,f804])).
fof(f804,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(cnf_transformation,[],[f21])).
fof(f21,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.ALT_ZERO')).
fof(f814,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f204])).
fof(f204,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.numeral.numeral_suc|split|0')).
fof(f895,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.BIT|49|')).
fof(f9832,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1379,f1035])).
fof(f1035,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X13,X6] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X6)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.ADD_SYM')).
fof(f1379,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1378,f804])).
fof(f1378,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f892,f828])).
fof(f892,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.SUC_ONE_ADD')).
fof(f1042,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X6,X13] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f22511,plain,(
  ( ! [X3] : (s('type.num.num','const.numeral.iDUB_1'(s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f22510,f18047])).
fof(f18047,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))) )),
  inference(forward_demodulation,[],[f17834,f1347])).
fof(f1347,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f828,f1343])).
fof(f1343,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1342,f1339])).
fof(f1339,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1338,f1335])).
fof(f1338,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f817,f804])).
fof(f817,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f57])).
fof(f57,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.numeral.numeral_distrib|split|13')).
fof(f1342,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f820,f804])).
fof(f820,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.TWO')).
fof(f17834,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f1048,f1370])).
fof(f1370,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1369,f804])).
fof(f1369,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f882,f828])).
fof(f882,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X13] : s('type.num.num',X13) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X13)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.MULT_LEFT_|49|')).
fof(f1048,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X13,X6] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X6))),s('type.num.num',X6)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.arithmetic.MULT|split|1')).
fof(f22510,plain,(
  ( ! [X0] : (s('type.num.num','const.numeral.iDUB_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f22300,f1370])).
fof(f22300,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f18051,f1335])).
fof(f18051,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f17879,f1433])).
fof(f1433,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f1432,f1035])).
fof(f1432,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1068,f829])).
fof(f829,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f222])).
fof(f222,axiom,(
  ! [X12] : s('type.num.num',X12) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X12)))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.numeral.iZ')).
fof(f1068,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))))))),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X6,X13] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num',X13))) = s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X13))),s('type.num.num',X13)))))))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.numeral.numeral_mult|split|3')).
fof(f17879,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f1048,f1035])).
fof(f896,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  file('thm.rat.rat_of_num_def_compute|split|2','thm.rat.rat_of_num_def|split|2')).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1055
% Passive clauses: 10013
% Generated clauses: 63565
% Final active clauses: 1015
% Final passive clauses: 8300
% Input formulas: 261
% Initial clauses: 506
% 
% Duplicate literals: 9
% Trivial inequalities: 22
% Fw subsumption resolutions: 685
% Fw demodulations: 11487
% Bw demodulations: 418
% Condensations: 18
% 
% Simple tautologies: 89
% Equational tautologies: 66
% Forward subsumptions: 39763
% Fw demodulations to eq. taut.: 377
% Bw demodulations to eq. taut.: 20
% 
% Unit resulting resolution: 17379
% Factoring: 1
% Forward superposition: 18634
% Backward superposition: 14151
% Self superposition: 219
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 26
% Split components: 57
% SAT solver clauses: 10807
% SAT solver unit clauses: 4843
% SAT solver binary clauses: 4693
% 
% Sat splits: 26
% Sat splitting refutations: 740
% 
% Memory used [KB]: 88527
% Time elapsed: 3.505 s
% ------------------------------
% ------------------------------
% Success in time 4.791 s
