Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f16751,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f16750,f285,f495])).
fof(f495,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f285,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.TC_RULES|split|0',aHL_FALSITY)).
fof(f16750,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f16749,f495])).
fof(f16749,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f16748,f3533])).
fof(f3533,plain,(
  s(bool,'F_0') = s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15),s(sK13,sK16)))),
  inference(resolution,[],[f278,f300])).
fof(f300,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.relation.TC_RULES|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f278,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15),s(sK13,sK16))))),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16)))) & ~p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15),s(sK13,sK16))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f126])).
fof(f126,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))))),
  file('thm.relation.TC_RULES|split|0','thm.relation.TC_RULES|split|0')).
fof(f16748,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15),s(sK13,sK16))))),
  inference(resolution,[],[f16737,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X2,X3,X1) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | sP0(X0,X2,X3,X1)) & (~sP0(X0,X2,X3,X1) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> sP0(X0,X2,X3,X1))),
  inference(definition_folding,[],[f142,f153])).
fof(f153,plain,(
  ! [X0,X2,X3,X1] : (sP0(X0,X2,X3,X1) <=> ! [X4] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X7,X19,X20] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X19),s(X6,X20)))) <=> ! [X14] : ((! [X8,X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X8))),s(X6,X9))))) & ! [X8,X9,X21] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X8))),s(X6,X9)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X9))),s(X6,X21))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X8))),s(X6,X21)))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X19))),s(X6,X20))))))),
  file('thm.relation.TC_RULES|split|0','thm.relation.TC_DEF')).
fof(f16737,plain,(
  sP0(sK13,sK15,sK16,sK14)),
  inference(duplicate_literal_removal,[],[f16662])).
fof(f16662,plain,(
  sP0(sK13,sK15,sK16,sK14) | sP0(sK13,sK15,sK16,sK14)),
  inference(resolution,[],[f6743,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X1))),s(X0,X2)))) | sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2,X3] : ((~sP0(X0,X1,X2,X3) | ! [X10] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK28(X10,X3,X2,X1,X0)))),s(X0,sK29(X10,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,sK28(X10,X3,X2,X1,X0)))),s(X0,sK29(X10,X3,X2,X1,X0)))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,sK30(X10,X3,X2,X1,X0)))),s(X0,sK31(X10,X3,X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,sK31(X10,X3,X2,X1,X0)))),s(X0,sK32(X10,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,sK30(X10,X3,X2,X1,X0)))),s(X0,sK32(X10,X3,X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X1))),s(X0,X2)))))) & ((! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X1))),s(X0,X2))))) | sP0(X0,X1,X2,X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29,sK30,sK31,sK32,sK33])],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3] : ((~sP0(X0,X1,X2,X3) | ! [X10] : (? [X14,X15] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X14))),s(X0,X15)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X14))),s(X0,X15))))) | ? [X11,X12,X13] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X11))),s(X0,X12)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X12))),s(X0,X13)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X11))),s(X0,X13))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X10),s(X0,X1))),s(X0,X2)))))) & (? [X4] : (! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X1))),s(X0,X2))))) | sP0(X0,X1,X2,X3)))),
  inference(rectify,[],[f214])).
fof(f214,plain,(
  ! [X0,X2,X3,X1] : ((~sP0(X0,X2,X3,X1) | ! [X4] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3)))))) & (? [X4] : (! [X8,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))) | sP0(X0,X2,X3,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f6743,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK33(sK14,X0,X1,sK13)),s(sK13,sK15))),s(sK13,sK16)))) | sP0(sK13,X1,X0,sK14)) )),
  inference(resolution,[],[f352,f277])).
fof(f277,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK15))),s(sK13,sK16))))),
  inference(cnf_transformation,[],[f173])).
fof(f352,plain,(
  ( ! [X2,X0,X8,X3,X1,X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK33(X3,X2,X1,X0)),s(X0,X8))),s(X0,X9)))) | sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f216])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 802
% Passive clauses: 6149
% Generated clauses: 16370
% Final active clauses: 792
% Final passive clauses: 5034
% Input formulas: 48
% Initial clauses: 210
% 
% Duplicate literals: 334
% Trivial inequalities: 45
% Fw subsumption resolutions: 1716
% Fw demodulations: 938
% Bw demodulations: 9
% 
% Simple tautologies: 241
% Equational tautologies: 34
% Forward subsumptions: 6588
% 
% Binary resolution: 6927
% Factoring: 828
% Forward superposition: 1581
% Backward superposition: 3538
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 425
% Split components: 851
% SAT solver clauses: 645
% SAT solver unit clauses: 10
% SAT solver binary clauses: 115
% SAT solver learnt clauses: 1
% 
% Sat splits: 425
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 644
% TWLsolver calls for satisfiability: 361
% 
% Memory used [KB]: 8443
% Time elapsed: 0.337 s
% ------------------------------
% ------------------------------
% Success in time 0.34 s
