Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rat
% (94159)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 1724
% Generated clauses: 21969
% Final active clauses: 109
% Final passive clauses: 1492
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 1122
% Trivial inequalities: 297
% Fw subsumption resolutions: 3141
% Fw demodulations: 2923
% Bw demodulations: 12
% 
% Simple tautologies: 349
% Equational tautologies: 65
% Forward subsumptions: 10365
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 98
% Forward superposition: 8550
% Backward superposition: 5500
% Self superposition: 279
% Equality factoring: 171
% Equality resolution: 24
% 
% Split clauses: 1450
% Split components: 2932
% SAT solver clauses: 2102
% SAT solver unit clauses: 23
% SAT solver binary clauses: 228
% 
% Sat splits: 1450
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2088
% TWLsolver calls for satisfiability: 128
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (94180)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 81
% Passive clauses: 281
% Generated clauses: 10933
% Final active clauses: 76
% Final passive clauses: 172
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 1496
% Trivial inequalities: 73
% Fw subsumption resolutions: 752
% Fw demodulations: 747
% Bw demodulations: 32
% Condensations: 1927
% 
% Simple tautologies: 164
% Equational tautologies: 48
% Forward subsumptions: 3892
% 
% Binary resolution: 30
% Forward superposition: 2964
% Backward superposition: 2481
% Self superposition: 169
% Equality factoring: 140
% Equality resolution: 8
% 
% Split clauses: 125
% Split components: 278
% SAT solver clauses: 852
% SAT solver unit clauses: 22
% SAT solver binary clauses: 68
% 
% Sat splits: 125
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 852
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 7547
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f30680,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f56,f57,f86,f90,f88,f500,f501,f507,f502,f503,f515,f504,f517,f510,f509,f527,f528,f529,f530,f531,f541,f567,f568,f569,f570,f571,f572,f540,f573,f574,f575,f538,f576,f535,f820,f822,f825,f829,f832,f867,f1222,f1206,f1192,f1187,f1186,f1177,f1231,f1176,f1175,f1256,f1257,f1153,f1242,f1148,f1258,f1147,f1139,f1243,f1310,f1421,f1434,f1656,f1814,f1816,f1817,f1806,f1818,f1820,f1821,f1822,f1823,f1825,f1788,f1703,f1827,f516,f1813,f1763,f1776,f1831,f1835,f1838,f1839,f1840,f511,f1842,f1844,f1815,f1431,f59,f1864,f1867,f1868,f1909,f2009,f2010,f2851,f2831,f2869,f2813,f2807,f2796,f2859,f2795,f2883,f2884,f2904,f2915,f2916,f2766,f2870,f2761,f2917,f2918,f2921,f2923,f1836,f3068,f3080,f3113,f3114,f3314,f3315,f3406,f3619,f3621,f3609,f3623,f3626,f514,f3627,f3628,f3591,f3466,f3616,f3819,f3824,f3825,f3826,f3828,f3830,f3831,f3832,f4746,f4748,f6009,f6010,f6041,f6040,f7582,f7604,f7753,f4750,f7755,f7787,f8933,f8940,f8941,f8945,f512,f9147,f9148,f9174,f9175,f9179,f12963,f15496,f15499,f15533,f15535,f15497,f20903,f20958,f21152,f21154,f21544,f21571,f21918,f21153,f22304,f23135,f2217,f2218,f539,f23181,f23205,f23206,f23212,f3618,f23214,f5020,f537,f23220,f4662,f4659,f2919,f534,f4609,f4608,f23213,f23215,f2258,f2255,f23232,f581,f23235,f23236,f23238,f23240,f23241,f23217,f3620,f23246,f23253,f23256,f23258,f23285,f25155,f25176,f25157,f25156,f25333,f25513,f25698,f25729,f25730,f25837,f26567,f536,f27013,f27360,f27444,f27445,f27449,f27453,f27560,f28123,f28134,f28136,f28138,f28140,f28142,f28146,f28150,f28256,f28590,f28675,f28676,f28680,f28684,f28790,f30560,f29984,f30586,f30532,f30507,f30490,f30595,f30596,f30587,f30488,f30487,f30456,f30584,f30452,f30614,f30615,f30616,f30617,f30401,f30618,f30382,f30381,f30620,f30353,f30597,f30660,f30677,f30678,f1230])).
fof(f1230,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) | $spl34),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f30678,plain,(
  ( ! [X61,X60] : (s('type.rat.rat',X60) = s('type.rat.rat',X61) | s(bool,'F_0') = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) ) | $spl134),
  inference(cnf_transformation,[],[f30678_D])).
fof(f30678_D,plain,(
  ( ! [X61,X60] : (s('type.rat.rat',X60) = s('type.rat.rat',X61) | s(bool,'F_0') = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f30677,plain,(
  ( ! [X61,X62,X60] : (s('type.rat.rat',X60) = s('type.rat.rat',X61) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X60),s('type.rat.rat',X61)))) | s(bool,'F_0') = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f30654,f53])).
fof(f53,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f36,f37])).
fof(f37,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_LEQ_CALCULATE',aHL_BOOL_CASES)).
fof(f36,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.RAT_LEQ_CALCULATE',aHL_FALSITY)).
fof(f30654,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,X62)) | s('type.rat.rat',X60) = s('type.rat.rat',X61) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X60),s('type.rat.rat',X61)))) | s(bool,'F_0') = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f44,f23309])).
fof(f23309,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23255,f51])).
fof(f51,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f37,f37])).
fof(f23255,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl10 | $spl30)),
  inference(forward_demodulation,[],[f508,f540])).
fof(f508,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl10),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f44,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) = s('type.rat.rat',X1) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & ((~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) & s('type.rat.rat',X0) != s('type.rat.rat',X1)) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) = s('type.rat.rat',X1))) & ((~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) & s('type.rat.rat',X0) != s('type.rat.rat',X1)) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) <=> (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X12] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))) <=> (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat',X12)))),
  file('thm.rat.RAT_LEQ_CALCULATE','thm.rat.rat_leq_def')).
fof(f30660,plain,(
  $false | ($spl3 | $spl7 | $spl83 | $spl85)),
  inference(subsumption_resolution,[],[f30659,f34])).
fof(f34,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_LEQ_CALCULATE',aHL_TRUTH)).
fof(f30659,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl7 | $spl83 | $spl85)),
  inference(forward_demodulation,[],[f30658,f22723])).
fof(f22723,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl7),
  inference(equality_resolution,[],[f21003])).
fof(f21003,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl7),
  inference(subsumption_resolution,[],[f9510,f20953])).
fof(f20953,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl7),
  inference(global_subsumption,[],[f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f5026,f6038,f7476,f90,f33,f9508,f9526,f18290,f18291,f18728,f20717,f16124])).
fof(f16124,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | $spl7),
  inference(equality_factoring,[],[f9526])).
fof(f20717,plain,(
  ( ! [X4,X5,X3] : (s('type.rat.rat',X3) = s('type.rat.rat',X4) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,X5) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f20683,f34])).
fof(f20683,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s('type.rat.rat',X3) = s('type.rat.rat',X4) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,X5) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f377,f51])).
fof(f18728,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X3),s('type.rat.rat',X4)))) )),
  inference(subsumption_resolution,[],[f18697,f36])).
fof(f18697,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X3),s('type.rat.rat',X4)))) )),
  inference(superposition,[],[f343,f37])).
fof(f18291,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) )),
  inference(subsumption_resolution,[],[f18258,f36])).
fof(f18258,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) )),
  inference(superposition,[],[f342,f37])).
fof(f18290,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | $spl7),
  inference(condensation,[],[f18289])).
fof(f18289,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl7),
  inference(subsumption_resolution,[],[f18257,f53])).
fof(f18257,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl7),
  inference(superposition,[],[f342,f9526])).
fof(f9526,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl7),
  inference(global_subsumption,[],[f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f5026,f6038,f7476,f90,f33,f9508])).
fof(f9508,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33,f51])).
fof(f33,plain,(
  s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f23])).
fof(f23,plain,(
  ? [X0,X1] : s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1] : s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X7))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X7))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X7))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X7))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X6)))))))),
  file('thm.rat.RAT_LEQ_CALCULATE','thm.rat.RAT_LEQ_CALCULATE')).
fof(f7476,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4)))))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))) )),
  inference(subsumption_resolution,[],[f7405,f34])).
fof(f7405,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4)))))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f42,f593])).
fof(f6038,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,g(s(X16,sK2(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,f(s(X16,sK2(X17,X18,bool,X16))))) )),
  inference(trivial_inequality_removal,[],[f6024])).
fof(f6024,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK2(X17,X18,bool,X16)))) | s(bool,'T_0') = s(bool,f(s(X16,sK2(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f471,f37])).
fof(f5026,plain,(
  ( ! [X127,X125,X126] : (s(X127,f(s(bool,'F_0'))) != s(X127,g(s(bool,'F_0'))) | s(fun(bool,X127),X125) = s(fun(bool,X127),X126) | s(X127,f(s(bool,'T_0'))) != s(X127,g(s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f5025])).
fof(f5025,plain,(
  ( ! [X127,X125,X126] : (s(X127,f(s(bool,'T_0'))) != s(X127,g(s(bool,'T_0'))) | s(fun(bool,X127),X125) = s(fun(bool,X127),X126) | s(fun(bool,X127),X125) = s(fun(bool,X127),X126) | s(X127,f(s(bool,'F_0'))) != s(X127,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f49,f465])).
fof(f471,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,f(s(X19,sK2(X20,X21,bool,X19)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK2(X20,X21,bool,X19))))) )),
  inference(superposition,[],[f49,f37])).
fof(f469,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) != s(bool,f(s(X11,sK2(X12,X13,bool,X11)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK2(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f49,f51])).
fof(f467,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK2(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK2(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f49,f51])).
fof(f377,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) = s('type.rat.rat',X1) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(condensation,[],[f376])).
fof(f376,plain,(
  ( ! [X6,X8,X7] : (s('type.rat.rat',X6) = s('type.rat.rat',X7) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X6),s('type.rat.rat',X7)))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f373,f53])).
fof(f373,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s('type.rat.rat',X6) = s('type.rat.rat',X7) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X6),s('type.rat.rat',X7)))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f44,f51])).
fof(f343,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)))) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f206,f34])).
fof(f206,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)))) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f45,f51])).
fof(f45,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f46,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat',X1) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f41,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9,X10] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int',X9) = s('type.integer.int',X10)))),
  file('thm.rat.RAT_LEQ_CALCULATE','thm.integer.INT_LE_LT')).
fof(f42,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f43,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f28])).
fof(f108,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(subsumption_resolution,[],[f107,f36])).
fof(f107,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f41,f37])).
fof(f342,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X19),s('type.integer.int',X20)))) | s(bool,X21) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f205,f34])).
fof(f205,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X19),s('type.integer.int',X20)))) | s(bool,X21) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f42,f51])).
fof(f459,plain,(
  ( ! [X6,X4,X5] : (s('type.integer.int',X4) = s('type.integer.int',X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,X6) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f452,f34])).
fof(f452,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int',X4) = s('type.integer.int',X5) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,X6) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X4),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f108,f51])).
fof(f720,plain,(
  ( ! [X6,X4,X5] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))) | s(bool,X6) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f713,f34])).
fof(f713,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))) | s(bool,X6) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f591,f51])).
fof(f664,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X3))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X2)))))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f658,f34])).
fof(f658,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X3))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X2)))))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f592,f51])).
fof(f592,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X3))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X2))))))))) )),
  inference(superposition,[],[f42,f40])).
fof(f591,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))))) )),
  inference(superposition,[],[f41,f40])).
fof(f40,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X7))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X7))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X6)))))))),
  file('thm.rat.RAT_LEQ_CALCULATE','thm.rat.RAT_LES_CALCULATE')).
fof(f586,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5))))) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f40,f51])).
fof(f588,plain,(
  ( ! [X12,X10,X11] : (s(bool,X12) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X10))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X11))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X10))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X11))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X11))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X10))))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f40,f51])).
fof(f593,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X4))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X5))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X4))))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f37,f40])).
fof(f1853,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X27)) ) | $spl7),
  inference(subsumption_resolution,[],[f360,f36])).
fof(f360,plain,(
  ( ! [X27] : (p(s(bool,'F_0')) | s(bool,X27) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X27)) ) | $spl7),
  inference(superposition,[],[f193,f101])).
fof(f465,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,sK2(X19,X20,X21,bool)) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(X21,f(s(bool,'F_0'))) != s(X21,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f49,f37])).
fof(f101,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f95,f37])).
fof(f519,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | $spl7),
  inference(unit_resulting_resolution,[],[f99,f47])).
fof(f463,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,f(s(bool,X14))) != s(X13,g(s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK2(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f49,f51])).
fof(f461,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK2(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f49,f51])).
fof(f193,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f34,f51])).
fof(f141,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f51,f51])).
fof(f99,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | $spl7),
  inference(unit_resulting_resolution,[],[f95,f65])).
fof(f95,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl7),
  inference(backward_demodulation,[],[f91,f33])).
fof(f91,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl7),
  inference(unit_resulting_resolution,[],[f90,f37])).
fof(f80,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0)))) )),
  inference(equality_resolution,[],[f65])).
fof(f73,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f71,f53])).
fof(f71,plain,(
  ( ! [X0] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))))) )),
  inference(equality_resolution,[],[f46])).
fof(f65,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(subsumption_resolution,[],[f64,f36])).
fof(f64,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int',X0) != s('type.integer.int',X1) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(superposition,[],[f43,f37])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_LEQ_CALCULATE',aHL_EXT)).
fof(f47,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ((s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) & (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X7))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X6))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X7))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X7))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X6))))))),
  file('thm.rat.RAT_LEQ_CALCULATE','thm.rat.RAT_EQ')).
fof(f48,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f9510,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f33,f51])).
fof(f30658,plain,(
  ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | ($spl3 | $spl7 | $spl83 | $spl85)),
  inference(subsumption_resolution,[],[f30657,f11786])).
fof(f11786,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | $spl85),
  inference(unit_resulting_resolution,[],[f4749,f47])).
fof(f4749,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | $spl85),
  inference(cnf_transformation,[],[f4749_D])).
fof(f4749_D,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f30657,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | ($spl3 | $spl7 | $spl83)),
  inference(subsumption_resolution,[],[f30637,f36])).
fof(f30637,plain,(
  p(s(bool,'F_0')) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | ($spl3 | $spl7 | $spl83)),
  inference(superposition,[],[f44,f23358])).
fof(f23358,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl3 | $spl7 | $spl83)),
  inference(forward_demodulation,[],[f23357,f15831])).
fof(f15831,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f59,f9526])).
fof(f23357,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl83)),
  inference(forward_demodulation,[],[f23340,f40])).
fof(f23340,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl83)),
  inference(unit_resulting_resolution,[],[f4747,f9526])).
fof(f4747,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl83),
  inference(cnf_transformation,[],[f4747_D])).
fof(f4747_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f30597,plain,(
  ( ! [X59] : (s(bool,X59) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'F_0') = s(bool,g(s(bool,X59)))) ) | $spl124),
  inference(cnf_transformation,[],[f30597_D])).
fof(f30597_D,plain,(
  ( ! [X59] : (s(bool,X59) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'F_0') = s(bool,g(s(bool,X59)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f30353,plain,(
  ( ! [X111,X110] : (s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,g(s(bool,X111))) | s(bool,X111) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl7 | $spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30328])).
fof(f30328,plain,(
  ( ! [X111,X110] : (s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,g(s(bool,X111))) | s(bool,X111) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl7 | $spl10 | $spl30)),
  inference(superposition,[],[f9526,f23309])).
fof(f30620,plain,(
  ( ! [X72] : (s(bool,X72) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'F_0') = s(bool,g(s(bool,X72)))) ) | $spl132),
  inference(cnf_transformation,[],[f30620_D])).
fof(f30620_D,plain,(
  ( ! [X72] : (s(bool,X72) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'F_0') = s(bool,g(s(bool,X72)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f30381,plain,(
  ( ! [X72,X71] : (s(bool,X72) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X72)))) ) | ($spl7 | $spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30218])).
fof(f30218,plain,(
  ( ! [X72,X71] : (s(bool,X72) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X72))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl7 | $spl10 | $spl30)),
  inference(superposition,[],[f21003,f23309])).
fof(f30382,plain,(
  ( ! [X70,X69] : (s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,g(s(bool,X70)))) ) | ($spl7 | $spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30217])).
fof(f30217,plain,(
  ( ! [X70,X69] : (s(bool,X70) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,g(s(bool,X70))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl7 | $spl10 | $spl30)),
  inference(superposition,[],[f9526,f23309])).
fof(f30618,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) | $spl130),
  inference(cnf_transformation,[],[f30618_D])).
fof(f30618_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f30401,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30202])).
fof(f30202,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f53,f23309])).
fof(f30617,plain,(
  ( ! [X136,X137] : (s('type.integer.int',X136) = s('type.integer.int',X137) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) ) | $spl128),
  inference(cnf_transformation,[],[f30617_D])).
fof(f30617_D,plain,(
  ( ! [X136,X137] : (s('type.integer.int',X136) = s('type.integer.int',X137) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f30616,plain,(
  ( ! [X138,X136,X137] : (s('type.integer.int',X136) = s('type.integer.int',X137) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X136),s('type.integer.int',X137)))) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) | s(bool,'T_0') = s(bool,X138)) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f30190,f53])).
fof(f30190,plain,(
  ( ! [X138,X136,X137] : (p(s(bool,X138)) | s('type.integer.int',X136) = s('type.integer.int',X137) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X136),s('type.integer.int',X137)))) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X136),s('type.integer.int',X137))))) | s(bool,'T_0') = s(bool,X138)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f41,f23309])).
fof(f30615,plain,(
  ( ! [X130,X131] : (s('type.integer.int',X130) != s('type.integer.int',X131) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X130),s('type.integer.int',X131)))))) ) | $spl126),
  inference(cnf_transformation,[],[f30615_D])).
fof(f30615_D,plain,(
  ( ! [X130,X131] : (s('type.integer.int',X130) != s('type.integer.int',X131) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X130),s('type.integer.int',X131)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f30614,plain,(
  ( ! [X132,X130,X131] : (s('type.integer.int',X130) != s('type.integer.int',X131) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X130),s('type.integer.int',X131))))) | s(bool,'T_0') = s(bool,X132)) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f30187,f53])).
fof(f30187,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,X132)) | s('type.integer.int',X130) != s('type.integer.int',X131) | s(bool,'F_0') = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X130),s('type.integer.int',X131))))) | s(bool,'T_0') = s(bool,X132)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f43,f23309])).
fof(f30452,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X77))) | s(bool,X77) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl7 | $spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30106])).
fof(f30106,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X77))) | s(bool,X77) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl7 | $spl10 | $spl30)),
  inference(superposition,[],[f23309,f9526])).
fof(f30584,plain,(
  ( ! [X71,X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X71)))) ) | $spl116),
  inference(cnf_transformation,[],[f30584_D])).
fof(f30584_D,plain,(
  ( ! [X71,X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X71)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f30456,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,g(s(bool,X50))) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,g(s(bool,X48)))) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30098])).
fof(f30098,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,g(s(bool,X50))) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X49)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f23309])).
fof(f30487,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,g(s(bool,X65))) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,g(s(bool,X64)))) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30043])).
fof(f30043,plain,(
  ( ! [X64,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,g(s(bool,X65))) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,g(s(bool,X64))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f23309])).
fof(f30488,plain,(
  ( ! [X59,X58] : (s(bool,X59) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'F_0') = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl7 | $spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30041])).
fof(f30041,plain,(
  ( ! [X59,X58] : (s(bool,X59) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'F_0') = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl7 | $spl10 | $spl30)),
  inference(superposition,[],[f23309,f9526])).
fof(f30587,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,g(s(bool,X69)))) ) | $spl120),
  inference(cnf_transformation,[],[f30587_D])).
fof(f30587_D,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,g(s(bool,X69)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f30596,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') = s(bool,X57) | s(bool,'F_0') = s(bool,g(s(bool,X57))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl10 | $spl30 | $spl32)),
  inference(forward_demodulation,[],[f30489,f540])).
fof(f30489,plain,(
  ( ! [X57,X56] : (s(bool,X57) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X57))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl10 | $spl30 | $spl32)),
  inference(duplicate_literal_removal,[],[f30040])).
fof(f30040,plain,(
  ( ! [X57,X56] : (s(bool,X57) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X57))) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | ($spl10 | $spl30 | $spl32)),
  inference(superposition,[],[f23309,f581])).
fof(f30595,plain,(
  ( ! [X55] : (s(bool,X55) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X55)))) ) | $spl122),
  inference(cnf_transformation,[],[f30595_D])).
fof(f30595_D,plain,(
  ( ! [X55] : (s(bool,X55) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X55)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f30490,plain,(
  ( ! [X54,X55] : (s(bool,X55) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X55))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl10 | $spl22 | $spl30)),
  inference(duplicate_literal_removal,[],[f30039])).
fof(f30039,plain,(
  ( ! [X54,X55] : (s(bool,X55) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X55))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | ($spl10 | $spl22 | $spl30)),
  inference(superposition,[],[f23309,f534])).
fof(f30507,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30026])).
fof(f30026,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f37])).
fof(f30532,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,g(s(bool,X66))) | s(bool,'F_0') = s(bool,g(s(bool,X68))) | s(bool,X66) = s(bool,X68)) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f30004])).
fof(f30004,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,g(s(bool,X66))) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') = s(bool,g(s(bool,X68))) | s(bool,X66) = s(bool,X68)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f23309])).
fof(f30586,plain,(
  ( ! [X70,X68] : (s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X68))) | s(bool,'T_0') = s(bool,X70)) ) | $spl118),
  inference(cnf_transformation,[],[f30586_D])).
fof(f30586_D,plain,(
  ( ! [X70,X68] : (s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X68))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f29984,plain,(
  ( ! [X70,X68,X69] : (s(bool,'F_0') = s(bool,X69) | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,g(s(bool,X68)))) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f23309])).
fof(f30560,plain,(
  ( ! [X70,X71,X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X71))) | s(bool,X69) = s(bool,X71)) ) | ($spl10 | $spl30)),
  inference(duplicate_literal_removal,[],[f29944])).
fof(f29944,plain,(
  ( ! [X70,X71,X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X71))) | s(bool,X69) = s(bool,X71)) ) | ($spl10 | $spl30)),
  inference(superposition,[],[f23309,f23309])).
fof(f28790,plain,(
  $false | ($spl27 | $spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28690,f34])).
fof(f28690,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl62 | $spl112)),
  inference(backward_demodulation,[],[f28689,f36])).
fof(f28689,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl62 | $spl112)),
  inference(condensation,[],[f28688])).
fof(f28688,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28687,f25156])).
fof(f28687,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl62)),
  inference(condensation,[],[f28686])).
fof(f28686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl27 | $spl62)),
  inference(condensation,[],[f28685])).
fof(f28685,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl27 | $spl62)),
  inference(subsumption_resolution,[],[f28596,f23379])).
fof(f23379,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl27),
  inference(superposition,[],[f537,f51])).
fof(f28596,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl62),
  inference(superposition,[],[f2870,f141])).
fof(f28684,plain,(
  $false | ($spl27 | $spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28683,f537])).
fof(f28683,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl62 | $spl112)),
  inference(duplicate_literal_removal,[],[f28682])).
fof(f28682,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl62 | $spl112)),
  inference(condensation,[],[f28681])).
fof(f28681,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28595,f25156])).
fof(f28595,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl62),
  inference(superposition,[],[f2870,f51])).
fof(f28680,plain,(
  $false | ($spl27 | $spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28679,f537])).
fof(f28679,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl62 | $spl112)),
  inference(duplicate_literal_removal,[],[f28678])).
fof(f28678,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl62 | $spl112)),
  inference(condensation,[],[f28677])).
fof(f28677,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl62 | $spl112)),
  inference(subsumption_resolution,[],[f28592,f25156])).
fof(f28592,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl62),
  inference(superposition,[],[f2870,f51])).
fof(f28676,plain,(
  $false | ($spl3 | $spl27 | $spl62)),
  inference(subsumption_resolution,[],[f28588,f537])).
fof(f28588,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl62)),
  inference(unit_resulting_resolution,[],[f59,f2870])).
fof(f28675,plain,(
  $false | ($spl3 | $spl27 | $spl62)),
  inference(subsumption_resolution,[],[f28589,f59])).
fof(f28589,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl27 | $spl62)),
  inference(unit_resulting_resolution,[],[f537,f2870])).
fof(f28590,plain,(
  $false | ($spl3 | $spl27 | $spl62)),
  inference(unit_resulting_resolution,[],[f59,f537,f2870])).
fof(f28256,plain,(
  $false | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28156,f34])).
fof(f28156,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl68)),
  inference(backward_demodulation,[],[f28155,f36])).
fof(f28155,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl68)),
  inference(condensation,[],[f28154])).
fof(f28154,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28153,f23379])).
fof(f28153,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl68),
  inference(duplicate_literal_removal,[],[f28152])).
fof(f28152,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl68),
  inference(condensation,[],[f28151])).
fof(f28151,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl68),
  inference(subsumption_resolution,[],[f28052,f51])).
fof(f28052,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'T_0') != s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl68),
  inference(superposition,[],[f2919,f141])).
fof(f28150,plain,(
  $false | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28149,f537])).
fof(f28149,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(duplicate_literal_removal,[],[f28148])).
fof(f28148,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(condensation,[],[f28147])).
fof(f28147,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl68),
  inference(subsumption_resolution,[],[f28051,f51])).
fof(f28051,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl68),
  inference(superposition,[],[f2919,f51])).
fof(f28146,plain,(
  $false | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28145,f537])).
fof(f28145,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(duplicate_literal_removal,[],[f28144])).
fof(f28144,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(condensation,[],[f28143])).
fof(f28143,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl68),
  inference(subsumption_resolution,[],[f28048,f51])).
fof(f28048,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl68),
  inference(superposition,[],[f2919,f51])).
fof(f28142,plain,(
  $false | ($spl3 | $spl27 | $spl68 | $spl88)),
  inference(subsumption_resolution,[],[f28141,f537])).
fof(f28141,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl68 | $spl88)),
  inference(forward_demodulation,[],[f28042,f6065])).
fof(f6065,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(fun(bool,bool),X0)))) ) | ($spl3 | $spl88)),
  inference(unit_resulting_resolution,[],[f59,f6041,f51])).
fof(f28042,plain,(
  ( ! [X0] : (s(bool,f(s(fun(bool,bool),X0))) = s(bool,g(s(bool,f(s(fun(bool,bool),X0)))))) ) | ($spl3 | $spl68 | $spl88)),
  inference(unit_resulting_resolution,[],[f6065,f2919])).
fof(f28140,plain,(
  $false | ($spl14 | $spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28139,f537])).
fof(f28139,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl14 | $spl68)),
  inference(forward_demodulation,[],[f28043,f511])).
fof(f28043,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,f(s(bool,'F_0'))))) | ($spl14 | $spl68)),
  inference(unit_resulting_resolution,[],[f511,f2919])).
fof(f28138,plain,(
  $false | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28137,f537])).
fof(f28137,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(forward_demodulation,[],[f28044,f73])).
fof(f28044,plain,(
  ( ! [X0] : (s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))) = s(bool,g(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f73,f2919])).
fof(f28136,plain,(
  $false | ($spl27 | $spl68)),
  inference(subsumption_resolution,[],[f28135,f537])).
fof(f28135,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl68),
  inference(forward_demodulation,[],[f28045,f80])).
fof(f28045,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0)))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f80,f2919])).
fof(f28134,plain,(
  $false | ($spl27 | $spl68)),
  inference(trivial_inequality_removal,[],[f28046])).
fof(f28046,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl27 | $spl68)),
  inference(unit_resulting_resolution,[],[f537,f2919])).
fof(f28123,plain,(
  $false | ($spl27 | $spl68)),
  inference(trivial_inequality_removal,[],[f28122])).
fof(f28122,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl27 | $spl68)),
  inference(duplicate_literal_removal,[],[f28099])).
fof(f28099,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') != s(bool,'T_0') | ($spl27 | $spl68)),
  inference(superposition,[],[f537,f2919])).
fof(f27560,plain,(
  $false | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27459,f34])).
fof(f27459,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(backward_demodulation,[],[f27458,f36])).
fof(f27458,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(condensation,[],[f27457])).
fof(f27457,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27456,f25156])).
fof(f27456,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl32 | $spl38)),
  inference(condensation,[],[f27455])).
fof(f27455,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl27 | $spl32 | $spl38)),
  inference(condensation,[],[f27454])).
fof(f27454,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl27 | $spl32 | $spl38)),
  inference(subsumption_resolution,[],[f27366,f26840])).
fof(f26840,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl32)),
  inference(condensation,[],[f26839])).
fof(f26839,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl27 | $spl32)),
  inference(duplicate_literal_removal,[],[f26806])).
fof(f26806,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl27 | $spl32)),
  inference(superposition,[],[f26568,f141])).
fof(f26568,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl27 | $spl32)),
  inference(subsumption_resolution,[],[f26443,f537])).
fof(f26443,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl27 | $spl32)),
  inference(superposition,[],[f537,f581])).
fof(f27366,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) = s(bool,f(s(bool,X10))) | s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl38),
  inference(superposition,[],[f1242,f141])).
fof(f27453,plain,(
  $false | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27452,f26568])).
fof(f27452,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl38 | $spl112)),
  inference(duplicate_literal_removal,[],[f27451])).
fof(f27451,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl38 | $spl112)),
  inference(condensation,[],[f27450])).
fof(f27450,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27365,f25156])).
fof(f27365,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl38),
  inference(superposition,[],[f1242,f51])).
fof(f27449,plain,(
  $false | ($spl27 | $spl32 | $spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27448,f26568])).
fof(f27448,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl38 | $spl112)),
  inference(duplicate_literal_removal,[],[f27447])).
fof(f27447,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl38 | $spl112)),
  inference(condensation,[],[f27446])).
fof(f27446,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl38 | $spl112)),
  inference(subsumption_resolution,[],[f27362,f25156])).
fof(f27362,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl38),
  inference(superposition,[],[f1242,f51])).
fof(f27445,plain,(
  $false | ($spl3 | $spl27 | $spl32 | $spl38)),
  inference(subsumption_resolution,[],[f27358,f26568])).
fof(f27358,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl38)),
  inference(unit_resulting_resolution,[],[f59,f1242])).
fof(f27444,plain,(
  $false | ($spl3 | $spl27 | $spl32 | $spl38)),
  inference(subsumption_resolution,[],[f27359,f59])).
fof(f27359,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl27 | $spl32 | $spl38)),
  inference(unit_resulting_resolution,[],[f26568,f1242])).
fof(f27360,plain,(
  $false | ($spl3 | $spl27 | $spl32 | $spl38)),
  inference(unit_resulting_resolution,[],[f59,f26568,f1242])).
fof(f27013,plain,(
  $false | ($spl27 | $spl32 | $spl96)),
  inference(subsumption_resolution,[],[f26911,f34])).
fof(f26911,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl32 | $spl96)),
  inference(backward_demodulation,[],[f26910,f36])).
fof(f26910,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl32 | $spl96)),
  inference(condensation,[],[f26909])).
fof(f26909,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl27 | $spl32 | $spl96)),
  inference(global_subsumption,[],[f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f141,f193,f461,f463,f33,f9508,f581,f537,f23379,f25081,f24925,f24923,f24917,f23616,f25155,f23781,f25176,f25202,f23877,f24443,f23925,f23928,f23931,f23976,f25284,f24164,f25434,f23646,f23614,f23585,f23527,f588,f586,f40,f25243,f25240,f720,f664,f592,f591,f25241,f25239,f25204,f25200,f18291,f459,f342,f108,f43,f42,f41,f25242,f25238,f25206,f25195,f20717,f18728,f377,f343,f46,f45,f44,f25690,f25193,f23978,f23927,f23879,f23783,f6038,f469,f467,f24921,f24919,f21153,f7476,f5026,f593,f471,f465,f36,f26387,f26389,f26568,f26797,f26799])).
fof(f26799,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl27 | $spl32)),
  inference(superposition,[],[f26568,f51])).
fof(f26797,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | ($spl27 | $spl32)),
  inference(superposition,[],[f26568,f141])).
fof(f26389,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl32),
  inference(superposition,[],[f581,f51])).
fof(f26387,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,X2)) ) | $spl32),
  inference(superposition,[],[f581,f141])).
fof(f24919,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK2(X2,X3,X4,bool)) | s(bool,X0) = s(bool,X1) | s(fun(bool,X4),X2) = s(fun(bool,X4),X3) | s(X4,f(s(bool,'F_0'))) != s(X4,g(s(bool,'F_0')))) )),
  inference(condensation,[],[f24918])).
fof(f24918,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,sK2(X147,X148,X149,bool)) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(fun(bool,X149),X147) = s(fun(bool,X149),X148) | s(X149,f(s(bool,'F_0'))) != s(X149,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f23586])).
fof(f23586,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,sK2(X147,X148,X149,bool)) | s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(fun(bool,X149),X147) = s(fun(bool,X149),X148) | s(X149,f(s(bool,'F_0'))) != s(X149,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f141,f465])).
fof(f24921,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X2))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X3))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X2))))))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3)))))) )),
  inference(condensation,[],[f24920])).
fof(f24920,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'T_0') = s(bool,X139) | s(bool,X140) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X137))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X138))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X138))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X137))))))) | s(bool,X140) = s(bool,X141) | s(bool,'T_0') = s(bool,X141) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X137))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X138)))))) )),
  inference(duplicate_literal_removal,[],[f23584])).
fof(f23584,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'T_0') = s(bool,X139) | s(bool,X140) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X137))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X138))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X138))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X137))))))) | s(bool,'T_0') = s(bool,X139) | s(bool,X140) = s(bool,X141) | s(bool,'T_0') = s(bool,X141) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X137))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X138)))))) )),
  inference(superposition,[],[f141,f593])).
fof(f23783,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (s(bool,X218) != s(bool,f(s(X215,sK2(X216,X217,bool,X215)))) | s(fun(X215,bool),X216) = s(fun(X215,bool),X217) | s(bool,X218) = s(bool,X219) | s(bool,'T_0') = s(bool,X219) | s(bool,X220) = s(bool,g(s(X215,sK2(X216,X217,bool,X215)))) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f49,f141])).
fof(f23879,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (s(bool,X218) != s(bool,f(s(X215,sK2(X216,X217,bool,X215)))) | s(fun(X215,bool),X216) = s(fun(X215,bool),X217) | s(bool,X219) = s(bool,g(s(X215,sK2(X216,X217,bool,X215)))) | s(bool,'T_0') = s(bool,X218) | s(bool,X219) = s(bool,X220) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f49,f141])).
fof(f23927,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (s(bool,X212) != s(bool,f(s(X209,sK2(X210,X211,bool,X209)))) | s(fun(X209,bool),X210) = s(fun(X209,bool),X211) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,g(s(X209,sK2(X210,X211,bool,X209)))) | s(bool,X213) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f49,f141])).
fof(f23978,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,'T_0') != s(bool,f(s(X185,sK2(X186,X187,bool,X185)))) | s(fun(X185,bool),X186) = s(fun(X185,bool),X187) | s(bool,X188) = s(bool,X189) | s(bool,X189) = s(bool,g(s(X185,sK2(X186,X187,bool,X185)))) | s(bool,X188) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f49,f141])).
fof(f25193,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25192])).
fof(f25192,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,X94) = s(bool,f(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f23853,f37])).
fof(f23853,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') != s(bool,X93) | s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,X94) = s(bool,f(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f471,f141])).
fof(f25690,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25689])).
fof(f25689,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'T_0') = s(bool,f(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f25642])).
fof(f25642,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'T_0') = s(bool,f(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f6038,f141])).
fof(f25195,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s(bool,X3) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25194])).
fof(f25194,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (s(bool,X120) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X117),s('type.rat.rat',X118))) | s(bool,'T_0') = s(bool,X120) | s(bool,X121) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X117),s('type.rat.rat',X118))) | s(bool,'T_0') = s(bool,X119) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f23860,f53])).
fof(f23860,plain,(
  ( ! [X121,X118,X122,X120,X119,X117] : (p(s(bool,X119)) | s(bool,X120) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X117),s('type.rat.rat',X118))) | s(bool,'T_0') = s(bool,X120) | s(bool,X121) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X117),s('type.rat.rat',X118))) | s(bool,'T_0') = s(bool,X119) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f343,f141])).
fof(f25206,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.rat.rat',X1) = s('type.rat.rat',X2) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) | s(bool,X3) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s(bool,X0) = s(bool,X3)) )),
  inference(condensation,[],[f25205])).
fof(f25205,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s('type.rat.rat',X179) = s('type.rat.rat',X180) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X179),s('type.rat.rat',X180)))) | s(bool,X182) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X179),s('type.rat.rat',X180))) | s(bool,'T_0') = s(bool,X181) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) )),
  inference(subsumption_resolution,[],[f23873,f53])).
fof(f23873,plain,(
  ( ! [X182,X180,X183,X181,X179] : (p(s(bool,X181)) | s('type.rat.rat',X179) = s('type.rat.rat',X180) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X179),s('type.rat.rat',X180)))) | s(bool,X182) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X179),s('type.rat.rat',X180))) | s(bool,'T_0') = s(bool,X181) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f44,f141])).
fof(f25238,plain,(
  ( ! [X94,X92,X95,X93,X96] : (s('type.rat.rat',X92) = s('type.rat.rat',X93) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X92),s('type.rat.rat',X93))) | s(bool,X94) = s(bool,X95) | s(bool,X95) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X92),s('type.rat.rat',X93))) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f23961,f34])).
fof(f23961,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,'T_0')) | s('type.rat.rat',X92) = s('type.rat.rat',X93) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X92),s('type.rat.rat',X93))) | s(bool,X94) = s(bool,X95) | s(bool,X95) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X92),s('type.rat.rat',X93))) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f377,f141])).
fof(f25242,plain,(
  ( ! [X158,X161,X159,X162,X160] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X158),s('type.rat.rat',X159)))) | s(bool,X160) = s(bool,X161) | s(bool,X161) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X158),s('type.rat.rat',X159))) | s(bool,X160) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(subsumption_resolution,[],[f23974,f34])).
fof(f23974,plain,(
  ( ! [X158,X161,X159,X162,X160] : (~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X158),s('type.rat.rat',X159)))) | s(bool,X160) = s(bool,X161) | s(bool,X161) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X158),s('type.rat.rat',X159))) | s(bool,X160) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f45,f141])).
fof(f25200,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s(bool,X3) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25199])).
fof(f25199,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s(bool,X134) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X131),s('type.integer.int',X132))) | s(bool,'T_0') = s(bool,X134) | s(bool,X135) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int',X132))) | s(bool,'T_0') = s(bool,X133) | s(bool,X135) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f23863,f53])).
fof(f23863,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (p(s(bool,X133)) | s(bool,X134) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X131),s('type.integer.int',X132))) | s(bool,'T_0') = s(bool,X134) | s(bool,X135) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X131),s('type.integer.int',X132))) | s(bool,'T_0') = s(bool,X133) | s(bool,X135) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f342,f141])).
fof(f25204,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.integer.int',X1) = s('type.integer.int',X2) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s(bool,X0) = s(bool,X3)) )),
  inference(condensation,[],[f25203])).
fof(f25203,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s('type.integer.int',X156) = s('type.integer.int',X157) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X156),s('type.integer.int',X157)))) | s(bool,X159) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X156),s('type.integer.int',X157))) | s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(subsumption_resolution,[],[f23868,f53])).
fof(f23868,plain,(
  ( ! [X158,X156,X159,X157,X160] : (p(s(bool,X158)) | s('type.integer.int',X156) = s('type.integer.int',X157) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X156),s('type.integer.int',X157)))) | s(bool,X159) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X156),s('type.integer.int',X157))) | s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f41,f141])).
fof(f25239,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s('type.integer.int',X114) = s('type.integer.int',X115) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X114),s('type.integer.int',X115))) | s(bool,X116) = s(bool,X117) | s(bool,X117) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X114),s('type.integer.int',X115))) | s(bool,X116) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f23965,f34])).
fof(f23965,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,'T_0')) | s('type.integer.int',X114) = s('type.integer.int',X115) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X114),s('type.integer.int',X115))) | s(bool,X116) = s(bool,X117) | s(bool,X117) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X114),s('type.integer.int',X115))) | s(bool,X116) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f108,f141])).
fof(f25241,plain,(
  ( ! [X142,X140,X138,X141,X139] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X138),s('type.integer.int',X139)))) | s(bool,X140) = s(bool,X141) | s(bool,X141) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(subsumption_resolution,[],[f23970,f34])).
fof(f23970,plain,(
  ( ! [X142,X140,X138,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X138),s('type.integer.int',X139)))) | s(bool,X140) = s(bool,X141) | s(bool,X141) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X138),s('type.integer.int',X139))) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f42,f141])).
fof(f25240,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X124))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X125))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X125))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X124))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X124))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X125)))))) | s(bool,X126) = s(bool,X127) | s(bool,X127) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X124))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X125))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X125))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X124))))))) | s(bool,X126) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(subsumption_resolution,[],[f23967,f34])).
fof(f23967,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X124))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X125))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X125))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X124))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X124))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X125)))))) | s(bool,X126) = s(bool,X127) | s(bool,X127) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X124))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X125))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X125))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X124))))))) | s(bool,X126) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f591,f141])).
fof(f25243,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X163))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X164))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X164))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X163)))))))) | s(bool,X165) = s(bool,X166) | s(bool,X166) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X163))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X164))))) | s(bool,X165) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(subsumption_resolution,[],[f23975,f34])).
fof(f23975,plain,(
  ( ! [X167,X165,X163,X166,X164] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X163))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X164))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X164))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X163)))))))) | s(bool,X165) = s(bool,X166) | s(bool,X166) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X163))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X164))))) | s(bool,X165) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f592,f141])).
fof(f23527,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s(bool,X144) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X142))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X143))))) | s(bool,X144) = s(bool,X145) | s(bool,'T_0') = s(bool,X145) | s(bool,X146) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X142))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X143))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X143))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X142))))))) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f141,f40])).
fof(f23585,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s(bool,X144) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X142))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X143))))) | s(bool,X145) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X142))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X143))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X143))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X142))))))) | s(bool,'T_0') = s(bool,X144) | s(bool,X145) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f141,f40])).
fof(f23614,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X132) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X130))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X131))))) | s(bool,X132) = s(bool,X133) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X130))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X131))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X131))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X130))))))) | s(bool,X133) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f141,f40])).
fof(f23646,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X87))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X88))))) | s(bool,X89) = s(bool,X90) | s(bool,X90) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X87))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X88))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X88))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X87))))))) | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f141,f40])).
fof(f25434,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(X1,f(s(bool,'T_0'))) != s(X1,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5026,f37])).
fof(f24164,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f141])).
fof(f25284,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f24157,f51])).
fof(f24157,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f141])).
fof(f23976,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(X176,f(s(bool,'T_0'))) != s(X176,g(s(bool,'T_0'))) | s(fun(bool,X176),X174) = s(fun(bool,X176),X175) | s(bool,X177) = s(bool,X178) | s(bool,X178) = s(bool,sK2(X174,X175,X176,bool)) | s(bool,X177) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f49,f141])).
fof(f23931,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X17) != s(bool,g(s(bool,X17))) | s(bool,X18) = s(bool,X19) | s(bool,X17) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl27),
  inference(superposition,[],[f537,f141])).
fof(f23928,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f34,f141])).
fof(f23925,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (s(X200,f(s(bool,X201))) != s(X200,g(s(bool,X201))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,X201) = s(bool,X202) | s(bool,'T_0') = s(bool,sK2(X198,X199,X200,bool)) | s(bool,X202) = s(bool,X203) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f49,f141])).
fof(f24443,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f24442])).
fof(f24442,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f23885])).
fof(f23885,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f53,f141])).
fof(f23877,plain,(
  ( ! [X208,X206,X204,X209,X207,X205] : (s(X206,f(s(bool,X207))) != s(X206,g(s(bool,X207))) | s(fun(bool,X206),X204) = s(fun(bool,X206),X205) | s(bool,X208) = s(bool,sK2(X204,X205,X206,bool)) | s(bool,'T_0') = s(bool,X207) | s(bool,X208) = s(bool,X209) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f49,f141])).
fof(f25202,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.integer.int',X1) != s('type.integer.int',X2) | s(bool,X3) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s(bool,X0) = s(bool,X3)) )),
  inference(condensation,[],[f25201])).
fof(f25201,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s('type.integer.int',X142) != s('type.integer.int',X143) | s(bool,X145) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X142),s('type.integer.int',X143))) | s(bool,'T_0') = s(bool,X144) | s(bool,X145) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(subsumption_resolution,[],[f23865,f53])).
fof(f23865,plain,(
  ( ! [X146,X144,X142,X145,X143] : (p(s(bool,X144)) | s('type.integer.int',X142) != s('type.integer.int',X143) | s(bool,X145) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X142),s('type.integer.int',X143))) | s(bool,'T_0') = s(bool,X144) | s(bool,X145) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f43,f141])).
fof(f23781,plain,(
  ( ! [X208,X206,X204,X209,X207,X205] : (s(X206,f(s(bool,X207))) != s(X206,g(s(bool,X207))) | s(fun(bool,X206),X204) = s(fun(bool,X206),X205) | s(bool,X207) = s(bool,X208) | s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,sK2(X204,X205,X206,bool)) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f49,f141])).
fof(f23616,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f141,f141])).
fof(f24917,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f24916])).
fof(f24916,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f23587])).
fof(f23587,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f141,f37])).
fof(f24923,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X2),s('type.integer.int',X2))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f24922])).
fof(f24922,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X129),s('type.integer.int',X129))) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(duplicate_literal_removal,[],[f23582])).
fof(f23582,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X129),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f141,f80])).
fof(f24925,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X2),s('type.rat.rat',X2))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f24924])).
fof(f24924,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') = s(bool,X126) | s(bool,X127) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X125),s('type.rat.rat',X125))) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(duplicate_literal_removal,[],[f23581])).
fof(f23581,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') = s(bool,X126) | s(bool,X127) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X125),s('type.rat.rat',X125))) | s(bool,'T_0') = s(bool,X126) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f141,f73])).
fof(f25081,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f25080])).
fof(f25080,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,X84) = s(bool,X88)) )),
  inference(duplicate_literal_removal,[],[f23514])).
fof(f23514,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,X84) = s(bool,X88)) )),
  inference(superposition,[],[f141,f141])).
fof(f536,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl25),
  inference(cnf_transformation,[],[f536_D])).
fof(f536_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f26567,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl7 | $spl32)),
  inference(forward_demodulation,[],[f26488,f22723])).
fof(f26488,plain,(
  ( ! [X47] : (s(bool,f(s(bool,'T_0'))) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl7 | $spl32)),
  inference(duplicate_literal_removal,[],[f26437])).
fof(f26437,plain,(
  ( ! [X47] : (s(bool,f(s(bool,'T_0'))) != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) ) | ($spl7 | $spl32)),
  inference(superposition,[],[f21003,f581])).
fof(f25837,plain,(
  $false | ($spl27 | $spl66)),
  inference(subsumption_resolution,[],[f25734,f34])).
fof(f25734,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl66)),
  inference(backward_demodulation,[],[f25733,f36])).
fof(f25733,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl66)),
  inference(subsumption_resolution,[],[f25732,f34])).
fof(f25732,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | ($spl27 | $spl66)),
  inference(condensation,[],[f25731])).
fof(f25731,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl27 | $spl66)),
  inference(subsumption_resolution,[],[f25700,f23379])).
fof(f25700,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl66),
  inference(superposition,[],[f2917,f51])).
fof(f25730,plain,(
  $false | ($spl27 | $spl66)),
  inference(subsumption_resolution,[],[f25696,f537])).
fof(f25696,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl66),
  inference(unit_resulting_resolution,[],[f34,f2917])).
fof(f25729,plain,(
  $false | ($spl27 | $spl66)),
  inference(subsumption_resolution,[],[f25697,f34])).
fof(f25697,plain,(
  ~p(s(bool,'T_0')) | ($spl27 | $spl66)),
  inference(unit_resulting_resolution,[],[f537,f2917])).
fof(f25698,plain,(
  $false | ($spl27 | $spl66)),
  inference(unit_resulting_resolution,[],[f34,f537,f2917])).
fof(f25513,plain,(
  $false | ($spl3 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f25512,f59])).
fof(f25512,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl30 | $spl44)),
  inference(backward_demodulation,[],[f25482,f540])).
fof(f25482,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl44),
  inference(unit_resulting_resolution,[],[f34,f1258])).
fof(f25333,plain,(
  $false | ($spl3 | $spl10 | $spl30 | $spl68 | $spl78)),
  inference(subsumption_resolution,[],[f25332,f59])).
fof(f25332,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl30 | $spl68 | $spl78)),
  inference(backward_demodulation,[],[f25330,f23255])).
fof(f25330,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl68 | $spl78)),
  inference(subsumption_resolution,[],[f25289,f2919])).
fof(f25289,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl78),
  inference(superposition,[],[f3618,f37])).
fof(f25156,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X4)) ) | $spl112),
  inference(cnf_transformation,[],[f25156_D])).
fof(f25156_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X4)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f25157,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl114),
  inference(cnf_transformation,[],[f25157_D])).
fof(f25157_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f25176,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25175])).
fof(f25175,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') != s(bool,X93) | s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f23805,f469])).
fof(f23805,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') != s(bool,X93) | s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,X94) = s(bool,f(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f471,f141])).
fof(f25155,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25154])).
fof(f25154,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') != s(bool,X93) | s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f23757,f469])).
fof(f23757,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'F_0') != s(bool,X93) | s(fun(X90,bool),X91) = s(fun(X90,bool),X92) | s(bool,'T_0') = s(bool,g(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,f(s(X90,sK2(X91,X92,bool,X90)))) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f471,f141])).
fof(f23285,plain,(
  $false | ($spl3 | $spl7 | $spl82)),
  inference(subsumption_resolution,[],[f23284,f36])).
fof(f23284,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl7 | $spl82)),
  inference(forward_demodulation,[],[f23283,f15831])).
fof(f23283,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | $spl82),
  inference(subsumption_resolution,[],[f23266,f34])).
fof(f23266,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | $spl82),
  inference(superposition,[],[f42,f4748])).
fof(f23258,plain,(
  $false | ($spl3 | $spl16 | $spl78)),
  inference(subsumption_resolution,[],[f23257,f59])).
fof(f23257,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl16 | $spl78)),
  inference(forward_demodulation,[],[f3618,f514])).
fof(f23256,plain,(
  $false | ($spl10 | $spl30 | $spl107)),
  inference(subsumption_resolution,[],[f23255,f23217])).
fof(f23253,plain,(
  $false | ($spl3 | $spl10 | $spl30 | $spl68 | $spl104)),
  inference(subsumption_resolution,[],[f23252,f59])).
fof(f23252,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl30 | $spl68 | $spl104)),
  inference(forward_demodulation,[],[f540,f23251])).
fof(f23251,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl10 | $spl68 | $spl104)),
  inference(forward_demodulation,[],[f508,f23248])).
fof(f23248,plain,(
  ( ! [X104] : (s(bool,X104) = s(bool,g(s(bool,X104)))) ) | ($spl68 | $spl104)),
  inference(subsumption_resolution,[],[f23215,f2919])).
fof(f23246,plain,(
  $false | ($spl27 | $spl62 | $spl70)),
  inference(subsumption_resolution,[],[f23245,f537])).
fof(f23245,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl62 | $spl70)),
  inference(condensation,[],[f23244])).
fof(f23244,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | ($spl62 | $spl70)),
  inference(subsumption_resolution,[],[f4658,f2870])).
fof(f4658,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl70),
  inference(superposition,[],[f3080,f37])).
fof(f3620,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,X4)))) ) | $spl80),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,X4)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f23217,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | $spl107),
  inference(cnf_transformation,[],[f23217_D])).
fof(f23217_D,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f23241,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl68 | $spl70)),
  inference(subsumption_resolution,[],[f4606,f2919])).
fof(f4606,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl70),
  inference(superposition,[],[f3080,f37])).
fof(f23240,plain,(
  $false | ($spl10 | $spl24 | $spl27)),
  inference(subsumption_resolution,[],[f23239,f537])).
fof(f23239,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl10 | $spl24)),
  inference(forward_demodulation,[],[f508,f535])).
fof(f23238,plain,(
  $false | ($spl3 | $spl24 | $spl30)),
  inference(subsumption_resolution,[],[f23237,f59])).
fof(f23237,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl24 | $spl30)),
  inference(backward_demodulation,[],[f535,f540])).
fof(f23236,plain,(
  $false | ($spl3 | $spl7 | $spl27 | $spl70 | $spl94)),
  inference(subsumption_resolution,[],[f537,f23223])).
fof(f23223,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl7 | $spl70 | $spl94)),
  inference(condensation,[],[f23222])).
fof(f23222,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl7 | $spl70 | $spl94)),
  inference(condensation,[],[f23221])).
fof(f23221,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,g(s(bool,X13))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl3 | $spl7 | $spl70 | $spl94)),
  inference(subsumption_resolution,[],[f4661,f23207])).
fof(f23207,plain,(
  ( ! [X81,X82] : (s(bool,X81) = s(bool,X82) | s(bool,X82) = s(bool,g(s(bool,X82))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl3 | $spl7 | $spl94)),
  inference(subsumption_resolution,[],[f16300,f15831])).
fof(f16300,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X81) = s(bool,X82) | s(bool,X82) = s(bool,g(s(bool,X82))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl7 | $spl94)),
  inference(superposition,[],[f15497,f9526])).
fof(f4661,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,X13) = s(bool,g(s(bool,X13))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl70),
  inference(superposition,[],[f3080,f51])).
fof(f23235,plain,(
  $false | ($spl3 | $spl30 | $spl54)),
  inference(subsumption_resolution,[],[f23234,f59])).
fof(f23234,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl30 | $spl54)),
  inference(forward_demodulation,[],[f540,f1836])).
fof(f581,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f581_D])).
fof(f581_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f23232,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl110),
  inference(cnf_transformation,[],[f23232_D])).
fof(f23232_D,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2255,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,X10) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl46),
  inference(superposition,[],[f1434,f51])).
fof(f2258,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,X16) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X16)) ) | $spl46),
  inference(superposition,[],[f1434,f51])).
fof(f23215,plain,(
  ( ! [X104] : (s(bool,X104) = s(bool,g(s(bool,X104))) | s(bool,'T_0') = s(bool,X104)) ) | $spl104),
  inference(cnf_transformation,[],[f23215_D])).
fof(f23215_D,plain,(
  ( ! [X104] : (s(bool,X104) = s(bool,g(s(bool,X104))) | s(bool,'T_0') = s(bool,X104)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f23213,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X108)) ) | $spl102),
  inference(cnf_transformation,[],[f23213_D])).
fof(f23213_D,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X108)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4608,plain,(
  ( ! [X4,X3] : (s(bool,X4) != s(bool,g(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl70),
  inference(superposition,[],[f3080,f51])).
fof(f4609,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,g(s(bool,'T_0'))) | s(bool,X5) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | $spl70),
  inference(superposition,[],[f3080,f51])).
fof(f534,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f534_D])).
fof(f534_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2919,plain,(
  ( ! [X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,'T_0') != s(bool,X32)) ) | $spl68),
  inference(cnf_transformation,[],[f2919_D])).
fof(f2919_D,plain,(
  ( ! [X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,'T_0') != s(bool,X32)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4659,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | s(bool,X8) = s(bool,g(s(bool,X8))) | s(bool,X9) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl70),
  inference(superposition,[],[f3080,f51])).
fof(f4662,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | s(bool,X14) = s(bool,g(s(bool,X14))) | s(bool,X15) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl70),
  inference(superposition,[],[f3080,f51])).
fof(f23220,plain,(
  ( ! [X111,X112,X110] : (s(bool,sK2(X110,X111,X112,bool)) = s(bool,g(s(bool,sK2(X110,X111,X112,bool)))) | s(X112,f(s(bool,'F_0'))) != s(X112,g(s(bool,'F_0'))) | s(fun(bool,X112),X110) = s(fun(bool,X112),X111)) ) | $spl108),
  inference(cnf_transformation,[],[f23220_D])).
fof(f23220_D,plain,(
  ( ! [X111,X112,X110] : (s(bool,sK2(X110,X111,X112,bool)) = s(bool,g(s(bool,sK2(X110,X111,X112,bool)))) | s(X112,f(s(bool,'F_0'))) != s(X112,g(s(bool,'F_0'))) | s(fun(bool,X112),X110) = s(fun(bool,X112),X111)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f537,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl27),
  inference(cnf_transformation,[],[f537_D])).
fof(f537_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f5020,plain,(
  ( ! [X111,X112,X110] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,sK2(X110,X111,X112,bool)) = s(bool,g(s(bool,sK2(X110,X111,X112,bool)))) | s(fun(bool,X112),X110) = s(fun(bool,X112),X111) | s(X112,f(s(bool,'F_0'))) != s(X112,g(s(bool,'F_0')))) ) | $spl70),
  inference(superposition,[],[f3080,f465])).
fof(f23214,plain,(
  ( ! [X104] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X104) = s(bool,g(s(bool,X104))) | s(bool,'T_0') = s(bool,X104)) ) | ($spl3 | $spl7 | $spl70)),
  inference(forward_demodulation,[],[f15984,f15831])).
fof(f15984,plain,(
  ( ! [X104] : (s(bool,g(s(bool,'T_0'))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X104) = s(bool,g(s(bool,X104))) | s(bool,'T_0') = s(bool,X104)) ) | ($spl7 | $spl70)),
  inference(superposition,[],[f3080,f9526])).
fof(f3618,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl78),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f23212,plain,(
  ( ! [X108] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X108) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X108)) ) | ($spl3 | $spl7 | $spl70)),
  inference(forward_demodulation,[],[f16110,f15831])).
fof(f16110,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,g(s(bool,'T_0'))) | s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,'T_0') = s(bool,X108)) ) | ($spl7 | $spl70)),
  inference(superposition,[],[f3080,f9526])).
fof(f23206,plain,(
  ( ! [X166] : (s(bool,X166) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X166) = s(bool,g(s(bool,X166)))) ) | $spl100),
  inference(cnf_transformation,[],[f23206_D])).
fof(f23206_D,plain,(
  ( ! [X166] : (s(bool,X166) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X166) = s(bool,g(s(bool,X166)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f23205,plain,(
  ( ! [X165,X166] : (s(bool,X166) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X166) = s(bool,g(s(bool,X166))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl7 | $spl94)),
  inference(subsumption_resolution,[],[f16334,f37])).
fof(f16334,plain,(
  ( ! [X165,X166] : (s(bool,'F_0') != s(bool,X165) | s(bool,X166) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X166) = s(bool,g(s(bool,X166))) | s(bool,'T_0') = s(bool,X165)) ) | ($spl7 | $spl94)),
  inference(superposition,[],[f15497,f9526])).
fof(f23181,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl98),
  inference(cnf_transformation,[],[f23181_D])).
fof(f23181_D,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f539,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2218,plain,(
  ( ! [X6,X5] : (s(bool,X6) != s(bool,f(s(bool,'T_0'))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | $spl46),
  inference(superposition,[],[f1434,f51])).
fof(f2217,plain,(
  ( ! [X4,X3] : (s(bool,X4) != s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl46),
  inference(superposition,[],[f1434,f51])).
fof(f23135,plain,(
  $false | ($spl3 | $spl7 | $spl26 | $spl85)),
  inference(subsumption_resolution,[],[f23134,f36])).
fof(f23134,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl7 | $spl26 | $spl85)),
  inference(forward_demodulation,[],[f23133,f15831])).
fof(f23133,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | ($spl7 | $spl26 | $spl85)),
  inference(subsumption_resolution,[],[f23018,f34])).
fof(f23018,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | ($spl7 | $spl26 | $spl85)),
  inference(superposition,[],[f592,f22949])).
fof(f22949,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl26 | $spl85)),
  inference(subsumption_resolution,[],[f22948,f11786])).
fof(f22948,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl26)),
  inference(subsumption_resolution,[],[f22833,f34])).
fof(f22833,plain,(
  ~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl26)),
  inference(superposition,[],[f377,f22053])).
fof(f22053,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl26)),
  inference(global_subsumption,[],[f22052,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f5026,f6038,f7476,f90,f33,f9508,f9526,f18290,f18291,f18728,f20717,f16124,f20953,f21003])).
fof(f22052,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,g(s(bool,X8))) | s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) ) | ($spl7 | $spl26)),
  inference(subsumption_resolution,[],[f9690,f3837])).
fof(f3837,plain,(
  ( ! [X70] : (s(bool,X70) = s(bool,g(s(bool,X70))) | s(bool,'F_0') = s(bool,X70)) ) | $spl26),
  inference(subsumption_resolution,[],[f2704,f34])).
fof(f2704,plain,(
  ( ! [X70] : (~p(s(bool,'T_0')) | s(bool,X70) = s(bool,g(s(bool,X70))) | s(bool,'F_0') = s(bool,X70)) ) | $spl26),
  inference(superposition,[],[f36,f631])).
fof(f631,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(superposition,[],[f538,f51])).
fof(f9690,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | s(bool,X8) = s(bool,g(s(bool,X8))) | s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) ) | ($spl7 | $spl26)),
  inference(superposition,[],[f90,f631])).
fof(f22304,plain,(
  $false | ($spl3 | $spl10 | $spl20 | $spl26 | $spl48)),
  inference(subsumption_resolution,[],[f22303,f59])).
fof(f22303,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl20 | $spl26 | $spl48)),
  inference(forward_demodulation,[],[f516,f21005])).
fof(f21005,plain,(
  ( ! [X249] : (s(bool,'T_0') = s(bool,f(s(bool,X249)))) ) | ($spl10 | $spl26 | $spl48)),
  inference(global_subsumption,[],[f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1587,f1577,f1656,f1806,f1795,f1788,f1703,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1854,f1855,f1864,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2813,f2807,f2484,f2486,f2796,f2795,f2881,f2883,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3702,f3079,f2349,f3827,f3837,f2739,f3682,f3681,f4059,f4061,f4753,f4949,f4945,f4878,f5029,f5026,f5827,f5134,f5136,f5137,f5139,f5172,f5928,f5319,f5320,f5414,f5415,f5423,f5426,f5427,f5428,f5429,f5959,f5545,f5546,f6038,f6052,f6330,f7476,f7584,f7582,f7527,f7575,f7645,f33,f9508,f9517,f9520,f11587,f13795,f15496,f15533,f15011,f17006,f17008,f18291,f18302,f18514,f18610,f18728,f18735,f20717,f20719,f20721,f17106])).
fof(f17106,plain,(
  ( ! [X249] : (~p(s(bool,X249)) | s(bool,'T_0') = s(bool,f(s(bool,X249))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f36,f5827])).
fof(f20721,plain,(
  ( ! [X52,X53,X51] : (s('type.rat.rat',X51) = s('type.rat.rat',X52) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X51),s('type.rat.rat',X52))) | s(bool,X53) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X51),s('type.rat.rat',X52))) | s(bool,'T_0') = s(bool,f(s(bool,X53)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f20699,f34])).
fof(f20699,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | s('type.rat.rat',X51) = s('type.rat.rat',X52) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X51),s('type.rat.rat',X52))) | s(bool,X53) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X51),s('type.rat.rat',X52))) | s(bool,'T_0') = s(bool,f(s(bool,X53)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f377,f1217])).
fof(f20719,plain,(
  ( ! [X30,X31,X32] : (s('type.rat.rat',X30) = s('type.rat.rat',X31) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X30),s('type.rat.rat',X31)))) ) | $spl26),
  inference(subsumption_resolution,[],[f20692,f34])).
fof(f20692,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s('type.rat.rat',X30) = s('type.rat.rat',X31) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X30),s('type.rat.rat',X31)))) ) | $spl26),
  inference(superposition,[],[f377,f631])).
fof(f18735,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s(bool,'T_0') = s(bool,f(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f18734])).
fof(f18734,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X62),s('type.rat.rat',X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,f(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X62),s('type.rat.rat',X63)))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f18712,f53])).
fof(f18712,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X64)) | s(bool,X65) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X62),s('type.rat.rat',X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,f(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X62),s('type.rat.rat',X63)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f343,f1217])).
fof(f18610,plain,(
  ( ! [X94,X92,X93] : (s(bool,X93) != s(bool,f(s(bool,X92))) | s(bool,X93) = s(bool,g(s(bool,X93))) | s(bool,X92) = s(bool,X94) | s(bool,'T_0') = s(bool,f(s(bool,X94)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f18411,f2739])).
fof(f18411,plain,(
  ( ! [X94,X92,X93] : (s(bool,X93) != s(bool,f(s(bool,X92))) | s(bool,X93) = s(bool,g(s(bool,X93))) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,X94) | s(bool,'T_0') = s(bool,f(s(bool,X94)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f11587,f1217])).
fof(f18514,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) != s(bool,f(s(bool,X64))) | s(bool,X62) = s(bool,g(s(bool,X62))) | s(bool,X62) = s(bool,X64) | s(bool,X63) = s(bool,g(s(bool,X63)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f18331])).
fof(f18331,plain,(
  ( ! [X64,X62,X63] : (s(bool,X63) != s(bool,f(s(bool,X64))) | s(bool,X62) = s(bool,g(s(bool,X62))) | s(bool,X62) = s(bool,X64) | s(bool,X62) = s(bool,g(s(bool,X62))) | s(bool,X63) = s(bool,g(s(bool,X63)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f11587,f2647])).
fof(f18302,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f18301])).
fof(f18301,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,X66) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X63),s('type.integer.int',X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X63),s('type.integer.int',X64)))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f18273,f53])).
fof(f18273,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | s(bool,X66) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X63),s('type.integer.int',X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X63),s('type.integer.int',X64)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f342,f1217])).
fof(f17008,plain,(
  ( ! [X428,X430,X429,X427] : (s(bool,X430) != s(bool,f(s(X427,sK2(X428,X429,bool,X427)))) | s(fun(X427,bool),X428) = s(fun(X427,bool),X429) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X427,sK2(X428,X429,bool,X427)))))) | s(bool,'T_0') = s(bool,f(s(bool,X430)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f5827])).
fof(f17006,plain,(
  ( ! [X420,X422,X421,X423] : (s(X422,f(s(bool,X423))) != s(X422,g(s(bool,X423))) | s(fun(bool,X422),X420) = s(fun(bool,X422),X421) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X420,X421,X422,bool)))) | s(bool,'T_0') = s(bool,f(s(bool,X423)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f5827])).
fof(f15011,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | $spl26),
  inference(equality_factoring,[],[f2647])).
fof(f13795,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1)))) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f1795])).
fof(f11587,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f1090])).
fof(f9520,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,f(s(bool,X12)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f33,f1217])).
fof(f9517,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,X9) = s(bool,g(s(bool,X9))) | s(bool,X9) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | $spl26),
  inference(superposition,[],[f33,f631])).
fof(f7645,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f37])).
fof(f7575,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f7574])).
fof(f7574,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f7546])).
fof(f7546,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f4949,f605])).
fof(f7527,plain,(
  ( ! [X81,X82] : (s(bool,'F_0') != s(bool,f(s(bool,X81))) | s(bool,X82) = s(bool,f(s(bool,X81))) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') = s(bool,X81)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f4949,f1264])).
fof(f7584,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f7583])).
fof(f7583,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,f(s(bool,X55)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f7516])).
fof(f7516,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,f(s(bool,X55)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f4949,f1217])).
fof(f6330,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,g(s(bool,X60))) | s(bool,X60) = s(bool,f(s(bool,X59))) | s(bool,'T_0') = s(bool,f(s(bool,X59)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f6237])).
fof(f6237,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,g(s(bool,X60))) | s(bool,X60) = s(bool,f(s(bool,X59))) | s(bool,'T_0') = s(bool,f(s(bool,X59))) | s(bool,'T_0') = s(bool,f(s(bool,X59)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f2720,f1137])).
fof(f6052,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X0,sK2(X1,X2,bool,X0))))))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f6051])).
fof(f6051,plain,(
  ( ! [X61,X59,X62,X60] : (s(fun(X59,bool),X60) = s(fun(X59,bool),X61) | s(bool,'T_0') = s(bool,g(s(X59,sK2(X60,X61,bool,X59)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X59,sK2(X60,X61,bool,X59))))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f6035,f37])).
fof(f6035,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') != s(bool,X62) | s(fun(X59,bool),X60) = s(fun(X59,bool),X61) | s(bool,'T_0') = s(bool,g(s(X59,sK2(X60,X61,bool,X59)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X59,sK2(X60,X61,bool,X59))))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f471,f1217])).
fof(f5546,plain,(
  ( ! [X198,X196,X199,X197] : (s(bool,'T_0') != s(bool,f(s(X196,sK2(X197,X198,bool,X196)))) | s(fun(X196,bool),X197) = s(fun(X196,bool),X198) | s(bool,X199) = s(bool,g(s(X196,sK2(X197,X198,bool,X196)))) | s(bool,'T_0') = s(bool,f(s(bool,X199)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5545,plain,(
  ( ! [X194,X192,X195,X193] : (s(X194,f(s(bool,'T_0'))) != s(X194,g(s(bool,'T_0'))) | s(fun(bool,X194),X192) = s(fun(bool,X194),X193) | s(bool,X195) = s(bool,sK2(X192,X193,X194,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X195)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5959,plain,(
  ( ! [X167,X165,X166] : (s('type.integer.int',X165) = s('type.integer.int',X166) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X166))) | s(bool,X167) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X165),s('type.integer.int',X166))) | s(bool,'T_0') = s(bool,f(s(bool,X167)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f5538,f34])).
fof(f5538,plain,(
  ( ! [X167,X165,X166] : (~p(s(bool,'T_0')) | s('type.integer.int',X165) = s('type.integer.int',X166) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X165),s('type.integer.int',X166))) | s(bool,X167) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X165),s('type.integer.int',X166))) | s(bool,'T_0') = s(bool,f(s(bool,X167)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f108,f1217])).
fof(f5429,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) != s(bool,X42) | s(bool,X42) = s(bool,g(s(bool,X42))) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f2739,f1217])).
fof(f5428,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) != s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,f(s(bool,X40)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f2734,f1217])).
fof(f5427,plain,(
  ( ! [X37,X35,X36] : (s(bool,X36) != s(bool,f(s(bool,X35))) | s(bool,'T_0') = s(bool,f(s(bool,X36))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,f(s(bool,X37)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1769,f1217])).
fof(f5426,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,f(s(bool,X33))) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,f(s(bool,X34)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f1217])).
fof(f5423,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,f(s(bool,X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,f(s(bool,X25)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f1217])).
fof(f5415,plain,(
  ( ! [X218,X220,X219,X217] : (s(bool,X220) != s(bool,f(s(X217,sK2(X218,X219,bool,X217)))) | s(fun(X217,bool),X218) = s(fun(X217,bool),X219) | s(bool,'T_0') = s(bool,g(s(X217,sK2(X218,X219,bool,X217)))) | s(bool,'T_0') = s(bool,f(s(bool,X220)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5414,plain,(
  ( ! [X216,X215,X213,X214] : (s(X215,f(s(bool,X216))) != s(X215,g(s(bool,X216))) | s(fun(bool,X215),X213) = s(fun(bool,X215),X214) | s(bool,'T_0') = s(bool,sK2(X213,X214,X215,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X216)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5320,plain,(
  ( ! [X218,X220,X219,X217] : (s(bool,X220) != s(bool,f(s(X217,sK2(X218,X219,bool,X217)))) | s(fun(X217,bool),X218) = s(fun(X217,bool),X219) | s(bool,'T_0') = s(bool,X220) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X217,sK2(X218,X219,bool,X217))))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5319,plain,(
  ( ! [X216,X215,X213,X214] : (s(X215,f(s(bool,X216))) != s(X215,g(s(bool,X216))) | s(fun(bool,X215),X213) = s(fun(bool,X215),X214) | s(bool,'T_0') = s(bool,X216) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X213,X214,X215,bool))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f1217])).
fof(f5928,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f5927])).
fof(f5927,plain,(
  ( ! [X187,X188,X186] : (s('type.integer.int',X186) != s('type.integer.int',X187) | s(bool,'T_0') = s(bool,X188) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X186),s('type.integer.int',X187)))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f5311,f53])).
fof(f5311,plain,(
  ( ! [X187,X188,X186] : (p(s(bool,X188)) | s('type.integer.int',X186) != s('type.integer.int',X187) | s(bool,'T_0') = s(bool,X188) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X186),s('type.integer.int',X187)))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f43,f1217])).
fof(f5172,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X62))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X63))))) | s(bool,X64) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X62))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X63))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X63))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X62))))))) | s(bool,'T_0') = s(bool,f(s(bool,X64)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f40])).
fof(f5139,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f51])).
fof(f5137,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,X4) = s(bool,X7)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f631])).
fof(f5136,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(bool,X3)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f1217])).
fof(f5134,plain,(
  ( ! [X90,X88,X89] : (s(bool,X90) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X88))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X89))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X88))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X89))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X89))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X88))))))) | s(bool,'T_0') = s(bool,f(s(bool,X90)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f40])).
fof(f5827,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1)))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f5826])).
fof(f5826,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,f(s(bool,X38))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,f(s(bool,X40)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f5069])).
fof(f5069,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,f(s(bool,X38))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,f(s(bool,X40)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1217,f1217])).
fof(f5029,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') = s(bool,f(s(bool,sK2(X83,X84,X85,bool)))) | s(fun(bool,X85),X83) = s(fun(bool,X85),X84) | s(X85,f(s(bool,'F_0'))) != s(X85,g(s(bool,'F_0')))) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f5011])).
fof(f5011,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(bool,sK2(X83,X84,X85,bool)))) | s(fun(bool,X85),X83) = s(fun(bool,X85),X84) | s(X85,f(s(bool,'F_0'))) != s(X85,g(s(bool,'F_0')))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f465])).
fof(f4878,plain,(
  ( ! [X76,X74,X75] : (s(bool,f(s(bool,X76))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X74))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X75))))) | s(bool,X76) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X74))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X75))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X75))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X74))))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X74))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X75))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X75))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X74)))))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1117,f40])).
fof(f4945,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,f(s(bool,X28))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f4859])).
fof(f4859,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,f(s(bool,X28))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X26)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1117,f631])).
fof(f4949,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f4850])).
fof(f4850,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1117,f37])).
fof(f4753,plain,(
  ( ! [X30,X31,X32] : (s('type.integer.int',X30) = s('type.integer.int',X31) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X30),s('type.integer.int',X31))) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X31)))) ) | $spl26),
  inference(subsumption_resolution,[],[f4745,f34])).
fof(f4745,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s('type.integer.int',X30) = s('type.integer.int',X31) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X30),s('type.integer.int',X31))) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X31)))) ) | $spl26),
  inference(superposition,[],[f108,f631])).
fof(f4061,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,g(s(bool,X7))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl26),
  inference(superposition,[],[f2739,f51])).
fof(f4059,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X2) = s(bool,g(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | $spl26),
  inference(superposition,[],[f2739,f631])).
fof(f3681,plain,(
  ( ! [X43,X41,X42] : (s(bool,X43) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X41))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X42))))) | s(bool,X43) = s(bool,g(s(bool,X43))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X41))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X42))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X42))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X41)))))))) ) | $spl26),
  inference(superposition,[],[f631,f40])).
fof(f3682,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X44))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X45))))) | s(bool,X46) = s(bool,g(s(bool,X46))) | s(bool,X46) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X44))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X45))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X45))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X44)))))))) ) | $spl26),
  inference(superposition,[],[f631,f40])).
fof(f2739,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2719])).
fof(f2719,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(equality_factoring,[],[f631])).
fof(f3827,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2142,f1761])).
fof(f2142,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f37])).
fof(f2349,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f2348,f584])).
fof(f2348,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2291,f1761])).
fof(f2291,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1769,f37])).
fof(f3079,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(forward_demodulation,[],[f3074,f538])).
fof(f3074,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3007])).
fof(f3007,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f2734,f37])).
fof(f3702,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X47))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X48))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X47))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X48))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X48))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X47)))))))) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f3683,f584])).
fof(f3683,plain,(
  ( ! [X47,X48] : (s(bool,f(s(bool,'T_0'))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X47))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X48))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X47))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X48))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X48))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X47)))))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f40])).
fof(f3561,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X36))))) | s(bool,X36) = s(bool,g(s(bool,X36)))) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f3548])).
fof(f3548,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X36))))) | s(bool,X36) = s(bool,g(s(bool,X36)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f2740])).
fof(f3570,plain,(
  ( ! [X43] : (s(bool,X43) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X43)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3521])).
fof(f3521,plain,(
  ( ! [X43] : (s(bool,X43) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X43))) | s(bool,'T_0') = s(bool,g(s(bool,X43)))) ) | $spl26),
  inference(superposition,[],[f2734,f2740])).
fof(f3638,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,X39) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X39)))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3517,f1761])).
fof(f3517,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,X39) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X39))))) | s(bool,'T_0') = s(bool,g(s(bool,X39)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f2740])).
fof(f3636,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,g(s(bool,X32)))) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f3571,f584])).
fof(f3571,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X32)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f3510])).
fof(f3510,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X32))) | s(bool,'T_0') = s(bool,g(s(bool,X32)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f2740])).
fof(f3574,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,g(s(bool,X13)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3500])).
fof(f3500,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,g(s(bool,X13))) | s(bool,'T_0') = s(bool,g(s(bool,X13)))) ) | $spl26),
  inference(superposition,[],[f53,f2740])).
fof(f3598,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,g(s(bool,X22))) | s(bool,X23) = s(bool,g(s(bool,X23)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3449])).
fof(f3449,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,g(s(bool,X22))) | s(bool,X23) = s(bool,g(s(bool,X23))) | s(bool,'T_0') = s(bool,g(s(bool,X22)))) ) | $spl26),
  inference(superposition,[],[f2740,f631])).
fof(f3035,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,g(s(bool,X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl26),
  inference(superposition,[],[f2734,f51])).
fof(f3033,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | $spl26),
  inference(superposition,[],[f2734,f631])).
fof(f2944,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X29),s('type.rat.rat',X30)))) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X31) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X29),s('type.rat.rat',X30)))) ) | $spl26),
  inference(subsumption_resolution,[],[f2943,f34])).
fof(f2943,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X29),s('type.rat.rat',X30)))) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X31) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X29),s('type.rat.rat',X30)))) ) | $spl26),
  inference(superposition,[],[f45,f631])).
fof(f2734,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2726])).
fof(f2726,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(equality_factoring,[],[f631])).
fof(f2720,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(equality_factoring,[],[f631])).
fof(f2740,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(trivial_inequality_removal,[],[f2718])).
fof(f2718,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,X0) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(equality_factoring,[],[f631])).
fof(f2717,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,f(s(X106,sK2(X107,X108,bool,X106)))) | s(fun(X106,bool),X107) = s(fun(X106,bool),X108) | s(bool,X109) = s(bool,g(s(bool,X109))) | s(bool,X109) = s(bool,g(s(X106,sK2(X107,X108,bool,X106))))) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2716,plain,(
  ( ! [X103,X105,X102,X104] : (s(X104,f(s(bool,'T_0'))) != s(X104,g(s(bool,'T_0'))) | s(fun(bool,X104),X102) = s(fun(bool,X104),X103) | s(bool,X105) = s(bool,g(s(bool,X105))) | s(bool,X105) = s(bool,sK2(X102,X103,X104,bool))) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2932,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X96),s('type.integer.int',X97)))) | s(bool,X98) = s(bool,g(s(bool,X98))) | s(bool,X98) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X96),s('type.integer.int',X97)))) ) | $spl26),
  inference(subsumption_resolution,[],[f2715,f34])).
fof(f2715,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X96),s('type.integer.int',X97)))) | s(bool,X98) = s(bool,g(s(bool,X98))) | s(bool,X98) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X96),s('type.integer.int',X97)))) ) | $spl26),
  inference(superposition,[],[f42,f631])).
fof(f2654,plain,(
  ( ! [X33,X34,X32] : (s(bool,X33) != s(bool,f(s(bool,X32))) | s(bool,'T_0') = s(bool,f(s(bool,X33))) | s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,X32) = s(bool,X34)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1769,f631])).
fof(f2653,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,f(s(bool,X30))) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X29) = s(bool,X31)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f631])).
fof(f2650,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,g(s(bool,X21))) | s(bool,X19) = s(bool,X21)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f631])).
fof(f2647,plain,(
  ( ! [X12,X11] : (s(bool,X11) = s(bool,g(s(bool,X11))) | s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,X11) = s(bool,X12)) ) | $spl26),
  inference(superposition,[],[f538,f631])).
fof(f2644,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | $spl26),
  inference(superposition,[],[f34,f631])).
fof(f2643,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) != s(bool,f(s(X122,sK2(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,X125) = s(bool,g(s(bool,X125))) | s(bool,'T_0') = s(bool,g(s(X122,sK2(X123,X124,bool,X122))))) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2642,plain,(
  ( ! [X121,X118,X120,X119] : (s(X120,f(s(bool,X121))) != s(X120,g(s(bool,X121))) | s(fun(bool,X120),X118) = s(fun(bool,X120),X119) | s(bool,X121) = s(bool,g(s(bool,X121))) | s(bool,'T_0') = s(bool,sK2(X118,X119,X120,bool))) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2588,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,X125) != s(bool,f(s(X122,sK2(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,g(s(X122,sK2(X123,X124,bool,X122)))) = s(bool,g(s(bool,g(s(X122,sK2(X123,X124,bool,X122)))))) | s(bool,'T_0') = s(bool,X125)) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2587,plain,(
  ( ! [X121,X118,X120,X119] : (s(X120,f(s(bool,X121))) != s(X120,g(s(bool,X121))) | s(fun(bool,X120),X118) = s(fun(bool,X120),X119) | s(bool,sK2(X118,X119,X120,bool)) = s(bool,g(s(bool,sK2(X118,X119,X120,bool)))) | s(bool,'T_0') = s(bool,X121)) ) | $spl26),
  inference(superposition,[],[f49,f631])).
fof(f2881,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X45))))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2529,f51])).
fof(f2529,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X45))))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f631])).
fof(f2486,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | $spl26),
  inference(superposition,[],[f631,f51])).
fof(f2484,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2839,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,g(s(bool,X0))) = s(bool,g(s(bool,g(s(bool,X0)))))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2838])).
fof(f2838,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,g(s(bool,X0))) = s(bool,g(s(bool,g(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(condensation,[],[f2425])).
fof(f2425,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,g(s(bool,X24))) = s(bool,g(s(bool,g(s(bool,X24))))) | s(bool,'T_0') = s(bool,X25)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2845,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2844])).
fof(f2844,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl26),
  inference(condensation,[],[f2419])).
fof(f2419,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,g(s(bool,X6)))) ) | $spl26),
  inference(superposition,[],[f631,f51])).
fof(f2417,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,X1)))) ) | $spl26),
  inference(superposition,[],[f631,f37])).
fof(f2314,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1769,f51])).
fof(f2133,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f51])).
fof(f1855,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26 | $spl48)),
  inference(backward_demodulation,[],[f1854,f561])).
fof(f1854,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26 | $spl48)),
  inference(subsumption_resolution,[],[f1813,f1761])).
fof(f1833,plain,(
  ( ! [X24] : (s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24)))))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1781,f1760])).
fof(f1781,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,f(s(bool,X24))) | s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24)))))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1706])).
fof(f1706,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,f(s(bool,X24))) | s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f1137])).
fof(f1760,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f1137])).
fof(f1761,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f1137])).
fof(f1769,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X25)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1738])).
fof(f1738,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,f(s(bool,X25)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f1137])).
fof(f1772,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,f(s(bool,X13)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1731])).
fof(f1731,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,f(s(bool,X13)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f53,f1137])).
fof(f1795,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1689])).
fof(f1689,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,f(s(bool,X17)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f605])).
fof(f1577,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,X13) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f193,f1264])).
fof(f1587,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1561])).
fof(f1561,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,X2)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1264,f37])).
fof(f1395,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f51])).
fof(f1394,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f605])).
fof(f1264,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,f(s(bool,X26))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1107,f34])).
fof(f1107,plain,(
  ( ! [X26] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,f(s(bool,X26))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f36,f605])).
fof(f561,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,X0))) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f508,f51])).
fof(f1087,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f34,f605])).
fof(f1213,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1212])).
fof(f1212,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f965])).
fof(f965,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,f(s(bool,X17))) = s(bool,f(s(bool,f(s(bool,X17))))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f1060,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X60) != s(bool,f(s(X57,sK2(X58,X59,bool,X57)))) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,g(s(X57,sK2(X58,X59,bool,X57)))) = s(bool,f(s(bool,g(s(X57,sK2(X58,X59,bool,X57)))))) | s(bool,'T_0') = s(bool,X60)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1059,plain,(
  ( ! [X54,X56,X55,X53] : (s(X55,f(s(bool,X56))) != s(X55,g(s(bool,X56))) | s(fun(bool,X55),X53) = s(fun(bool,X55),X54) | s(bool,sK2(X53,X54,X55,bool)) = s(bool,f(s(bool,sK2(X53,X54,X55,bool)))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1114,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') != s(bool,f(s(X45,sK2(X46,X47,bool,X45)))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,X48) = s(bool,f(s(bool,X48))) | s(bool,X48) = s(bool,g(s(X45,sK2(X46,X47,bool,X45))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1086,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X60) != s(bool,f(s(X57,sK2(X58,X59,bool,X57)))) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,X60) = s(bool,f(s(bool,X60))) | s(bool,'T_0') = s(bool,g(s(X57,sK2(X58,X59,bool,X57))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1131,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1123])).
fof(f1123,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f605])).
fof(f1117,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f605])).
fof(f1137,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(trivial_inequality_removal,[],[f1115])).
fof(f1115,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(equality_factoring,[],[f605])).
fof(f1113,plain,(
  ( ! [X43,X41,X44,X42] : (s(X43,f(s(bool,'T_0'))) != s(X43,g(s(bool,'T_0'))) | s(fun(bool,X43),X41) = s(fun(bool,X43),X42) | s(bool,X44) = s(bool,f(s(bool,X44))) | s(bool,X44) = s(bool,sK2(X41,X42,X43,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1091,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,X14) = s(bool,X15)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f584,f605])).
fof(f1090,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,X12) = s(bool,X13)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f538,f605])).
fof(f1085,plain,(
  ( ! [X54,X56,X55,X53] : (s(X55,f(s(bool,X56))) != s(X55,g(s(bool,X56))) | s(fun(bool,X55),X53) = s(fun(bool,X55),X54) | s(bool,X56) = s(bool,f(s(bool,X56))) | s(bool,'T_0') = s(bool,sK2(X53,X54,X55,bool))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f49,f605])).
fof(f1009,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f51])).
fof(f1008,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f1217,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1216])).
fof(f1216,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(condensation,[],[f961])).
fof(f961,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,f(s(bool,X6)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f51])).
fof(f959,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,X1)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f37])).
fof(f605,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f584,f51])).
fof(f584,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl10 | $spl26)),
  inference(backward_demodulation,[],[f538,f508])).
fof(f21153,plain,(
  ( ! [X103,X101] : (s(bool,'F_0') != s(bool,X101) | s(bool,X101) = s(bool,X103) | s(bool,X103) = s(bool,f(s(bool,X103)))) ) | $spl96),
  inference(cnf_transformation,[],[f21153_D])).
fof(f21153_D,plain,(
  ( ! [X103,X101] : (s(bool,'F_0') != s(bool,X101) | s(bool,X101) = s(bool,X103) | s(bool,X103) = s(bool,f(s(bool,X103)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f21918,plain,(
  ( ! [X152,X151,X153] : (s(bool,'F_0') != s(bool,X151) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,f(s(bool,X153))) | s(bool,X151) = s(bool,X153)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f11218,f21550])).
fof(f21550,plain,(
  ( ! [X144,X142,X143] : (s(bool,X143) != s(bool,g(s(bool,X142))) | s(bool,'T_0') = s(bool,X143) | s(bool,X144) = s(bool,f(s(bool,X144))) | s(bool,X142) = s(bool,X144)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f12234,f1394])).
fof(f12234,plain,(
  ( ! [X144,X142,X143] : (s(bool,X143) != s(bool,g(s(bool,X142))) | s(bool,X143) = s(bool,f(s(bool,X142))) | s(bool,'T_0') = s(bool,X143) | s(bool,X144) = s(bool,f(s(bool,X144))) | s(bool,X142) = s(bool,X144)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f2720,f1091])).
fof(f11218,plain,(
  ( ! [X152,X151,X153] : (s(bool,'F_0') != s(bool,X151) | s(bool,X152) = s(bool,g(s(bool,X151))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,f(s(bool,X153))) | s(bool,X151) = s(bool,X153)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f1090])).
fof(f21571,plain,(
  ( ! [X185,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,'T_0') = s(bool,X184) | s(bool,X185) = s(bool,f(s(bool,X185))) | s(bool,X183) = s(bool,X185)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f12172,f1394])).
fof(f12172,plain,(
  ( ! [X185,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,X184) = s(bool,f(s(bool,X183))) | s(bool,'T_0') = s(bool,X184) | s(bool,X185) = s(bool,f(s(bool,X185))) | s(bool,X183) = s(bool,X185)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f1091])).
fof(f21544,plain,(
  ( ! [X185,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,'T_0') = s(bool,X184) | s(bool,X185) = s(bool,f(s(bool,X185))) | s(bool,X183) = s(bool,X185)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f12251,f1394])).
fof(f12251,plain,(
  ( ! [X185,X184,X183] : (s(bool,'F_0') != s(bool,X183) | s(bool,X184) = s(bool,f(s(bool,X183))) | s(bool,'T_0') = s(bool,X184) | s(bool,X185) = s(bool,f(s(bool,X185))) | s(bool,X183) = s(bool,X185)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f1091])).
fof(f21154,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X98) | s(bool,X99) = s(bool,g(s(bool,X99))) | s(bool,X100) = s(bool,f(s(bool,X100))) | s(bool,X98) = s(bool,X100)) ) | ($spl10 | $spl26 | $spl94)),
  inference(subsumption_resolution,[],[f16305,f18611])).
fof(f18611,plain,(
  ( ! [X99,X100,X98] : (s(bool,X99) != s(bool,f(s(bool,X98))) | s(bool,X99) = s(bool,g(s(bool,X99))) | s(bool,X100) = s(bool,f(s(bool,X100))) | s(bool,X98) = s(bool,X100)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f18413,f2739])).
fof(f18413,plain,(
  ( ! [X99,X100,X98] : (s(bool,X99) != s(bool,f(s(bool,X98))) | s(bool,X99) = s(bool,g(s(bool,X99))) | s(bool,'T_0') = s(bool,X99) | s(bool,X100) = s(bool,f(s(bool,X100))) | s(bool,X98) = s(bool,X100)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f11587,f605])).
fof(f16305,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X98) | s(bool,X99) = s(bool,f(s(bool,X98))) | s(bool,X99) = s(bool,g(s(bool,X99))) | s(bool,X100) = s(bool,f(s(bool,X100))) | s(bool,X98) = s(bool,X100)) ) | ($spl10 | $spl26 | $spl94)),
  inference(superposition,[],[f15497,f1091])).
fof(f21152,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') != s(bool,X101) | s(bool,X102) = s(bool,g(s(bool,X102))) | s(bool,X103) = s(bool,f(s(bool,X103))) | s(bool,X101) = s(bool,X103)) ) | ($spl10 | $spl26 | $spl94)),
  inference(subsumption_resolution,[],[f16306,f18611])).
fof(f16306,plain,(
  ( ! [X103,X101,X102] : (s(bool,'F_0') != s(bool,X101) | s(bool,X102) = s(bool,f(s(bool,X101))) | s(bool,X102) = s(bool,g(s(bool,X102))) | s(bool,X103) = s(bool,f(s(bool,X103))) | s(bool,X101) = s(bool,X103)) ) | ($spl10 | $spl26 | $spl94)),
  inference(superposition,[],[f15497,f1091])).
fof(f20958,plain,(
  $false | ($spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(subsumption_resolution,[],[f20957,f36])).
fof(f20957,plain,(
  p(s(bool,'F_0')) | ($spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(backward_demodulation,[],[f20955,f20901])).
fof(f20901,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | ($spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(subsumption_resolution,[],[f20765,f34])).
fof(f20765,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | ($spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(superposition,[],[f592,f20716])).
fof(f20716,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(subsumption_resolution,[],[f20715,f11786])).
fof(f20715,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(subsumption_resolution,[],[f20681,f34])).
fof(f20681,plain,(
  ~p(s(bool,'T_0')) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f377,f9679])).
fof(f9679,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(unit_resulting_resolution,[],[f1815,f90,f4949])).
fof(f20955,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(global_subsumption,[],[f20954,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f5026,f6038,f7476,f90,f33,f9508,f9526,f18290,f18291,f18728,f20717,f16124,f20953])).
fof(f20954,plain,(
  ( ! [X123] : (s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X123)) ) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f16243,f1815])).
fof(f16243,plain,(
  ( ! [X123] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X123)) ) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(subsumption_resolution,[],[f16122,f8237])).
fof(f8237,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f8120,f584])).
fof(f8120,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26 | $spl50)),
  inference(condensation,[],[f8119])).
fof(f8119,plain,(
  ( ! [X50,X49] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X50)) ) | ($spl10 | $spl26 | $spl50)),
  inference(duplicate_literal_removal,[],[f7940])).
fof(f7940,plain,(
  ( ! [X50,X49] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,X50))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X49)) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f1131,f1971])).
fof(f1971,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) ) | $spl50),
  inference(superposition,[],[f1815,f51])).
fof(f16122,plain,(
  ( ! [X123] : (s(bool,'T_0') = s(bool,f(s(bool,X123))) | s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X123)) ) | ($spl7 | $spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f8904,f9526])).
fof(f8904,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,f(s(bool,'F_0')))) ) | ($spl10 | $spl26 | $spl50)),
  inference(condensation,[],[f8903])).
fof(f8903,plain,(
  ( ! [X24,X23] : (s(bool,X24) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | ($spl10 | $spl26 | $spl50)),
  inference(duplicate_literal_removal,[],[f8610])).
fof(f8610,plain,(
  ( ! [X24,X23] : (s(bool,X24) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f8177,f1217])).
fof(f8177,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(condensation,[],[f8176])).
fof(f8176,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl50),
  inference(duplicate_literal_removal,[],[f7848])).
fof(f7848,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(superposition,[],[f1971,f37])).
fof(f20903,plain,(
  $false | ($spl5 | $spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(subsumption_resolution,[],[f20902,f36])).
fof(f20902,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl7 | $spl10 | $spl26 | $spl50 | $spl85)),
  inference(forward_demodulation,[],[f20901,f9805])).
fof(f9805,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl5 | $spl50)),
  inference(forward_demodulation,[],[f9762,f1815])).
fof(f9762,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl5 | $spl50)),
  inference(unit_resulting_resolution,[],[f87,f8177])).
fof(f87,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl5),
  inference(cnf_transformation,[],[f87_D])).
fof(f87_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f15497,plain,(
  ( ! [X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,X188) = s(bool,X190) | s(bool,X190) = s(bool,g(s(bool,X190)))) ) | $spl94),
  inference(cnf_transformation,[],[f15497_D])).
fof(f15497_D,plain,(
  ( ! [X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,X188) = s(bool,X190) | s(bool,X190) = s(bool,g(s(bool,X190)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f15535,plain,(
  ( ! [X191,X192,X193] : (s(bool,'F_0') != s(bool,X191) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(subsumption_resolution,[],[f15534,f1131])).
fof(f15534,plain,(
  ( ! [X191,X192,X193] : (s(bool,X192) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,X191) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(forward_demodulation,[],[f14683,f3651])).
fof(f3651,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | ($spl10 | $spl26 | $spl76)),
  inference(subsumption_resolution,[],[f3616,f3636])).
fof(f14683,plain,(
  ( ! [X191,X192,X193] : (s(bool,'F_0') != s(bool,X191) | s(bool,X192) = s(bool,f(s(bool,g(s(bool,X191))))) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f7688,f2647])).
fof(f7688,plain,(
  ( ! [X74,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,X74) = s(bool,f(s(bool,X73))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f7575,f1815])).
fof(f15533,plain,(
  ( ! [X189,X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,'T_0') = s(bool,X189) | s(bool,X190) = s(bool,g(s(bool,X190))) | s(bool,X188) = s(bool,X190)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14682,f3033])).
fof(f14682,plain,(
  ( ! [X189,X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,X189) = s(bool,g(s(bool,X188))) | s(bool,'T_0') = s(bool,X189) | s(bool,X190) = s(bool,g(s(bool,X190))) | s(bool,X188) = s(bool,X190)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f2647])).
fof(f15499,plain,(
  ( ! [X191,X192,X193] : (s(bool,'F_0') != s(bool,X191) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(subsumption_resolution,[],[f15498,f1131])).
fof(f15498,plain,(
  ( ! [X191,X192,X193] : (s(bool,X192) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,X191) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(forward_demodulation,[],[f14603,f3651])).
fof(f14603,plain,(
  ( ! [X191,X192,X193] : (s(bool,'F_0') != s(bool,X191) | s(bool,X192) = s(bool,f(s(bool,g(s(bool,X191))))) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(bool,X193))) | s(bool,X191) = s(bool,X193)) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f7688,f2647])).
fof(f15496,plain,(
  ( ! [X189,X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,'T_0') = s(bool,X189) | s(bool,X190) = s(bool,g(s(bool,X190))) | s(bool,X188) = s(bool,X190)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f14602,f3033])).
fof(f14602,plain,(
  ( ! [X189,X190,X188] : (s(bool,'F_0') != s(bool,X188) | s(bool,X189) = s(bool,g(s(bool,X188))) | s(bool,'T_0') = s(bool,X189) | s(bool,X190) = s(bool,g(s(bool,X190))) | s(bool,X188) = s(bool,X190)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f7575,f2647])).
fof(f12963,plain,(
  $false | ($spl5 | $spl50 | $spl82)),
  inference(subsumption_resolution,[],[f12962,f36])).
fof(f12962,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl50 | $spl82)),
  inference(forward_demodulation,[],[f12961,f9805])).
fof(f12961,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | $spl82),
  inference(subsumption_resolution,[],[f12861,f34])).
fof(f12861,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) | $spl82),
  inference(superposition,[],[f42,f4748])).
fof(f9179,plain,(
  $false | $spl84),
  inference(subsumption_resolution,[],[f9178,f73])).
fof(f9178,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))) | $spl84),
  inference(forward_demodulation,[],[f9177,f9109])).
fof(f9109,plain,(
  s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))) | $spl84),
  inference(unit_resulting_resolution,[],[f4750,f48])).
fof(f9177,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl84),
  inference(forward_demodulation,[],[f9176,f80])).
fof(f9176,plain,(
  s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))))) | $spl84),
  inference(forward_demodulation,[],[f33,f4750])).
fof(f9175,plain,(
  $false | ($spl7 | $spl50 | $spl84)),
  inference(global_subsumption,[],[f8840,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f90,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f33,f5026,f6038,f7476,f4750,f9109])).
fof(f8840,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,'F_0')))) ) | $spl50),
  inference(equality_factoring,[],[f8177])).
fof(f9174,plain,(
  $false | ($spl7 | $spl15 | $spl84)),
  inference(global_subsumption,[],[f512,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f90,f91,f95,f99,f141,f193,f461,f463,f519,f101,f36,f465,f1853,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f33,f5026,f6038,f7476,f4750,f9109])).
fof(f9148,plain,(
  $false | ($spl4 | $spl84)),
  inference(subsumption_resolution,[],[f9145,f73])).
fof(f9145,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))) | ($spl4 | $spl84)),
  inference(backward_demodulation,[],[f9109,f3833])).
fof(f3833,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl4),
  inference(forward_demodulation,[],[f33,f88])).
fof(f9147,plain,(
  $false | ($spl3 | $spl6 | $spl84)),
  inference(subsumption_resolution,[],[f9146,f59])).
fof(f9146,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl84)),
  inference(forward_demodulation,[],[f9144,f73])).
fof(f9144,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))))) | ($spl6 | $spl84)),
  inference(backward_demodulation,[],[f9109,f89])).
fof(f89,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl6),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f512,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f8945,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f8944,f538])).
fof(f8944,plain,(
  ( ! [X71] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl10 | $spl26 | $spl50)),
  inference(subsumption_resolution,[],[f8731,f8237])).
fof(f8731,plain,(
  ( ! [X71] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X71))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f3022,f8177])).
fof(f3022,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X26)))) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f2734,f1815])).
fof(f8941,plain,(
  ( ! [X59] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl26 | $spl50)),
  inference(forward_demodulation,[],[f8863,f538])).
fof(f8863,plain,(
  ( ! [X59] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl26 | $spl50)),
  inference(duplicate_literal_removal,[],[f8723])).
fof(f8723,plain,(
  ( ! [X59] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f2734,f8177])).
fof(f8940,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,'F_0')))) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(condensation,[],[f8939])).
fof(f8939,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X57)) ) | ($spl10 | $spl26 | $spl50 | $spl76)),
  inference(forward_demodulation,[],[f8864,f3651])).
fof(f8864,plain,(
  ( ! [X57,X58] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,X58))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X57)) ) | ($spl26 | $spl50)),
  inference(duplicate_literal_removal,[],[f8722])).
fof(f8722,plain,(
  ( ! [X57,X58] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,X58))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f2720,f8177])).
fof(f8933,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl10 | $spl26 | $spl50)),
  inference(subsumption_resolution,[],[f8715,f8237])).
fof(f8715,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X47))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl10 | $spl26 | $spl50)),
  inference(superposition,[],[f1761,f8177])).
fof(f7787,plain,(
  $false | ($spl6 | $spl92)),
  inference(subsumption_resolution,[],[f7786,f36])).
fof(f7786,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl92)),
  inference(forward_demodulation,[],[f7766,f89])).
fof(f7766,plain,(
  p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | $spl92),
  inference(unit_resulting_resolution,[],[f7755,f45])).
fof(f7755,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | $spl92),
  inference(cnf_transformation,[],[f7755_D])).
fof(f7755_D,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4750,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | $spl84),
  inference(cnf_transformation,[],[f4750_D])).
fof(f4750_D,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7753,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | $spl4),
  inference(subsumption_resolution,[],[f7734,f34])).
fof(f7734,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1)))))) | $spl4),
  inference(superposition,[],[f591,f88])).
fof(f7604,plain,(
  ( ! [X58,X56] : (s(bool,'F_0') != s(bool,f(s(bool,X56))) | s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,g(s(bool,X58)))) ) | $spl90),
  inference(cnf_transformation,[],[f7604_D])).
fof(f7604_D,plain,(
  ( ! [X58,X56] : (s(bool,'F_0') != s(bool,f(s(bool,X56))) | s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,g(s(bool,X58)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7582,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') != s(bool,f(s(bool,X56))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,g(s(bool,X58))) | s(bool,X56) = s(bool,X58)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f7517])).
fof(f7517,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') != s(bool,f(s(bool,X56))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,g(s(bool,X58))) | s(bool,X56) = s(bool,X58)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f4949,f631])).
fof(f6040,plain,(
  ( ! [X43,X44] : (s(fun(fun(bool,bool),bool),X43) = s(fun(fun(bool,bool),bool),X44) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X43,X44,bool,fun(bool,bool)))))) ) | $spl86),
  inference(cnf_transformation,[],[f6040_D])).
fof(f6040_D,plain,(
  ( ! [X43,X44] : (s(fun(fun(bool,bool),bool),X43) = s(fun(fun(bool,bool),bool),X44) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X43,X44,bool,fun(bool,bool)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6041,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X45)))) ) | $spl88),
  inference(cnf_transformation,[],[f6041_D])).
fof(f6041_D,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X45)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6010,plain,(
  ( ! [X47,X48,X46] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X48))) | s(fun(fun(bool,bool),bool),X46) = s(fun(fun(bool,bool),bool),X47) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X46,X47,bool,fun(bool,bool)))))) ) | $spl12),
  inference(superposition,[],[f471,f510])).
fof(f6009,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X45))) | s(fun(fun(bool,bool),bool),X43) = s(fun(fun(bool,bool),bool),X44) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK2(X43,X44,bool,fun(bool,bool)))))) ) | $spl12),
  inference(superposition,[],[f471,f510])).
fof(f4748,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl82),
  inference(cnf_transformation,[],[f4748_D])).
fof(f4748_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4746,plain,(
  s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl4),
  inference(subsumption_resolution,[],[f4734,f34])).
fof(f4734,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl4),
  inference(superposition,[],[f108,f88])).
fof(f3832,plain,(
  $false | ($spl3 | $spl4 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f88,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079,f2349,f3827,f59,f33])).
fof(f3818,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(condensation,[],[f3817])).
fof(f3817,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X39)) ) | ($spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(forward_demodulation,[],[f3816,f3651])).
fof(f3816,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,X39) | s(bool,'F_0') = s(bool,g(s(bool,X39)))) ) | ($spl10 | $spl26 | $spl50 | $spl80)),
  inference(forward_demodulation,[],[f3787,f1815])).
fof(f3787,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,f(s(bool,X39))) | s(bool,'F_0') = s(bool,g(s(bool,X39)))) ) | ($spl10 | $spl26 | $spl80)),
  inference(superposition,[],[f1587,f3620])).
fof(f3754,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,g(s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl80),
  inference(superposition,[],[f3620,f51])).
fof(f3703,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X57))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X58))))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X57))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X58))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X58))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X57)))))))) ) | ($spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f3688,f1815])).
fof(f3688,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,f(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X57))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X58))))))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X57))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X58))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X58))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X57)))))))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1587,f40])).
fof(f2506,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X31) = s(bool,f(s(bool,X30)))) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f631,f1815])).
fof(f2473,plain,(
  ( ! [X43,X44] : (s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,g(s(bool,X44))) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | ($spl26 | $spl50)),
  inference(superposition,[],[f631,f1815])).
fof(f1826,plain,(
  ( ! [X24] : (s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24)))))) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1781,f1260])).
fof(f1591,plain,(
  ( ! [X25] : (s(bool,f(s(bool,X25))) = s(bool,f(s(bool,f(s(bool,X25))))) | s(bool,'F_0') = s(bool,X25)) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1585,f37])).
fof(f1585,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,f(s(bool,X25))) = s(bool,f(s(bool,f(s(bool,X25))))) | s(bool,'F_0') = s(bool,X25)) ) | ($spl3 | $spl10 | $spl26)),
  inference(superposition,[],[f1260,f1264])).
fof(f1584,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,f(s(bool,X24))) = s(bool,f(s(bool,f(s(bool,X24))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl3 | $spl10 | $spl26)),
  inference(superposition,[],[f1259,f1264])).
fof(f1473,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | ($spl3 | $spl10 | $spl26)),
  inference(superposition,[],[f1260,f51])).
fof(f1472,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl3 | $spl10 | $spl26)),
  inference(superposition,[],[f1260,f605])).
fof(f1259,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | s(bool,X32) = s(bool,f(s(bool,X32)))) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1074,f59])).
fof(f1074,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | s(bool,X32) = s(bool,f(s(bool,X32))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f36,f605])).
fof(f1260,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,f(s(bool,X33)))) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1075,f59])).
fof(f1075,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl3 | $spl10 | $spl26)),
  inference(superposition,[],[f59,f605])).
fof(f3831,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079,f2349,f3827,f59,f33])).
fof(f3830,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f33,f34,f37,f48,f47,f49,f51,f53,f56,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079,f2349,f3827,f59])).
fof(f3828,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f33,f34,f37,f48,f47,f49,f51,f53,f56,f59,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079,f2349,f3827])).
fof(f3826,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f33,f34,f37,f48,f47,f49,f51,f53,f56,f59,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079,f2349])).
fof(f3825,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f33,f34,f37,f48,f47,f49,f51,f53,f56,f59,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818,f3079])).
fof(f3824,plain,(
  $false | ($spl3 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(global_subsumption,[],[f33,f34,f37,f48,f47,f49,f51,f53,f56,f59,f65,f71,f73,f80,f86,f141,f193,f461,f463,f508,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f538,f584,f605,f631,f1222,f959,f1217,f1206,f1192,f1186,f1008,f1009,f1177,f1176,f1256,f1153,f1148,f1260,f1085,f1090,f1091,f1113,f1137,f1117,f1131,f1086,f1114,f1059,f1060,f1259,f1213,f1087,f561,f1264,f36,f1421,f1394,f1395,f1472,f1473,f1587,f1577,f1584,f1591,f1656,f1814,f1816,f1806,f1818,f1795,f1821,f1822,f1823,f1788,f1703,f1826,f1772,f1769,f1761,f1760,f1763,f1776,f1833,f1835,f1838,f1839,f1840,f1842,f1844,f1431,f465,f1815,f593,f588,f586,f40,f591,f592,f664,f720,f459,f342,f108,f43,f42,f41,f46,f45,f44,f343,f377,f467,f469,f471,f1971,f2133,f2314,f2851,f2417,f2845,f2839,f2831,f2869,f2813,f2807,f2473,f2484,f2486,f2796,f2795,f2506,f2881,f2883,f2904,f2915,f2587,f2588,f2766,f2761,f2918,f2921,f2923,f2642,f2643,f2644,f2647,f2650,f2653,f2654,f2932,f2716,f2717,f2740,f2720,f2734,f2944,f3068,f3022,f3033,f3035,f3113,f3314,f3406,f3609,f3598,f3591,f3466,f3574,f3636,f3638,f3570,f3561,f3620,f3651,f3702,f3703,f3754,f3818])).
fof(f3819,plain,(
  $false | ($spl7 | $spl10 | $spl26 | $spl50 | $spl76 | $spl80)),
  inference(subsumption_resolution,[],[f3818,f1429])).
fof(f1429,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1428,f584])).
fof(f1428,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | ($spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1373,f101])).
fof(f1373,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl10 | $spl26)),
  inference(unit_resulting_resolution,[],[f95,f1131])).
fof(f3616,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(cnf_transformation,[],[f3616_D])).
fof(f3616_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3466,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,g(s(bool,X29))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl26),
  inference(superposition,[],[f2740,f37])).
fof(f3591,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3459])).
fof(f3459,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) ) | $spl26),
  inference(superposition,[],[f2740,f631])).
fof(f3628,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3454,f2924])).
fof(f2924,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X77) = s(bool,g(s(bool,X77)))) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2622,f1429])).
fof(f2622,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X77) = s(bool,g(s(bool,X77))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f1429,f631])).
fof(f3454,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | $spl26),
  inference(superposition,[],[f2740,f51])).
fof(f3627,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3453,f2924])).
fof(f3453,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f2740,f51])).
fof(f514,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl16),
  inference(cnf_transformation,[],[f514_D])).
fof(f514_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3626,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3451,f2924])).
fof(f3451,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f2740,f37])).
fof(f3623,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3421,f2924])).
fof(f3421,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl26),
  inference(superposition,[],[f2740,f37])).
fof(f3609,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X17)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3415])).
fof(f3415,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,g(s(bool,X17))) | s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | $spl26),
  inference(superposition,[],[f2740,f631])).
fof(f3621,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,g(s(bool,X6))) | s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3409,f2924])).
fof(f3409,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,g(s(bool,X6))) | s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | $spl26),
  inference(superposition,[],[f2740,f51])).
fof(f3619,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,g(s(bool,X4))) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f3408,f2924])).
fof(f3408,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,g(s(bool,X4))) | s(bool,'T_0') = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl26),
  inference(superposition,[],[f2740,f51])).
fof(f3406,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f2740,f37])).
fof(f3315,plain,(
  ( ! [X24,X23] : (s('type.rat.rat',X23) = s('type.rat.rat',X24) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) ) | $spl74),
  inference(cnf_transformation,[],[f3315_D])).
fof(f3315_D,plain,(
  ( ! [X24,X23] : (s('type.rat.rat',X23) = s('type.rat.rat',X24) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3314,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',X23) = s('type.rat.rat',X24) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X23),s('type.rat.rat',X24)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) | s(bool,'T_0') = s(bool,X25)) ) | $spl26),
  inference(subsumption_resolution,[],[f3307,f53])).
fof(f3307,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X25)) | s('type.rat.rat',X23) = s('type.rat.rat',X24) | ~p(s(bool,'const.rat.rat_leq_2'(s('type.rat.rat',X23),s('type.rat.rat',X24)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))) | s(bool,'T_0') = s(bool,X25)) ) | $spl26),
  inference(superposition,[],[f44,f631])).
fof(f3114,plain,(
  ( ! [X24,X23] : (s('type.integer.int',X23) = s('type.integer.int',X24) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))) = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) ) | $spl72),
  inference(cnf_transformation,[],[f3114_D])).
fof(f3114_D,plain,(
  ( ! [X24,X23] : (s('type.integer.int',X23) = s('type.integer.int',X24) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))) = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3113,plain,(
  ( ! [X24,X23,X25] : (s('type.integer.int',X23) = s('type.integer.int',X24) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int',X24)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))) = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) | s(bool,'T_0') = s(bool,X25)) ) | $spl26),
  inference(subsumption_resolution,[],[f3106,f53])).
fof(f3106,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X25)) | s('type.integer.int',X23) = s('type.integer.int',X24) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X23),s('type.integer.int',X24)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))) = s(bool,g(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X23),s('type.integer.int',X24))))) | s(bool,'T_0') = s(bool,X25)) ) | $spl26),
  inference(superposition,[],[f41,f631])).
fof(f3080,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,g(s(bool,'T_0'))) | s(bool,X18) = s(bool,g(s(bool,X18)))) ) | $spl70),
  inference(cnf_transformation,[],[f3080_D])).
fof(f3080_D,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,g(s(bool,'T_0'))) | s(bool,X18) = s(bool,g(s(bool,X18)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3068,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,X18)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f3016])).
fof(f3016,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | $spl26),
  inference(superposition,[],[f2734,f631])).
fof(f1836,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,f(s(bool,X6)))) ) | $spl54),
  inference(cnf_transformation,[],[f1836_D])).
fof(f1836_D,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,f(s(bool,X6)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2923,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') != s(bool,X48) | s(bool,'T_0') = s(bool,f(s(bool,X47))) | s(bool,X48) = s(bool,g(s(bool,X48)))) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f2922,f584])).
fof(f2922,plain,(
  ( ! [X47,X48] : (s(bool,X48) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X47))) | s(bool,X48) = s(bool,g(s(bool,X48)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2608,f1761])).
fof(f2608,plain,(
  ( ! [X47,X48] : (s(bool,X48) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X47))) | s(bool,X48) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1769,f631])).
fof(f2921,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,'T_0') = s(bool,f(s(bool,X45))) | s(bool,X46) = s(bool,g(s(bool,X46)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2607,f1761])).
fof(f2607,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,'T_0') = s(bool,f(s(bool,X45))) | s(bool,X46) = s(bool,g(s(bool,X46))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1761,f631])).
fof(f2918,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,g(s(bool,X32)))) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f2759,f584])).
fof(f2759,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,g(s(bool,X32)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f2600])).
fof(f2600,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f631])).
fof(f2917,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl66),
  inference(cnf_transformation,[],[f2917_D])).
fof(f2917_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2761,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2594])).
fof(f2594,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | $spl26),
  inference(superposition,[],[f53,f631])).
fof(f2870,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,g(s(bool,X53))) | s(bool,'F_0') = s(bool,X53)) ) | $spl62),
  inference(cnf_transformation,[],[f2870_D])).
fof(f2870_D,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,g(s(bool,X53))) | s(bool,'F_0') = s(bool,X53)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2766,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2589])).
fof(f2589,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f37,f631])).
fof(f2916,plain,(
  ( ! [X107,X108] : (s('type.integer.int',X107) != s('type.integer.int',X108) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108)))))) ) | $spl64),
  inference(cnf_transformation,[],[f2916_D])).
fof(f2916_D,plain,(
  ( ! [X107,X108] : (s('type.integer.int',X107) != s('type.integer.int',X108) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2915,plain,(
  ( ! [X109,X107,X108] : (s('type.integer.int',X107) != s('type.integer.int',X108) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))))) | s(bool,'T_0') = s(bool,X109)) ) | $spl26),
  inference(subsumption_resolution,[],[f2584,f53])).
fof(f2584,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X109)) | s('type.integer.int',X107) != s('type.integer.int',X108) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))) = s(bool,g(s(bool,'const.integer.int_le_2'(s('type.integer.int',X107),s('type.integer.int',X108))))) | s(bool,'T_0') = s(bool,X109)) ) | $spl26),
  inference(superposition,[],[f43,f631])).
fof(f2904,plain,(
  ( ! [X94,X95] : (s(bool,X94) = s(bool,g(s(bool,X94))) | s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X94)) ) | ($spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f2773,f1815])).
fof(f2773,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X94) | s(bool,f(s(bool,X94))) = s(bool,g(s(bool,f(s(bool,X94)))))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f2577])).
fof(f2577,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X94) | s(bool,f(s(bool,X94))) = s(bool,g(s(bool,f(s(bool,X94))))) | s(bool,'T_0') = s(bool,X95)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1587,f631])).
fof(f2884,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,g(s(bool,X51))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl26 | $spl56)),
  inference(subsumption_resolution,[],[f2555,f53])).
fof(f2555,plain,(
  ( ! [X52,X51] : (p(s(bool,X52)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,g(s(bool,X51))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl26 | $spl56)),
  inference(superposition,[],[f1867,f631])).
fof(f2883,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,f(s(bool,X42))) | s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,g(s(bool,X41)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f2550,f1761])).
fof(f2550,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,f(s(bool,X42))) | s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,g(s(bool,X41))) | s(bool,'T_0') = s(bool,X42)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1587,f631])).
fof(f2795,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'F_0') = s(bool,X29)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2505])).
fof(f2505,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'F_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X28)) ) | $spl26),
  inference(superposition,[],[f631,f37])).
fof(f2859,plain,(
  ( ! [X34,X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,X34) | s(bool,X34) = s(bool,g(s(bool,X34)))) ) | $spl60),
  inference(cnf_transformation,[],[f2859_D])).
fof(f2859_D,plain,(
  ( ! [X34,X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,X32) = s(bool,X34) | s(bool,X34) = s(bool,g(s(bool,X34)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2796,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X19))) | s(bool,X17) = s(bool,X19) | s(bool,X17) = s(bool,g(s(bool,X17)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2503])).
fof(f2503,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X19))) | s(bool,X17) = s(bool,X19) | s(bool,X17) = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X18)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2807,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,X28) | s(bool,X28) = s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2465])).
fof(f2465,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,X28) | s(bool,X28) = s(bool,g(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X26)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2813,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2456])).
fof(f2456,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(superposition,[],[f631,f37])).
fof(f2869,plain,(
  ( ! [X54,X53] : (s(bool,X53) = s(bool,g(s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl10 | $spl26 | $spl50)),
  inference(forward_demodulation,[],[f2822,f1815])).
fof(f2822,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,f(s(bool,X53))) = s(bool,g(s(bool,f(s(bool,X53))))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f2448])).
fof(f2448,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,f(s(bool,X53))) = s(bool,g(s(bool,f(s(bool,X53))))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f631,f1587])).
fof(f2831,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X29) = s(bool,X31)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2438])).
fof(f2438,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,X31))) | s(bool,X29) = s(bool,X31)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2851,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,X32) = s(bool,X34)) ) | $spl26),
  inference(duplicate_literal_removal,[],[f2400])).
fof(f2400,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,g(s(bool,X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,g(s(bool,X34))) | s(bool,X32) = s(bool,X34)) ) | $spl26),
  inference(superposition,[],[f631,f631])).
fof(f2010,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl58),
  inference(cnf_transformation,[],[f2010_D])).
fof(f2010_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2009,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X17) | p(s(bool,X16))) ) | ($spl10 | $spl26 | $spl48 | $spl50)),
  inference(duplicate_literal_removal,[],[f2008])).
fof(f2008,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X17) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17)) ) | ($spl10 | $spl26 | $spl48 | $spl50)),
  inference(forward_demodulation,[],[f1987,f1854])).
fof(f1987,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) ) | $spl50),
  inference(superposition,[],[f193,f1815])).
fof(f1909,plain,(
  $false | $spl0),
  inference(equality_resolution,[],[f1899])).
fof(f1899,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat',X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f1878,f36])).
fof(f1878,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.rat.rat',X0) != s('type.rat.rat',X1)) ) | $spl0),
  inference(backward_demodulation,[],[f57,f46])).
fof(f1868,plain,(
  $false | ($spl2 | $spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f58,f1429])).
fof(f58,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl2),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1867,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) | $spl56),
  inference(cnf_transformation,[],[f1867_D])).
fof(f1867_D,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1864,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) | ($spl10 | $spl26 | $spl48)),
  inference(duplicate_literal_removal,[],[f1858])).
fof(f1858,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | ($spl10 | $spl26 | $spl48)),
  inference(backward_demodulation,[],[f1854,f1577])).
fof(f59,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f59_D])).
fof(f59_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1431,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(forward_demodulation,[],[f1427,f584])).
fof(f1427,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1375])).
fof(f1375,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f37])).
fof(f1815,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,f(s(bool,X4)))) ) | $spl50),
  inference(cnf_transformation,[],[f1815_D])).
fof(f1815_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,f(s(bool,X4)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1844,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1843,f1761])).
fof(f1843,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1658,f1760])).
fof(f1658,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f51])).
fof(f1842,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1841,f1761])).
fof(f1841,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1659,f1760])).
fof(f1659,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f51])).
fof(f511,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl14),
  inference(cnf_transformation,[],[f511_D])).
fof(f511_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1840,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1668,f1760])).
fof(f1668,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f37])).
fof(f1839,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1691,f1760])).
fof(f1691,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f37])).
fof(f1838,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1837,f1760])).
fof(f1837,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1693,f1761])).
fof(f1693,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f51])).
fof(f1835,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5)))) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1834,f1760])).
fof(f1834,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1694,f1761])).
fof(f1694,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f51])).
fof(f1831,plain,(
  ( ! [X18] : (s(bool,f(s(bool,X18))) = s(bool,f(s(bool,f(s(bool,X18))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | $spl52),
  inference(cnf_transformation,[],[f1831_D])).
fof(f1831_D,plain,(
  ( ! [X18] : (s(bool,f(s(bool,X18))) = s(bool,f(s(bool,f(s(bool,X18))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1776,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,f(s(bool,X14))) = s(bool,f(s(bool,f(s(bool,X14)))))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1722])).
fof(f1722,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,f(s(bool,X14))) = s(bool,f(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f605])).
fof(f1763,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,f(s(bool,X18))) = s(bool,f(s(bool,f(s(bool,X18))))) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1753])).
fof(f1753,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,f(s(bool,X18))) = s(bool,f(s(bool,f(s(bool,X18))))) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,f(s(bool,X18)))) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f1137])).
fof(f1813,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl48),
  inference(cnf_transformation,[],[f1813_D])).
fof(f1813_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f516,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl20),
  inference(cnf_transformation,[],[f516_D])).
fof(f516_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1827,plain,(
  ( ! [X31] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1714,f101])).
fof(f1714,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,f(s(bool,X31))) | s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f1137,f333])).
fof(f333,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f151,f59])).
fof(f151,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X16)) ) | $spl7),
  inference(superposition,[],[f51,f101])).
fof(f1703,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,f(s(bool,X22))) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f37])).
fof(f1788,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X10) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1697])).
fof(f1697,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X10) = s(bool,f(s(bool,X10))) | s(bool,X10) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f605])).
fof(f1825,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1824,f101])).
fof(f1824,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,X9) = s(bool,f(s(bool,X9)))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1696,f1260])).
fof(f1696,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,X9) = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f1137,f333])).
fof(f1823,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1694,f1260])).
fof(f1822,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1693,f1260])).
fof(f1821,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1691,f1260])).
fof(f1820,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X31) = s(bool,f(s(bool,X31)))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1819,f101])).
fof(f1819,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1679,f1260])).
fof(f1679,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f1137,f333])).
fof(f1818,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1668,f1260])).
fof(f1806,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1663])).
fof(f1663,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,f(s(bool,X12))) | s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f605])).
fof(f1817,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(forward_demodulation,[],[f1661,f101])).
fof(f1661,plain,(
  ( ! [X9] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(bool,'T_0') = s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f1137,f333])).
fof(f1816,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1659,f1260])).
fof(f1814,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl3 | $spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1658,f1260])).
fof(f1656,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1137,f37])).
fof(f1434,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(bool,'T_0'))) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) | $spl46),
  inference(cnf_transformation,[],[f1434_D])).
fof(f1434_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(bool,'T_0'))) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1421,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,f(s(bool,X13)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1382])).
fof(f1382,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f1131,f605])).
fof(f1310,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f1265,f34])).
fof(f1265,plain,(
  ~p(s(bool,'T_0')) | $spl34),
  inference(backward_demodulation,[],[f1230,f36])).
fof(f1243,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl40),
  inference(cnf_transformation,[],[f1243_D])).
fof(f1243_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1139,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,X36) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1111])).
fof(f1111,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,f(s(bool,X36))) | s(bool,X36) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f333,f605])).
fof(f1147,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1068])).
fof(f1068,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f333,f605])).
fof(f1258,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl44),
  inference(cnf_transformation,[],[f1258_D])).
fof(f1258_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1148,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1066])).
fof(f1066,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f53,f605])).
fof(f1242,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl38),
  inference(cnf_transformation,[],[f1242_D])).
fof(f1242_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1153,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1061])).
fof(f1061,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f37,f605])).
fof(f1257,plain,(
  ( ! [X43,X44] : (s('type.integer.int',X43) != s('type.integer.int',X44) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44)))))) ) | $spl42),
  inference(cnf_transformation,[],[f1257_D])).
fof(f1257_D,plain,(
  ( ! [X43,X44] : (s('type.integer.int',X43) != s('type.integer.int',X44) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44)))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1256,plain,(
  ( ! [X45,X43,X44] : (s('type.integer.int',X43) != s('type.integer.int',X44) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl10 | $spl26)),
  inference(subsumption_resolution,[],[f1056,f53])).
fof(f1056,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X45)) | s('type.integer.int',X43) != s('type.integer.int',X44) | s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))) = s(bool,f(s(bool,'const.integer.int_le_2'(s('type.integer.int',X43),s('type.integer.int',X44))))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f43,f605])).
fof(f1175,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,f(s(bool,X30))) | s(bool,X30) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1026])).
fof(f1026,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,f(s(bool,X30))) | s(bool,X30) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f605,f333])).
fof(f1176,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,X23)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1024])).
fof(f1024,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,f(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f37])).
fof(f1231,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) | $spl36),
  inference(cnf_transformation,[],[f1231_D])).
fof(f1231_D,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1177,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f1022])).
fof(f1022,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f1186,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f995])).
fof(f995,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f1187,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f993])).
fof(f993,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl3 | $spl7 | $spl10 | $spl26)),
  inference(superposition,[],[f605,f333])).
fof(f1192,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f988])).
fof(f988,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f37])).
fof(f1206,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f976])).
fof(f976,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f1222,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl10 | $spl26)),
  inference(duplicate_literal_removal,[],[f948])).
fof(f948,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl10 | $spl26)),
  inference(superposition,[],[f605,f605])).
fof(f867,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f834,f34])).
fof(f834,plain,(
  ~p(s(bool,'T_0')) | ($spl22 | $spl26)),
  inference(backward_demodulation,[],[f807,f36])).
fof(f807,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl22 | $spl26)),
  inference(duplicate_literal_removal,[],[f745])).
fof(f745,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl22 | $spl26)),
  inference(superposition,[],[f534,f538])).
fof(f832,plain,(
  $false | ($spl3 | $spl7 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f831,f59])).
fof(f831,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f830,f538])).
fof(f830,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl7 | $spl22)),
  inference(subsumption_resolution,[],[f735,f95])).
fof(f735,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl22)),
  inference(superposition,[],[f534,f101])).
fof(f829,plain,(
  $false | ($spl3 | $spl7 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f828,f59])).
fof(f828,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl7 | $spl22 | $spl26)),
  inference(condensation,[],[f827])).
fof(f827,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f826,f538])).
fof(f826,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl22)),
  inference(forward_demodulation,[],[f812,f101])).
fof(f812,plain,(
  ( ! [X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl22)),
  inference(duplicate_literal_removal,[],[f728])).
fof(f728,plain,(
  ( ! [X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl22)),
  inference(superposition,[],[f534,f333])).
fof(f825,plain,(
  $false | ($spl3 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f824,f59])).
fof(f824,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl22 | $spl26)),
  inference(condensation,[],[f823])).
fof(f823,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f817,f538])).
fof(f817,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f723])).
fof(f723,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f534,f37])).
fof(f822,plain,(
  $false | ($spl3 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f821,f59])).
fof(f821,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f721,f538])).
fof(f721,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f59,f534])).
fof(f820,plain,(
  $false | ($spl3 | $spl7 | $spl22 | $spl26)),
  inference(subsumption_resolution,[],[f819,f59])).
fof(f819,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl22 | $spl26)),
  inference(forward_demodulation,[],[f818,f538])).
fof(f818,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl7 | $spl22)),
  inference(forward_demodulation,[],[f722,f101])).
fof(f722,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | ($spl7 | $spl22)),
  inference(unit_resulting_resolution,[],[f95,f534])).
fof(f535,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl24),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f576,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f51,f508])).
fof(f538,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl26),
  inference(cnf_transformation,[],[f538_D])).
fof(f538_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f575,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f51,f508])).
fof(f574,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl10),
  inference(superposition,[],[f51,f508])).
fof(f573,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f51,f508])).
fof(f540,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl30),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f572,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl10),
  inference(superposition,[],[f37,f508])).
fof(f571,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl10),
  inference(superposition,[],[f508,f37])).
fof(f570,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl10),
  inference(superposition,[],[f508,f51])).
fof(f569,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f508,f51])).
fof(f568,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl10),
  inference(superposition,[],[f508,f51])).
fof(f567,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f508,f51])).
fof(f541,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl31),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f531,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl11),
  inference(superposition,[],[f509,f37])).
fof(f530,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl11),
  inference(superposition,[],[f509,f51])).
fof(f529,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f509,f51])).
fof(f528,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl11),
  inference(superposition,[],[f509,f51])).
fof(f527,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f509,f51])).
fof(f509,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl11),
  inference(cnf_transformation,[],[f509_D])).
fof(f509_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f510,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl12),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f517,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl21),
  inference(cnf_transformation,[],[f517_D])).
fof(f517_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f504,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f485,f37])).
fof(f485,plain,(
  ( ! [X28,X29,X27] : (s(X29,f(s(bool,'F_0'))) != s(X29,g(s(bool,'F_0'))) | s(fun(bool,X29),X27) = s(fun(bool,X29),X28) | s(X29,f(s(bool,'T_0'))) != s(X29,g(s(bool,'T_0')))) ) | ($spl3 | $spl7)),
  inference(duplicate_literal_removal,[],[f484])).
fof(f484,plain,(
  ( ! [X28,X29,X27] : (s(X29,f(s(bool,'T_0'))) != s(X29,g(s(bool,'T_0'))) | s(fun(bool,X29),X27) = s(fun(bool,X29),X28) | s(fun(bool,X29),X27) = s(fun(bool,X29),X28) | s(X29,f(s(bool,'F_0'))) != s(X29,g(s(bool,'F_0')))) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f49,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool)) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0')))) ) | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f460,f101])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) != s(X2,g(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f49,f333])).
fof(f515,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl18),
  inference(cnf_transformation,[],[f515_D])).
fof(f515_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f503,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f485,f51])).
fof(f502,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f485,f51])).
fof(f507,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(cnf_transformation,[],[f507_D])).
fof(f507_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f501,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f485,f51])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f485,f51])).
fof(f88,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) | $spl4),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f90,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | $spl7),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f86,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_leq_2'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK0))),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',sK1))))) | s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK1))))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',sK1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',sK0)))))))),
  inference(superposition,[],[f33,f37])).
fof(f57,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f37])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 215
% Passive clauses: 1036
% Generated clauses: 33219
% Final active clauses: 46
% Final passive clauses: 153
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 2803
% Trivial inequalities: 387
% Fw subsumption resolutions: 2701
% Fw demodulations: 2796
% Bw demodulations: 498
% Condensations: 4333
% Global subsumptions: 20
% 
% Simple tautologies: 404
% Equational tautologies: 442
% Forward subsumptions: 18183
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 48
% 
% Unit resulting resolution: 1207
% Forward superposition: 7648
% Backward superposition: 7140
% Self superposition: 422
% Equality factoring: 239
% Equality resolution: 36
% 
% Split clauses: 141
% Split components: 292
% SAT solver clauses: 1406
% SAT solver unit clauses: 27
% SAT solver binary clauses: 162
% 
% Sat splits: 141
% Sat splitting refutations: 66
% 
% Memory used [KB]: 12281
% Time elapsed: 1.615 s
% ------------------------------
% ------------------------------
% Success in time 2.932 s
