Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% (49780)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 2376
% Generated clauses: 15432
% Final active clauses: 115
% Final passive clauses: 2128
% Input formulas: 20
% Initial clauses: 48
% 
% Duplicate literals: 701
% Trivial inequalities: 339
% Fw subsumption resolutions: 2371
% Fw demodulations: 1867
% Bw demodulations: 9
% 
% Simple tautologies: 390
% Equational tautologies: 47
% Forward subsumptions: 5671
% 
% Binary resolution: 113
% Forward superposition: 3991
% Backward superposition: 5430
% Self superposition: 190
% Equality factoring: 143
% Equality resolution: 43
% 
% Split clauses: 992
% Split components: 2090
% SAT solver clauses: 1575
% SAT solver unit clauses: 18
% SAT solver binary clauses: 133
% 
% Sat splits: 991
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1575
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (49783)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 683
% Generated clauses: 8544
% Final active clauses: 92
% Final passive clauses: 507
% Input formulas: 20
% Initial clauses: 48
% 
% Duplicate literals: 649
% Trivial inequalities: 190
% Fw subsumption resolutions: 354
% Fw demodulations: 700
% Bw demodulations: 83
% Condensations: 1052
% Global subsumptions: 4
% 
% Simple tautologies: 236
% Equational tautologies: 41
% Forward subsumptions: 2012
% 
% Binary resolution: 63
% Forward superposition: 1930
% Backward superposition: 2690
% Self superposition: 130
% Equality factoring: 117
% Equality resolution: 20
% 
% Split clauses: 61
% Split components: 131
% SAT solver clauses: 1116
% SAT solver unit clauses: 52
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 8
% 
% Sat splits: 61
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1100
% TWLsolver calls for satisfiability: 223
% 
% Memory used [KB]: 9722
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f3339,plain,(
  $false),
  inference(subsumption_resolution,[],[f3338,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0),s(sK0,X1))) = s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1),s(sK0,X0)))) )),
  inference(subsumption_resolution,[],[f589,f87])).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.symmetric_inv_identity',aHL_TRUTH)).
fof(f589,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0),s(sK0,X1))) = s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1),s(sK0,X0)))) )),
  inference(superposition,[],[f153,f161])).
fof(f161,plain,(
  s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(subsumption_resolution,[],[f158,f89])).
fof(f89,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.symmetric_inv_identity',aHL_FALSITY)).
fof(f158,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(superposition,[],[f85,f90])).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.symmetric_inv_identity',aHL_BOOL_CASES)).
fof(f85,plain,(
  p(s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  p(s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & s(fun(sK0,fun(sK0,bool)),sK1) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f48])).
fof(f48,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & s(fun(X0,fun(X0,bool)),X1) != s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7))))),
  file('thm.relation.symmetric_inv_identity','thm.relation.symmetric_inv_identity')).
fof(f153,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X5))) = s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X4)))) )),
  inference(forward_demodulation,[],[f151,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X9,X10] : s(bool,'const.relation.inv_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X9),s(X6,X10))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X9)))),
  file('thm.relation.symmetric_inv_identity','thm.relation.inv_DEF')).
fof(f151,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X4))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(backward_demodulation,[],[f114,f108])).
fof(f108,plain,(
  ( ! [X4,X0,X5,X1] : (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK2(X1,X0)))),s(X0,sK3(X1,X0)))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK2(X1,X0)))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X4)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2)))) & (? [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) != s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))) | p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X2))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> ! [X9,X10] : s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X9))),s(X6,X10))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X10))),s(X6,X9))))),
  file('thm.relation.symmetric_inv_identity','thm.relation.symmetric_def')).
fof(f3338,plain,(
  s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),sK0,bool)),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool))))) != s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool))),s(sK0,sK5(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),sK0,bool))))),
  inference(forward_demodulation,[],[f3279,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X19,X20,X21] : s(bool,'const.relation.inv_3'(s(fun(X6,fun(X6,bool)),X19),s(X6,X20),s(X6,X21))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X19))),s(X6,X20))),s(X6,X21)))),
  file('thm.relation.symmetric_inv_identity','ARITY_1_3_const.relation.inv')).
fof(f3279,plain,(
  s(bool,'const.relation.inv_3'(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),sK0,bool)),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool))))) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool))))),s(sK0,sK5(happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))),sK0,bool))))),
  inference(unit_resulting_resolution,[],[f515,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,happ(s(fun(X0,bool),X3),s(X0,sK5(happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)),X3,X0,bool)))) != s(bool,'const.relation.inv_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK5(happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)),X3,X0,bool)),s(X0,X2))) | s(fun(X0,bool),X3) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2)))) )),
  inference(superposition,[],[f118,f114])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK5(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK5(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK5(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK5(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11,X6] : ! [X12,X13] : (s(fun(X6,X11),X12) = s(fun(X6,X11),X13) <=> ! [X9] : s(X11,happ(s(fun(X6,X11),X12),s(X6,X9))) = s(X11,happ(s(fun(X6,X11),X13),s(X6,X9))))),
  file('thm.relation.symmetric_inv_identity','thm.bool.FUN_EQ_THM')).
fof(f515,plain,(
  s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool))))) != s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK5(sK1,'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)),sK0,fun(sK0,bool)))))),
  inference(unit_resulting_resolution,[],[f86,f118])).
fof(f86,plain,(
  s(fun(sK0,fun(sK0,bool)),sK1) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(cnf_transformation,[],[f59])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 373
% Generated clauses: 3340
% Final active clauses: 47
% Final passive clauses: 206
% Input formulas: 20
% Initial clauses: 48
% 
% Duplicate literals: 192
% Trivial inequalities: 31
% Fw subsumption resolutions: 132
% Fw demodulations: 316
% Bw demodulations: 91
% Condensations: 303
% 
% Simple tautologies: 95
% Equational tautologies: 49
% Forward subsumptions: 1745
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 173
% Forward superposition: 1009
% Backward superposition: 804
% Self superposition: 52
% Equality factoring: 36
% Equality resolution: 6
% 
% Split clauses: 22
% Split components: 49
% SAT solver clauses: 440
% SAT solver unit clauses: 76
% SAT solver binary clauses: 73
% 
% Sat splits: 22
% Sat splitting refutations: 21
% 
% Memory used [KB]: 3582
% Time elapsed: 0.206 s
% ------------------------------
% ------------------------------
% Success in time 1.534 s
