Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (68392)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1031
% Passive clauses: 10550
% Generated clauses: 25832
% Final active clauses: 1005
% Final passive clauses: 9202
% Input formulas: 46
% Initial clauses: 191
% 
% Duplicate literals: 460
% Trivial inequalities: 72
% Fw subsumption resolutions: 2760
% Fw demodulations: 1465
% Bw demodulations: 25
% 
% Simple tautologies: 267
% Equational tautologies: 35
% Forward subsumptions: 9596
% 
% Binary resolution: 11004
% Factoring: 487
% Forward superposition: 3802
% Backward superposition: 5381
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 599
% Split components: 1204
% SAT solver clauses: 950
% SAT solver unit clauses: 62
% SAT solver binary clauses: 245
% 
% Sat splits: 599
% 
% TWLsolver clauses: 950
% TWLsolver calls for satisfiability: 484
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (68413)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 452
% Passive clauses: 2925
% Generated clauses: 13150
% Final active clauses: 390
% Final passive clauses: 2437
% Input formulas: 46
% Initial clauses: 191
% 
% Duplicate literals: 258
% Trivial inequalities: 91
% Fw subsumption resolutions: 971
% Fw demodulations: 849
% Bw demodulations: 59
% Condensations: 290
% Global subsumptions: 13
% 
% Simple tautologies: 331
% Equational tautologies: 18
% Forward subsumptions: 3992
% Bw demodulations to eq. taut.: 12
% 
% Binary resolution: 3314
% Factoring: 45
% Forward superposition: 2405
% Backward superposition: 1610
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 151
% Split components: 307
% SAT solver clauses: 3588
% SAT solver unit clauses: 191
% SAT solver binary clauses: 88
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 11
% 
% Sat splits: 151
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3535
% TWLsolver calls for satisfiability: 1754
% 
% Memory used [KB]: 14200
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% (68440)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2435
% Passive clauses: 21333
% Generated clauses: 116347
% Final active clauses: 1899
% Final passive clauses: 3813
% Input formulas: 46
% Initial clauses: 191
% Discarded non-redundant clauses: 6007
% 
% Duplicate literals: 590
% Trivial inequalities: 39
% Fw subsumption resolutions: 2339
% Fw demodulations: 27841
% Bw demodulations: 6929
% Condensations: 505
% Global subsumptions: 272
% 
% Simple tautologies: 85
% Equational tautologies: 148
% Forward subsumptions: 60163
% Bw demodulations to eq. taut.: 7
% 
% Unit resulting resolution: 69093
% Factoring: 1
% Forward superposition: 5593
% Backward superposition: 2574
% Self superposition: 31
% Equality factoring: 30
% Equality resolution: 1
% 
% Split clauses: 217
% Split components: 439
% SAT solver clauses: 23705
% SAT solver unit clauses: 6317
% SAT solver binary clauses: 11400
% 
% Sat splits: 217
% Sat splitting refutations: 1663
% 
% Memory used [KB]: 74199
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f17083,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f546,f547,f548,f549,f671,f672,f673,f674,f809,f810,f913,f946,f1053,f947,f1055,f948,f1056,f1058,f949,f1059,f950,f1060,f951,f1062,f952,f1064,f1456,f1455,f3357,f3358,f3359,f3360,f3471,f3472,f653,f3473,f3474,f3861,f3722,f3721,f3862,f4959,f4960,f657,f5747,f5748,f779,f7621,f7623,f8231,f8354,f8355,f8356,f8367,f8386,f8387,f8388,f8398,f8518,f8519,f8625,f8650,f8651,f8757,f8935,f8937,f8936,f8939,f8941,f9060,f9059,f9058,f9046,f9045,f9044,f9036,f9035,f9034,f9130,f9131,f9132,f9133,f9134,f9135,f9909,f8394,f9910,f8363,f10028,f8396,f10029,f8365,f8361,f10672,f10698,f10673,f10700,f10674,f10701,f10677,f10705,f10679,f10708,f10858,f10846,f10845,f10838,f10915,f10916,f10917,f10918,f12514,f12515,f12619,f12519,f12620,f12629,f7622,f12630,f12734,f13426,f13427,f3465,f14124,f14125,f3579,f14813,f14814,f3966,f658,f14898,f654,f655,f17068])).
fof(f17068,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17067,f418])).
fof(f418,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.REAL_IMP_LE_SUP',aHL_FALSITY)).
fof(f17067,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(forward_demodulation,[],[f16961,f889])).
fof(f889,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))),
  inference(resolution,[],[f254,f403])).
fof(f403,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.EQ_CLAUSES|split|2')).
fof(f254,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)))) & ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK15))))) & (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f113])).
fof(f113,plain,(
  ? [X0,X1] : (? [X5] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ? [X3] : ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) & ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ? [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ? [X3] : ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) & ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ? [X3] : ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) & ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & ? [X9] : ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9))))) & ? [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & ? [X9] : ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9))))) & ? [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6)))))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.real.REAL_IMP_LE_SUP')).
fof(f16961,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl10),
  inference(resolution,[],[f443,f779])).
fof(f443,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) )),
  inference(resolution,[],[f253,f257])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7,X14,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X9))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.real.REAL_LE_TRANS')).
fof(f253,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(cnf_transformation,[],[f159])).
fof(f655,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) ) | $spl4),
  inference(cnf_transformation,[],[f655_D])).
fof(f655_D,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f654,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0))))) ) | $spl2),
  inference(cnf_transformation,[],[f654_D])).
fof(f654_D,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f14898,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl34),
  inference(resolution,[],[f3966,f256])).
fof(f256,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X2,X0)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK17(X2,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X2,X0)),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f160])).
fof(f160,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : ? [X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f115])).
fof(f115,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X15] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X15),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X15),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) => ! [X14] : (p(s(bool,happ(s(fun('type.realax.real',bool),X15),s('type.realax.real',X14)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X15))))))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.real.REAL_SUP_UBOUND_LE')).
fof(f658,plain,(
  ( ! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) ) | $spl8),
  inference(cnf_transformation,[],[f658_D])).
fof(f658_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3966,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X8,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl34),
  inference(cnf_transformation,[],[f3966_D])).
fof(f3966_D,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X8,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f14814,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X12,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X13,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl0),
  inference(subsumption_resolution,[],[f14591,f441])).
fof(f441,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f381])).
fof(f381,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : ! [X7] : (s(X12,X7) = s(X12,X7) <=> p(s(bool,'T_0')))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.REFL_CLAUSE')).
fof(f14591,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X12,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X13,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl0),
  inference(superposition,[],[f256,f3805])).
fof(f3805,plain,(
  ( ! [X220] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X220,sK12))))) ) | $spl0),
  inference(resolution,[],[f653,f407])).
fof(f407,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.EQ_CLAUSES|split|0')).
fof(f14813,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X11,sK12)),s('type.realax.real',X11))))) ) | $spl0),
  inference(subsumption_resolution,[],[f14590,f441])).
fof(f14590,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X11,sK12)),s('type.realax.real',X11))))) ) | $spl0),
  inference(superposition,[],[f256,f3805])).
fof(f3579,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl32),
  inference(cnf_transformation,[],[f3579_D])).
fof(f3579_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f14125,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(subsumption_resolution,[],[f13905,f441])).
fof(f13905,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(superposition,[],[f256,f3557])).
fof(f3557,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real'))))),
  inference(resolution,[],[f551,f407])).
fof(f551,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real')))))),
  inference(resolution,[],[f250,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X1,X0)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12] : ! [X15] : (~? [X7] : p(s(bool,happ(s(fun(X12,bool),X15),s(X12,X7)))) <=> ! [X7] : ~p(s(bool,happ(s(fun(X12,bool),X15),s(X12,X7)))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.NOT_EXISTS_THM')).
fof(f250,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))))),
  inference(cnf_transformation,[],[f159])).
fof(f14124,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f13904,f441])).
fof(f13904,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f256,f3557])).
fof(f3465,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl30),
  inference(cnf_transformation,[],[f3465_D])).
fof(f3465_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f13427,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(subsumption_resolution,[],[f13207,f441])).
fof(f13207,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(superposition,[],[f256,f3443])).
fof(f3443,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real'))))),
  inference(resolution,[],[f550,f407])).
fof(f550,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real')))))),
  inference(resolution,[],[f250,f431])).
fof(f431,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X1,X0)))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f13426,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f13206,f441])).
fof(f13206,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f256,f3443])).
fof(f12734,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl74),
  inference(cnf_transformation,[],[f12734_D])).
fof(f12734_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f12630,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X10,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(resolution,[],[f553,f256])).
fof(f553,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(sK12,X9,'type.realax.real'))))) | p(s(bool,X9))) )),
  inference(resolution,[],[f250,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : ! [X15,X16] : ((p(s(bool,X15)) | ? [X7] : p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X7))))) <=> ? [X7] : (p(s(bool,X15)) | p(s(bool,happ(s(fun(X12,bool),X16),s(X12,X7))))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f7622,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl36),
  inference(cnf_transformation,[],[f7622_D])).
fof(f7622_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f12629,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X8,sK12)),s('type.realax.real',X8))))) )),
  inference(resolution,[],[f553,f256])).
fof(f12620,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(sK12,X17,'type.realax.real'))))) | p(s(bool,X17))) ) | $spl72),
  inference(cnf_transformation,[],[f12620_D])).
fof(f12620_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(sK12,X17,'type.realax.real'))))) | p(s(bool,X17))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f12519,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(sK12,X17,'type.realax.real'))))) | p(s(bool,X17))) )),
  inference(resolution,[],[f552,f430])).
fof(f552,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(sK12,X8,'type.realax.real'))))) | p(s(bool,X8))) )),
  inference(resolution,[],[f250,f428])).
fof(f428,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f12619,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl70),
  inference(cnf_transformation,[],[f12619_D])).
fof(f12619_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f12515,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X9,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X10,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(resolution,[],[f552,f256])).
fof(f12514,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X8,sK12)),s('type.realax.real',X8))))) )),
  inference(resolution,[],[f552,f256])).
fof(f10918,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f10840,f7622])).
fof(f10840,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(resolution,[],[f7622,f422])).
fof(f422,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.NOT_CLAUSES|split|2')).
fof(f10917,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f10839,f7622])).
fof(f10839,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(resolution,[],[f7622,f420])).
fof(f420,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.NOT_CLAUSES|split|1')).
fof(f10916,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f10837,f7622])).
fof(f10837,plain,(
  ( ! [X222] : (~p(s(bool,X222))) ) | $spl36),
  inference(resolution,[],[f7622,f416])).
fof(f416,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) => ~p(s(bool,X13)))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_F')).
fof(f10915,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f10835,f7622])).
fof(f10835,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl36),
  inference(resolution,[],[f7622,f388])).
fof(f388,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_CLAUSES|split|4')).
fof(f10838,plain,(
  $false | $spl36),
  inference(resolution,[],[f7622,f418])).
fof(f10845,plain,(
  $false | ($spl17 | $spl36)),
  inference(resolution,[],[f7622,f1055])).
fof(f10846,plain,(
  $false | ($spl21 | $spl36)),
  inference(resolution,[],[f7622,f1058])).
fof(f10858,plain,(
  $false | $spl36),
  inference(resolution,[],[f7622,f254])).
fof(f10708,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl69),
  inference(cnf_transformation,[],[f10708_D])).
fof(f10708_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f10679,plain,(
  ( ! [X38] : (p(s(bool,X38)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) | $spl10),
  inference(resolution,[],[f7804,f5886])).
fof(f5886,plain,(
  ( ! [X114,X115] : (sP6('const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X114,X115)) ) | $spl10),
  inference(resolution,[],[f779,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f7804,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f7798,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f7798,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f7769,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f58,f153,f152,f151,f150])).
fof(f151,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f153,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f58,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X8,X20,X6] : ((p(s(bool,X6)) <=> s(bool,X8) = s(bool,X20)) <=> ((p(s(bool,X6)) | p(s(bool,X20)) | p(s(bool,X8))) & (p(s(bool,X6)) | ~p(s(bool,X8)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X8)) | ~p(s(bool,X6))) & (p(s(bool,X8)) | ~p(s(bool,X20)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.sat.sat_2|49|_prove_ltb')).
fof(f7769,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f7767,f418])).
fof(f7767,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f7749,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f186])).
fof(f186,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f153])).
fof(f7749,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f7739,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f151])).
fof(f7739,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f7736,f319])).
fof(f7736,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f7732,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f193])).
fof(f7732,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f7726,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f7726,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f7725,f407])).
fof(f7725,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f7722,f7712])).
fof(f7712,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7707,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f7707,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f2166,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f2166,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) )),
  inference(superposition,[],[f2159,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f196])).
fof(f2159,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f2152,f253])).
fof(f2152,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16))))),
  inference(superposition,[],[f521,f407])).
fof(f521,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK16)))),
  inference(resolution,[],[f253,f402])).
fof(f402,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f7722,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f7713,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f7713,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7708,f321])).
fof(f7708,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f2166,f310])).
fof(f10705,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl67),
  inference(cnf_transformation,[],[f10705_D])).
fof(f10705_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f10677,plain,(
  ( ! [X36] : (p(s(bool,X36)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) | $spl6),
  inference(resolution,[],[f7804,f5098])).
fof(f5098,plain,(
  ( ! [X114,X115] : (sP6('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X114,X115)) ) | $spl6),
  inference(resolution,[],[f657,f327])).
fof(f10701,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X33,sK12))))) ) | $spl64),
  inference(cnf_transformation,[],[f10701_D])).
fof(f10701_D,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X33,sK12))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f10674,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X33,sK12))))) ) | $spl0),
  inference(resolution,[],[f7804,f3783])).
fof(f3783,plain,(
  ( ! [X182,X183,X181] : (sP6(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X181,sK12))),X182,X183)) ) | $spl0),
  inference(resolution,[],[f653,f327])).
fof(f10700,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real')))) | $spl63),
  inference(cnf_transformation,[],[f10700_D])).
fof(f10700_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real')))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f10673,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real'))))) )),
  inference(resolution,[],[f7804,f3535])).
fof(f3535,plain,(
  ( ! [X123,X122] : (sP6(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(sK12,'type.realax.real'))),X122,X123)) )),
  inference(resolution,[],[f551,f327])).
fof(f10698,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real')))) | $spl61),
  inference(cnf_transformation,[],[f10698_D])).
fof(f10698_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real')))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f10672,plain,(
  ( ! [X30] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real'))))) )),
  inference(resolution,[],[f7804,f3421])).
fof(f3421,plain,(
  ( ! [X123,X122] : (sP6(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(sK12,'type.realax.real'))),X122,X123)) )),
  inference(resolution,[],[f550,f327])).
fof(f8361,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl40),
  inference(cnf_transformation,[],[f8361_D])).
fof(f8361_D,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8365,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | $spl45),
  inference(cnf_transformation,[],[f8365_D])).
fof(f8365_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f10029,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | ~p(s(bool,X9))) ) | $spl21),
  inference(resolution,[],[f4186,f7845])).
fof(f7845,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f7844,f325])).
fof(f7844,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f7843,f388])).
fof(f7843,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f7836,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f7836,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f2168,f314])).
fof(f2168,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) )),
  inference(superposition,[],[f2159,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f4186,plain,(
  ( ! [X99,X98] : (sP8(X98,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))),X99)) ) | $spl21),
  inference(resolution,[],[f1058,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f190])).
fof(f8396,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | $spl50),
  inference(cnf_transformation,[],[f8396_D])).
fof(f8396_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10028,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | ~p(s(bool,X8))) ) | $spl21),
  inference(resolution,[],[f4186,f7803])).
fof(f7803,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f7797,f325])).
fof(f7797,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f7769,f330])).
fof(f8363,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | $spl43),
  inference(cnf_transformation,[],[f8363_D])).
fof(f8363_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f9910,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | ~p(s(bool,X9))) ) | $spl17),
  inference(resolution,[],[f4051,f7845])).
fof(f4051,plain,(
  ( ! [X99,X98] : (sP8(X98,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real'))),X99)) ) | $spl17),
  inference(resolution,[],[f1055,f318])).
fof(f8394,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | $spl48),
  inference(cnf_transformation,[],[f8394_D])).
fof(f8394_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9909,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | ~p(s(bool,X8))) ) | $spl17),
  inference(resolution,[],[f4051,f7803])).
fof(f9135,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9033,f8361])).
fof(f9033,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f8361,f421])).
fof(f421,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f235])).
fof(f9134,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9032,f8361])).
fof(f9032,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f8361,f419])).
fof(f419,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f234])).
fof(f9133,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9031,f8361])).
fof(f9031,plain,(
  ( ! [X224] : (p(s(bool,X224))) ) | $spl40),
  inference(resolution,[],[f8361,f399])).
fof(f399,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_CLAUSES|split|1')).
fof(f9132,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9029,f8361])).
fof(f9029,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f8361,f396])).
fof(f396,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_CLAUSES|split|2')).
fof(f9131,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9027,f8361])).
fof(f9027,plain,(
  ( ! [X221] : (p(s(bool,X221))) ) | $spl40),
  inference(resolution,[],[f8361,f393])).
fof(f393,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_CLAUSES|split|3')).
fof(f9130,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f9026,f8361])).
fof(f9026,plain,(
  ( ! [X220] : (p(s(bool,X220))) ) | $spl40),
  inference(resolution,[],[f8361,f390])).
fof(f390,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.IMP_CLAUSES|split|0')).
fof(f9034,plain,(
  $false | $spl40),
  inference(resolution,[],[f8361,f423])).
fof(f423,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.REAL_IMP_LE_SUP',aHL_TRUTH)).
fof(f9035,plain,(
  $false | $spl40),
  inference(resolution,[],[f8361,f424])).
fof(f424,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.TRUTH')).
fof(f9036,plain,(
  $false | $spl40),
  inference(resolution,[],[f8361,f441])).
fof(f9044,plain,(
  $false | ($spl0 | $spl40)),
  inference(resolution,[],[f8361,f653])).
fof(f9045,plain,(
  $false | $spl40),
  inference(resolution,[],[f8361,f550])).
fof(f9046,plain,(
  $false | $spl40),
  inference(resolution,[],[f8361,f551])).
fof(f9058,plain,(
  $false | ($spl18 | $spl40)),
  inference(resolution,[],[f8361,f1056])).
fof(f9059,plain,(
  $false | ($spl6 | $spl40)),
  inference(resolution,[],[f8361,f657])).
fof(f9060,plain,(
  $false | ($spl10 | $spl40)),
  inference(resolution,[],[f8361,f779])).
fof(f8941,plain,(
  $false | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f4001,f1057])).
fof(f1057,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))))) | $spl20),
  inference(cnf_transformation,[],[f1057_D])).
fof(f1057_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4001,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))))) | $spl17),
  inference(resolution,[],[f1055,f434])).
fof(f434,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : ! [X15] : (~! [X7] : p(s(bool,happ(s(fun(X12,bool),X15),s(X12,X7)))) <=> ? [X7] : ~p(s(bool,happ(s(fun(X12,bool),X15),s(X12,X7)))))),
  file('thm.real.REAL_IMP_LE_SUP','thm.bool.NOT_FORALL_THM')).
fof(f8939,plain,(
  $false | ($spl13 | $spl16)),
  inference(subsumption_resolution,[],[f8918,f418])).
fof(f8918,plain,(
  p(s(bool,'F_0')) | ($spl13 | $spl16)),
  inference(superposition,[],[f8520,f1708])).
fof(f1708,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13))) | $spl13),
  inference(resolution,[],[f913,f403])).
fof(f8520,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl16),
  inference(resolution,[],[f1054,f433])).
fof(f433,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f1054,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real'))))) | $spl16),
  inference(cnf_transformation,[],[f1054_D])).
fof(f1054_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real'))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8936,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl58),
  inference(cnf_transformation,[],[f8936_D])).
fof(f8936_D,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8937,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X9,sK12)),s('type.realax.real',X9))))) ) | $spl16),
  inference(subsumption_resolution,[],[f8826,f8520])).
fof(f8826,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X9,sK12)),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl16),
  inference(resolution,[],[f8520,f256])).
fof(f8935,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X7,sK12)),s('type.realax.real',X7))))) ) | $spl16),
  inference(subsumption_resolution,[],[f8825,f8520])).
fof(f8825,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X7,sK12)),s('type.realax.real',X7))))) ) | $spl16),
  inference(resolution,[],[f8520,f256])).
fof(f8757,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl56),
  inference(cnf_transformation,[],[f8757_D])).
fof(f8757_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8651,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl20),
  inference(resolution,[],[f1057,f256])).
fof(f8650,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) ) | $spl20),
  inference(resolution,[],[f1057,f256])).
fof(f8625,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl54),
  inference(cnf_transformation,[],[f8625_D])).
fof(f8625_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8519,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl16),
  inference(resolution,[],[f1054,f256])).
fof(f8518,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) ) | $spl16),
  inference(resolution,[],[f1054,f256])).
fof(f8398,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl52),
  inference(cnf_transformation,[],[f8398_D])).
fof(f8398_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8388,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~p(s(bool,X11))) )),
  inference(resolution,[],[f7803,f861])).
fof(f861,plain,(
  ( ! [X103,X104] : (sP8(X103,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X104)) )),
  inference(resolution,[],[f254,f318])).
fof(f8387,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | ~p(s(bool,X10))) ) | $spl21),
  inference(resolution,[],[f7803,f4186])).
fof(f8386,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | ~p(s(bool,X9))) ) | $spl17),
  inference(resolution,[],[f7803,f4051])).
fof(f8367,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl47),
  inference(cnf_transformation,[],[f8367_D])).
fof(f8367_D,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f8356,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~p(s(bool,X11))) )),
  inference(resolution,[],[f7845,f861])).
fof(f8355,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))) | ~p(s(bool,X10))) ) | $spl21),
  inference(resolution,[],[f7845,f4186])).
fof(f8354,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))) | ~p(s(bool,X9))) ) | $spl17),
  inference(resolution,[],[f7845,f4051])).
fof(f8231,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f8018,f441])).
fof(f8018,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl0),
  inference(superposition,[],[f256,f3943])).
fof(f3943,plain,(
  ( ! [X213] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X213,sK12)),s('type.realax.real',sK15)))) ) | $spl0),
  inference(resolution,[],[f3718,f407])).
fof(f3718,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X0,sK12)),s('type.realax.real',sK15))))) ) | $spl0),
  inference(resolution,[],[f653,f251])).
fof(f251,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK15))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f7623,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f7623_D])).
fof(f7623_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7621,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f7616])).
fof(f7616,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f2165,f331])).
fof(f2165,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) )),
  inference(superposition,[],[f2159,f328])).
fof(f779,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl10),
  inference(cnf_transformation,[],[f779_D])).
fof(f779_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5748,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(subsumption_resolution,[],[f5530,f441])).
fof(f5530,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(superposition,[],[f256,f757])).
fof(f757,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16)))),
  inference(resolution,[],[f252,f407])).
fof(f252,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16))))),
  inference(cnf_transformation,[],[f159])).
fof(f5747,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f5529,f441])).
fof(f5529,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f256,f757])).
fof(f657,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl6),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f4960,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(subsumption_resolution,[],[f4742,f441])).
fof(f4742,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(superposition,[],[f256,f632])).
fof(f632,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)))),
  inference(resolution,[],[f250,f407])).
fof(f4959,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f4741,f441])).
fof(f4741,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f256,f632])).
fof(f3862,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl0),
  inference(resolution,[],[f3718,f256])).
fof(f3721,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X7,sK12)),s('type.realax.real',X7))))) ) | $spl0),
  inference(resolution,[],[f653,f256])).
fof(f3722,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X8,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X9,sK12)),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl0),
  inference(resolution,[],[f653,f256])).
fof(f3861,plain,(
  $false | ($spl0 | $spl8)),
  inference(resolution,[],[f3718,f658])).
fof(f3474,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(resolution,[],[f551,f256])).
fof(f3473,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f551,f256])).
fof(f653,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) ) | $spl0),
  inference(cnf_transformation,[],[f653_D])).
fof(f653_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f3472,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f551,f255])).
fof(f255,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK17(X2,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f3471,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) )),
  inference(resolution,[],[f551,f255])).
fof(f3360,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(resolution,[],[f550,f256])).
fof(f3359,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f550,f256])).
fof(f3358,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f550,f255])).
fof(f3357,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) )),
  inference(resolution,[],[f550,f255])).
fof(f1455,plain,(
  $false | $spl4),
  inference(resolution,[],[f655,f250])).
fof(f1456,plain,(
  $false | $spl4),
  inference(resolution,[],[f655,f252])).
fof(f1064,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | $spl28),
  inference(cnf_transformation,[],[f1064_D])).
fof(f1064_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real',sK15)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f952,plain,(
  ( ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) )),
  inference(resolution,[],[f251,f432])).
fof(f1062,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | $spl26),
  inference(cnf_transformation,[],[f1062_D])).
fof(f1062_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real',sK15)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f951,plain,(
  ( ! [X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) )),
  inference(resolution,[],[f251,f431])).
fof(f1060,plain,(
  ( ! [X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X7,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X7))) ) | $spl24),
  inference(cnf_transformation,[],[f1060_D])).
fof(f1060_D,plain,(
  ( ! [X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X7,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X7))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f950,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(sK12,X7,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) )),
  inference(resolution,[],[f251,f430])).
fof(f1059,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X5,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f1059_D])).
fof(f1059_D,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X5,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f949,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(sK12,X5,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X5)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6))))) )),
  inference(resolution,[],[f251,f428])).
fof(f1058,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))))) | $spl21),
  inference(cnf_transformation,[],[f1058_D])).
fof(f1058_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real'))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1056,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK15))))) ) | $spl18),
  inference(cnf_transformation,[],[f1056_D])).
fof(f1056_D,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK15))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f948,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12,'type.realax.real')))))) )),
  inference(resolution,[],[f251,f434])).
fof(f1055,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real'))))) | $spl17),
  inference(cnf_transformation,[],[f1055_D])).
fof(f1055_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real'))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f947,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(sK12,'type.realax.real')))))) )),
  inference(resolution,[],[f251,f433])).
fof(f1053,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X0,sK12)),s('type.realax.real',sK15))))) ) | $spl14),
  inference(cnf_transformation,[],[f1053_D])).
fof(f1053_D,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X0,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f946,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X0,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) )),
  inference(resolution,[],[f251,f255])).
fof(f913,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | $spl13),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f810,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X2,sK12)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f254,f256])).
fof(f809,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X0,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) )),
  inference(resolution,[],[f254,f255])).
fof(f674,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(resolution,[],[f252,f256])).
fof(f673,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f252,f256])).
fof(f672,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f252,f255])).
fof(f671,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) )),
  inference(resolution,[],[f252,f255])).
fof(f549,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) )),
  inference(resolution,[],[f250,f256])).
fof(f548,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X5,sK12)),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f250,f256])).
fof(f547,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f250,f255])).
fof(f546,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK17(X1,sK12)))))) )),
  inference(resolution,[],[f250,f255])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 672
% Passive clauses: 4344
% Generated clauses: 16737
% Final active clauses: 607
% Final passive clauses: 2712
% Input formulas: 46
% Initial clauses: 188
% 
% Selected by SInE selection: 43
% SInE iterations: 1
% 
% Duplicate literals: 533
% Trivial inequalities: 93
% Fw subsumption resolutions: 2673
% Fw demodulations: 1649
% Bw demodulations: 29
% 
% Simple tautologies: 1045
% Equational tautologies: 1
% Forward subsumptions: 6594
% 
% Binary resolution: 7286
% Forward superposition: 1071
% Backward superposition: 3256
% Equality resolution: 2
% 
% Split clauses: 67
% Split components: 157
% SAT solver clauses: 97
% SAT solver unit clauses: 20
% SAT solver binary clauses: 33
% 
% Sat splits: 67
% Sat splitting refutations: 29
% 
% Memory used [KB]: 5117
% Time elapsed: 0.221 s
% ------------------------------
% ------------------------------
% Success in time 11.944 s
