Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (72873)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 212
% Passive clauses: 3699
% Generated clauses: 27637
% Final active clauses: 188
% Final passive clauses: 3136
% Input formulas: 69
% Initial clauses: 137
% 
% Duplicate literals: 484
% Trivial inequalities: 371
% Fw subsumption resolutions: 2368
% Fw demodulations: 5284
% Bw demodulations: 13
% 
% Simple tautologies: 403
% Equational tautologies: 45
% Forward subsumptions: 14919
% Fw demodulations to eq. taut.: 18
% 
% Binary resolution: 146
% Factoring: 1
% Forward superposition: 9105
% Backward superposition: 9680
% Self superposition: 280
% Equality factoring: 77
% Equality resolution: 56
% 
% Split clauses: 158
% Split components: 326
% SAT solver clauses: 337
% SAT solver unit clauses: 10
% SAT solver binary clauses: 79
% 
% Sat splits: 158
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 337
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 21875
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (73051)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 218
% Passive clauses: 1950
% Generated clauses: 20801
% Final active clauses: 166
% Final passive clauses: 1556
% Input formulas: 69
% Initial clauses: 137
% 
% Duplicate literals: 265
% Trivial inequalities: 195
% Fw subsumption resolutions: 550
% Fw demodulations: 3379
% Bw demodulations: 204
% Condensations: 482
% Global subsumptions: 16
% 
% Simple tautologies: 134
% Equational tautologies: 37
% Forward subsumptions: 10677
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 142
% Factoring: 2
% Forward superposition: 5913
% Backward superposition: 6700
% Self superposition: 206
% Equality factoring: 56
% Equality resolution: 36
% 
% Split clauses: 59
% Split components: 143
% SAT solver clauses: 2946
% SAT solver unit clauses: 174
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 33
% SAT solver learnt literals: 5
% 
% Sat splits: 59
% Sat splitting refutations: 27
% 
% TWLsolver clauses: 2881
% TWLsolver calls for satisfiability: 714
% 
% Memory used [KB]: 27632
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f13197,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f251,f406,f404,f250,f405,f403,f1625,f1639,f1716,f1718,f1720,f2552,f2566,f2565,f3308,f3365,f3750,f3758,f3752,f4196,f4258,f1638,f6918,f6970,f6968,f6966,f7001,f7032,f7040,f7067,f7030,f7028,f7369,f7370,f7377,f7374,f7372,f7378,f10520,f3759,f10521,f10528,f10530,f10532,f10536,f10538,f10542,f10544,f10547,f10549,f10551,f10553,f10555,f10557,f10559,f10561,f10583,f10611,f10855,f10856,f11507,f11512,f11515,f11516,f11465,f11519,f11511,f11509,f11527,f11529,f12902,f13157,f13161,f13089,f13090,f13167,f13166,f13164,f13171,f13176,f13181,f13182,f13114,f13186,f13175,f13173])).
fof(f13173,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | $spl60),
  inference(cnf_transformation,[],[f13173_D])).
fof(f13173_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13175,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | $spl62),
  inference(cnf_transformation,[],[f13175_D])).
fof(f13175_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13186,plain,(
  ( ! [X35,X34] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X34))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X35))) | s('type.extreal.extreal',X34) = s('type.extreal.extreal',X35)) ) | $spl68),
  inference(cnf_transformation,[],[f13186_D])).
fof(f13186_D,plain,(
  ( ! [X35,X34] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X34))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X35))) | s('type.extreal.extreal',X34) = s('type.extreal.extreal',X35)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13114,plain,(
  ( ! [X35,X34] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X34))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',X35))) | s('type.extreal.extreal',X34) = s('type.extreal.extreal',X35) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(superposition,[],[f357,f12965])).
fof(f12965,plain,(
  s('type.extreal.extreal',sK3) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f246,f247,f405,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | ((s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))) & (s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2)) | (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2)) => (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2)) => (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X11,X12] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12)) => (s('type.extreal.extreal',X10) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12))) <=> s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X12)))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.eq_sub_ladd')).
fof(f247,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK1)),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK2) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1)))) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f191])).
fof(f191,plain,(
  ? [X0,X1,X2,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ? [X0,X1,X2,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2) & ((s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ? [X0,X1,X2,X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) <~> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ? [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2)) & (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) <~> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X2) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X2)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) <=> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X3))) <=> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X8) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X8)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X9))) <=> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X8) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X8)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X9))) <=> s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.eq_add_sub_switch')).
fof(f246,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK1)),
  inference(cnf_transformation,[],[f192])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ((s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)) & (s('type.extreal.extreal',X1) != s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)))),
  inference(ennf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal',X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11,X12] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X10) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X10)) => (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X12))) <=> s('type.extreal.extreal',X11) = s('type.extreal.extreal',X12)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.EXTREAL_EQ_LADD')).
fof(f13182,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal',X25) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X24) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK3)))) ) | $spl66),
  inference(cnf_transformation,[],[f13182_D])).
fof(f13182_D,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal',X25) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X24) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK3)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f13181,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal',X25) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X24) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f13180,f13037])).
fof(f13037,plain,(
  s('type.extreal.extreal',sK3) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | $spl2),
  inference(forward_demodulation,[],[f12970,f7695])).
fof(f7695,plain,(
  ( ! [X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK1)))))) )),
  inference(superposition,[],[f7510,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.add_comm')).
fof(f7510,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK1)))) )),
  inference(forward_demodulation,[],[f7438,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : ! [X6,X25] : s(X13,X6) = s(X13,'const.ind_type.FCONS_3'(s(X13,X6),s(fun('type.num.num',X13),X25),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.eq_add_sub_switch','thm.ind_type.FCONS|split|0')).
fof(f7438,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK1))),s(fun('type.num.num','type.extreal.extreal'),X1),s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK1)))) )),
  inference(unit_resulting_resolution,[],[f246,f247,f325,f353])).
fof(f12970,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))))) | $spl2),
  inference(unit_resulting_resolution,[],[f246,f247,f405,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal',X1) != s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f13180,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal',X25) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))))) | s('type.extreal.extreal',X24) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f13109,f337])).
fof(f13109,plain,(
  ( ! [X24,X25] : (s('type.extreal.extreal',X24) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X25),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X25) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(superposition,[],[f354,f12965])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X2)) )),
  inference(cnf_transformation,[],[f235])).
fof(f13176,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal',X19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X18) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK3)))) ) | $spl64),
  inference(cnf_transformation,[],[f13176_D])).
fof(f13176_D,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal',X19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X18) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK3)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13171,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal',X19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X18) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f13170,f13037])).
fof(f13170,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal',X19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))))) | s('type.extreal.extreal',X18) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f13106,f337])).
fof(f13106,plain,(
  ( ! [X19,X18] : (s('type.extreal.extreal',X18) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal',X19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X18),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1)))) ) | $spl2),
  inference(superposition,[],[f353,f12965])).
fof(f13164,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | $spl54),
  inference(cnf_transformation,[],[f13164_D])).
fof(f13164_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13166,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | $spl56),
  inference(cnf_transformation,[],[f13166_D])).
fof(f13166_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13167,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK3) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',X0))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X0)) ) | $spl58),
  inference(cnf_transformation,[],[f13167_D])).
fof(f13167_D,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK3) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',X0))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X0)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13090,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK3) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',X1))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))) ) | $spl2),
  inference(superposition,[],[f357,f12965])).
fof(f13089,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK3) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',X0))) | s('type.extreal.extreal',sK1) = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))) ) | $spl2),
  inference(superposition,[],[f357,f12965])).
fof(f13161,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f13160,f404])).
fof(f13160,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | $spl2),
  inference(forward_demodulation,[],[f13081,f11103])).
fof(f11103,plain,(
  ( ! [X8,X7] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',X8)))))) )),
  inference(superposition,[],[f351,f7609])).
fof(f7609,plain,(
  ( ! [X1] : (s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK2)))))) )),
  inference(superposition,[],[f7509,f337])).
fof(f7509,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK2)))) )),
  inference(forward_demodulation,[],[f7439,f325])).
fof(f7439,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.ind_type.FCONS_3'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK2))),s(fun('type.num.num','type.extreal.extreal'),X1),s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK2)))) )),
  inference(unit_resulting_resolution,[],[f248,f249,f325,f353])).
fof(f249,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK2)),
  inference(cnf_transformation,[],[f192])).
fof(f248,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK2)),
  inference(cnf_transformation,[],[f192])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11,X12] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X12))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))),s('type.extreal.extreal',X12)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.add_assoc')).
fof(f13081,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f249,f12965,f358])).
fof(f13157,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f13156,f404])).
fof(f13156,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | $spl2),
  inference(forward_demodulation,[],[f13084,f11103])).
fof(f13084,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))))) | $spl2),
  inference(unit_resulting_resolution,[],[f248,f249,f12965,f358])).
fof(f12902,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12849,f11103])).
fof(f12849,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f7753,f8224])).
fof(f8224,plain,(
  ( ! [X14] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X14))) | s('type.extreal.extreal',sK3) = s('type.extreal.extreal',X14)) ) | $spl0),
  inference(subsumption_resolution,[],[f8223,f248])).
fof(f8223,plain,(
  ( ! [X14] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X14))) | s('type.extreal.extreal',sK3) = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2)) ) | $spl0),
  inference(subsumption_resolution,[],[f8074,f249])).
fof(f8074,plain,(
  ( ! [X14] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X14))) | s('type.extreal.extreal',sK3) = s('type.extreal.extreal',X14) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2)) ) | $spl0),
  inference(superposition,[],[f357,f403])).
fof(f7753,plain,(
  s('type.extreal.extreal',sK3) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))),s('type.extreal.extreal',sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f246,f247,f406,f354])).
fof(f11529,plain,(
  $false | ($spl0 | $spl21 | $spl23 | $spl44)),
  inference(subsumption_resolution,[],[f11528,f248])).
fof(f11528,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK2) | ($spl0 | $spl21 | $spl23 | $spl44)),
  inference(backward_demodulation,[],[f11509,f7764])).
fof(f7764,plain,(
  s('type.extreal.extreal',sK2) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | ($spl0 | $spl21 | $spl23)),
  inference(unit_resulting_resolution,[],[f6965,f6967,f403,f354])).
fof(f6967,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK3) | $spl23),
  inference(cnf_transformation,[],[f6967_D])).
fof(f6967_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK3) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f6965,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK3) | $spl21),
  inference(cnf_transformation,[],[f6965_D])).
fof(f6965_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK3) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f11527,plain,(
  $false | ($spl0 | $spl21 | $spl23 | $spl46)),
  inference(subsumption_resolution,[],[f11526,f249])).
fof(f11526,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK2) | ($spl0 | $spl21 | $spl23 | $spl46)),
  inference(backward_demodulation,[],[f11511,f7764])).
fof(f11509,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | $spl44),
  inference(cnf_transformation,[],[f11509_D])).
fof(f11509_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f11511,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | $spl46),
  inference(cnf_transformation,[],[f11511_D])).
fof(f11511_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f11519,plain,(
  ( ! [X31,X32] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X31))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X32))) | s('type.extreal.extreal',X31) = s('type.extreal.extreal',X32)) ) | $spl52),
  inference(cnf_transformation,[],[f11519_D])).
fof(f11519_D,plain,(
  ( ! [X31,X32] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X31))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X32))) | s('type.extreal.extreal',X31) = s('type.extreal.extreal',X32)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f11465,plain,(
  ( ! [X31,X32] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X31))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',X32))) | s('type.extreal.extreal',X31) = s('type.extreal.extreal',X32) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3)))) ) | ($spl0 | $spl21 | $spl23)),
  inference(superposition,[],[f357,f7764])).
fof(f11516,plain,(
  ( ! [X21,X22] : (s('type.extreal.extreal',X22) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X21),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X21) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',sK2)))) ) | $spl50),
  inference(cnf_transformation,[],[f11516_D])).
fof(f11516_D,plain,(
  ( ! [X21,X22] : (s('type.extreal.extreal',X22) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X21),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X21) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',sK2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f11515,plain,(
  ( ! [X21,X22] : (s('type.extreal.extreal',X22) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X21),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X21) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3)))) ) | ($spl0 | $spl21 | $spl23)),
  inference(forward_demodulation,[],[f11460,f7764])).
fof(f11460,plain,(
  ( ! [X21,X22] : (s('type.extreal.extreal',X21) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X22),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X22) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X21),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3)))) ) | ($spl0 | $spl21 | $spl23)),
  inference(superposition,[],[f354,f7764])).
fof(f11512,plain,(
  ( ! [X15,X16] : (s('type.extreal.extreal',X16) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X15) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK2)))) ) | $spl48),
  inference(cnf_transformation,[],[f11512_D])).
fof(f11512_D,plain,(
  ( ! [X15,X16] : (s('type.extreal.extreal',X16) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X15) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK2)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11507,plain,(
  ( ! [X15,X16] : (s('type.extreal.extreal',X16) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X15) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3)))) ) | ($spl0 | $spl21 | $spl23)),
  inference(forward_demodulation,[],[f11457,f7764])).
fof(f11457,plain,(
  ( ! [X15,X16] : (s('type.extreal.extreal',X15) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X16),s('type.extreal.extreal',sK2))) | s('type.extreal.extreal',X16) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))),s('type.extreal.extreal',sK3)))) ) | ($spl0 | $spl21 | $spl23)),
  inference(superposition,[],[f353,f7764])).
fof(f10856,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f10537,f407])).
fof(f407,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f257,f252])).
fof(f252,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.eq_add_sub_switch',aHL_TRUTH)).
fof(f257,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f45])).
fof(f45,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.NOT_CLAUSES|split|1')).
fof(f10537,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f415,f3759,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.IMP_ANTISYM_AX')).
fof(f415,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f407,f267])).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.EQ_CLAUSES|split|1')).
fof(f10855,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f10543,f407])).
fof(f10543,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(unit_resulting_resolution,[],[f415,f3759,f344])).
fof(f10611,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f421,f415,f3759,f573])).
fof(f573,plain,(
  ( ! [X12,X11] : (p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X11)) )),
  inference(forward_demodulation,[],[f572,f319])).
fof(f319,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : ! [X10] : s(X13,X10) = s(X13,'const.combin.I_1'(s(X13,X10)))),
  file('thm.extreal.eq_add_sub_switch','thm.combin.I_THM')).
fof(f572,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,'const.combin.I_1'(s(bool,X11))) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f564,f252])).
fof(f564,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X12) != s(bool,'const.combin.I_1'(s(bool,X11))) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f394,f439])).
fof(f439,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f263,f319])).
fof(f263,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.eq_add_sub_switch',aHL_BOOL_CASES)).
fof(f394,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f382])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X10,X38,X11,X39] : ((s(bool,X10) = s(bool,X38) & (p(s(bool,X38)) => s(bool,X11) = s(bool,X39))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X38)) => p(s(bool,X39)))))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.IMP_CONG')).
fof(f421,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f261,f268])).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f261,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : p(s(bool,'const.relation.WF_1'(s(fun(X13,fun(X13,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.extreal.eq_add_sub_switch','thm.relation.WF_EMPTY_REL')).
fof(f10583,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f415,f415,f3759,f482])).
fof(f482,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f475,f407])).
fof(f475,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f344,f263])).
fof(f10561,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f325,f3759,f394])).
fof(f10559,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f332,f3759,f394])).
fof(f332,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : ! [X23,X24,X25] : s(X13,X24) = s(X13,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X13,X23),s(X13,X24),s(fun('type.realax.real',X13),X25)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_case_def|split|1')).
fof(f10557,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f331,f3759,f394])).
fof(f331,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : ! [X23,X24,X25] : s(X13,X23) = s(X13,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X13,X23),s(X13,X24),s(fun('type.realax.real',X13),X25)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_case_def|split|0')).
fof(f10555,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f408,f3759,f394])).
fof(f408,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f320,f319])).
fof(f320,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : ! [X30] : s(X13,'const.combin.I_1'(s(X13,X30))) = s(X13,happ(s(fun(X13,X13),'const.combin.I_0'),s(X13,X30)))),
  file('thm.extreal.eq_add_sub_switch','ARITY_0_1_const.combin.I')).
fof(f10553,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f319,f3759,f394])).
fof(f10551,plain,(
  $false | ($spl14 | $spl16)),
  inference(unit_resulting_resolution,[],[f252,f3772,f3759,f394])).
fof(f3772,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f3752,f268])).
fof(f10549,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f252,f2579,f3759,f394])).
fof(f2579,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f2565,f268])).
fof(f10547,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f421,f3759,f394])).
fof(f10544,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f407,f415,f3759,f344])).
fof(f10542,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f3759,f415,f3759,f344])).
fof(f10538,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f407,f415,f3759,f344])).
fof(f10536,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f3759,f415,f3759,f344])).
fof(f10532,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f415,f3759,f272])).
fof(f272,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.EQ_CLAUSES|split|2')).
fof(f10530,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f421,f3759,f267])).
fof(f10528,plain,(
  $false | $spl16),
  inference(unit_resulting_resolution,[],[f252,f3759])).
fof(f10521,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) | $spl42),
  inference(cnf_transformation,[],[f10521_D])).
fof(f10521_D,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3759,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl16),
  inference(cnf_transformation,[],[f3759_D])).
fof(f3759_D,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f10520,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) )),
  inference(subsumption_resolution,[],[f10490,f271])).
fof(f271,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f10490,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f350,f663])).
fof(f663,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f657,f252])).
fof(f657,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f571,f421])).
fof(f571,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f570,f319])).
fof(f570,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f563,f252])).
fof(f563,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f347,f439])).
fof(f347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f350,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X26,X13] : ! [X33,X34,X25] : (s(fun(X13,X26),X25) = s(fun(X13,X26),'const.relation.WFREC_2'(s(fun(X13,fun(X13,bool)),X34),s(fun(fun(X13,X26),fun(X13,X26)),X33))) => (p(s(bool,'const.relation.WF_1'(s(fun(X13,fun(X13,bool)),X34)))) => ! [X10] : s(X26,happ(s(fun(X13,X26),X25),s(X13,X10))) = s(X26,happ(s(fun(X13,X26),happ(s(fun(fun(X13,X26),fun(X13,X26)),X33),s(fun(X13,X26),'const.relation.RESTRICT_3'(s(fun(X13,X26),X25),s(fun(X13,fun(X13,bool)),X34),s(X13,X10))))),s(X13,X10)))))),
  file('thm.extreal.eq_add_sub_switch','thm.relation.WFREC_COROLLARY')).
fof(f7378,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))))) ) | $spl40),
  inference(cnf_transformation,[],[f7378_D])).
fof(f7378_D,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7372,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | $spl36),
  inference(cnf_transformation,[],[f7372_D])).
fof(f7372_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f7374,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | $spl38),
  inference(cnf_transformation,[],[f7374_D])).
fof(f7374_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7377,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))))) ) | $spl24),
  inference(forward_demodulation,[],[f7376,f6970])).
fof(f7376,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f7375,f6970])).
fof(f7375,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f7340,f7333])).
fof(f7333,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3)))))) ) | $spl24),
  inference(superposition,[],[f341,f6970])).
fof(f341,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X1)))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10,X11] : s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X11)))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.sub_rneg')).
fof(f7340,plain,(
  ( ! [X7] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(superposition,[],[f343,f6970])).
fof(f343,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : (((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1))) | s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1))) => s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1))) => s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10,X11] : (((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X10) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X10)) | (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11))) => s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X11))))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.neg_sub')).
fof(f7370,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6)))))) ) | $spl34),
  inference(cnf_transformation,[],[f7370_D])).
fof(f7370_D,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7369,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6)))))) ) | $spl24),
  inference(forward_demodulation,[],[f7368,f6970])).
fof(f7368,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f7367,f6970])).
fof(f7367,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f7339,f7333])).
fof(f7339,plain,(
  ( ! [X6] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))),s('type.extreal.extreal',X6))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2)))))) ) | $spl24),
  inference(superposition,[],[f342,f6970])).
fof(f342,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f7028,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | $spl26),
  inference(cnf_transformation,[],[f7028_D])).
fof(f7028_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f7030,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | $spl28),
  inference(cnf_transformation,[],[f7030_D])).
fof(f7030_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7067,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK1))))) | $spl32),
  inference(cnf_transformation,[],[f7067_D])).
fof(f7067_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK1))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f7040,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | ($spl2 | $spl22)),
  inference(superposition,[],[f342,f7038])).
fof(f7038,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | ($spl2 | $spl22)),
  inference(backward_demodulation,[],[f6968,f405])).
fof(f7032,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK1))))) | $spl30),
  inference(cnf_transformation,[],[f7032_D])).
fof(f7032_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK1))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f7001,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK0))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0) | ($spl2 | $spl20)),
  inference(superposition,[],[f342,f6972])).
fof(f6972,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f6966,f405])).
fof(f6966,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK3) | $spl20),
  inference(cnf_transformation,[],[f6966_D])).
fof(f6966_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK3) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6968,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK3) | $spl22),
  inference(cnf_transformation,[],[f6968_D])).
fof(f6968_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK3) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6970,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | $spl24),
  inference(cnf_transformation,[],[f6970_D])).
fof(f6970_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6918,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK1),s('type.extreal.extreal',sK3))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK3) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK3) | $spl2),
  inference(superposition,[],[f342,f405])).
fof(f1638,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f1638_D])).
fof(f1638_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4258,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X35)) ) | $spl18),
  inference(cnf_transformation,[],[f4258_D])).
fof(f4258_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X35)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4196,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f4108])).
fof(f4108,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | $spl14),
  inference(superposition,[],[f692,f3799])).
fof(f3799,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl14),
  inference(subsumption_resolution,[],[f3778,f407])).
fof(f3778,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl14),
  inference(superposition,[],[f3752,f1441])).
fof(f1441,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f663,f663])).
fof(f692,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f681,f407])).
fof(f681,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f573,f263])).
fof(f3752,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl14),
  inference(cnf_transformation,[],[f3752_D])).
fof(f3752_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3758,plain,(
  ( ! [X43] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f3735,f408])).
fof(f3735,plain,(
  ( ! [X43] : (~p(s(bool,X43)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f326,f408])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : ! [X29,X10] : (p(s(bool,happ(s(fun(X13,bool),X29),s(X13,X10)))) => p(s(bool,happ(s(fun(X13,bool),X29),s(X13,'const.min.@_1'(s(fun(X13,bool),X29)))))))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.SELECT_AX')).
fof(f3750,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f3749,f408])).
fof(f3749,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3723,f252])).
fof(f3723,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f326,f524])).
fof(f524,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f408,f263])).
fof(f3365,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(cnf_transformation,[],[f3365_D])).
fof(f3365_D,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'T_0') = s(bool,X34)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3308,plain,(
  ( ! [X33,X34] : (s(bool,X34) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X34)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3233])).
fof(f3233,plain,(
  ( ! [X33,X34] : (s(bool,X34) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f692,f2606])).
fof(f2606,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl8),
  inference(subsumption_resolution,[],[f2585,f407])).
fof(f2585,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f2565,f1441])).
fof(f2565,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl8),
  inference(cnf_transformation,[],[f2565_D])).
fof(f2565_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2566,plain,(
  ( ! [X39] : (~p(s(bool,sK4(X39,'const.combin.I_0',bool)))) ) | $spl10),
  inference(cnf_transformation,[],[f2566_D])).
fof(f2566_D,plain,(
  ( ! [X39] : (~p(s(bool,sK4(X39,'const.combin.I_0',bool)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2552,plain,(
  ( ! [X39] : (~p(s(bool,sK4(X39,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f328,f408])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13] : ! [X29,X16] : (! [X10] : (s(X13,X10) = s(X13,X16) => p(s(bool,happ(s(fun(X13,bool),X29),s(X13,X10))))) => p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),X29)))))),
  file('thm.extreal.eq_add_sub_switch','thm.bool.JRH_INDUCT_UTIL')).
fof(f1720,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f251,f250,f249,f248,f247,f246,f252,f254,f255,f407,f261,f262,f263,f268,f267,f272,f271,f319,f408,f325,f326,f328,f327,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f341,f343,f342,f347,f344,f348,f349,f350,f351,f352,f354,f353,f358,f357,f360,f359,f364,f363,f366,f365,f373,f374,f409,f410,f394,f415,f421,f437,f439,f482,f524,f568,f569,f571,f573,f591,f609,f635,f636,f662,f663,f692,f696,f697,f698,f700,f785,f823,f825,f803,f873,f914,f916,f893,f932,f933,f934,f935,f936,f952,f962,f963,f964,f965,f966,f1332,f1337,f1338,f1339,f1345,f1365,f1379,f1380,f1381,f1382,f1383,f1436,f1437,f1438,f1439,f1440,f1441,f1602,f1603,f1604,f1605,f1606,f1625,f1619,f1620,f1621,f1622,f1623,f1638,f1653,f1664,f1713,f1714,f1715,f1717,f1719])).
fof(f1719,plain,(
  ( ! [X51,X49] : (~p(s(bool,X49)) | p(s(bool,X51))) ) | $spl4),
  inference(subsumption_resolution,[],[f1684,f1714])).
fof(f1684,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | s(bool,X51) != s(bool,sK4(X49,X50,bool)) | p(s(bool,X51))) ) | $spl4),
  inference(superposition,[],[f394,f1653])).
fof(f1717,plain,(
  ( ! [X36] : (~p(s(bool,X36))) ) | $spl4),
  inference(subsumption_resolution,[],[f1679,f1714])).
fof(f1679,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | s(bool,'F_0') != s(bool,sK4(X36,X37,bool))) ) | $spl4),
  inference(superposition,[],[f271,f1653])).
fof(f1715,plain,(
  ( ! [X32] : (p(s(bool,X32))) ) | $spl4),
  inference(subsumption_resolution,[],[f1677,f1714])).
fof(f1677,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | s(bool,'T_0') != s(bool,sK4(X32,X33,bool))) ) | $spl4),
  inference(superposition,[],[f267,f1653])).
fof(f1714,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,sK4(X6,X7,bool))) ) | $spl4),
  inference(global_subsumption,[],[f251,f250,f249,f248,f247,f246,f252,f254,f255,f407,f261,f262,f263,f268,f267,f272,f271,f319,f408,f325,f326,f328,f327,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f341,f343,f342,f347,f344,f348,f349,f350,f351,f352,f354,f353,f358,f357,f360,f359,f364,f363,f366,f365,f373,f374,f409,f410,f394,f415,f421,f437,f439,f482,f524,f568,f569,f571,f573,f591,f609,f635,f636,f662,f663,f692,f696,f697,f698,f700,f785,f823,f825,f803,f873,f914,f916,f893,f932,f933,f934,f935,f936,f952,f962,f963,f964,f965,f966,f1332,f1337,f1338,f1339,f1345,f1365,f1379,f1380,f1381,f1382,f1383,f1436,f1437,f1438,f1439,f1440,f1441,f1602,f1603,f1604,f1605,f1606,f1625,f1619,f1620,f1621,f1622,f1623,f1638,f1653,f1664,f1713])).
fof(f1713,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(global_subsumption,[],[f251,f250,f249,f248,f247,f246,f252,f254,f255,f407,f261,f262,f263,f268,f267,f272,f271,f319,f408,f325,f326,f328,f327,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f341,f343,f342,f347,f344,f348,f349,f350,f351,f352,f354,f353,f358,f357,f360,f359,f364,f363,f366,f365,f373,f374,f409,f410,f394,f415,f421,f437,f439,f482,f524,f568,f569,f571,f573,f591,f609,f635,f636,f662,f663,f692,f696,f697,f698,f700,f785,f823,f825,f803,f873,f914,f916,f893,f932,f933,f934,f935,f936,f952,f962,f963,f964,f965,f966,f1332,f1337,f1338,f1339,f1345,f1365,f1379,f1380,f1381,f1382,f1383,f1436,f1437,f1438,f1439,f1440,f1441,f1602,f1603,f1604,f1605,f1606,f1625,f1619,f1620,f1621,f1622,f1623,f1638,f1653,f1664])).
fof(f1664,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f1653,f437])).
fof(f1653,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f1645,f407])).
fof(f1645,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1638,f327])).
fof(f1623,plain,(
  ( ! [X45,X46,X44] : (s(bool,X44) != s(bool,X46) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X44),s(fun('type.num.num',bool),X45),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f692,f325])).
fof(f1622,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X41) != s(bool,X43) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X40),s(bool,X41),s(fun('type.realax.real',bool),X42))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f692,f332])).
fof(f1621,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X36) != s(bool,X39) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X36),s(bool,X37),s(fun('type.realax.real',bool),X38))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f692,f331])).
fof(f1620,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X34))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f692,f408])).
fof(f1619,plain,(
  ( ! [X33,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f692,f319])).
fof(f1606,plain,(
  ( ! [X43,X44,X42] : (s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X42),s(fun('type.num.num',bool),X43),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f692,f325])).
fof(f1605,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X39) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X38),s(bool,X39),s(fun('type.realax.real',bool),X40)))) )),
  inference(superposition,[],[f692,f332])).
fof(f1604,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X34) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X34),s(bool,X35),s(fun('type.realax.real',bool),X36)))) )),
  inference(superposition,[],[f692,f331])).
fof(f1603,plain,(
  ( ! [X33,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X32)))) )),
  inference(superposition,[],[f692,f408])).
fof(f1602,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X30)))) )),
  inference(superposition,[],[f692,f319])).
fof(f1440,plain,(
  ( ! [X43,X44,X42] : (s(bool,X42) = s(bool,'const.relation.WF_1'(s(fun(X44,fun(X44,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X42),s(fun('type.num.num',bool),X43),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f663,f325])).
fof(f1439,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X39) = s(bool,'const.relation.WF_1'(s(fun(X41,fun(X41,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X38),s(bool,X39),s(fun('type.realax.real',bool),X40)))) )),
  inference(superposition,[],[f663,f332])).
fof(f1438,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X34) = s(bool,'const.relation.WF_1'(s(fun(X37,fun(X37,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X34),s(bool,X35),s(fun('type.realax.real',bool),X36)))) )),
  inference(superposition,[],[f663,f331])).
fof(f1437,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,'const.relation.WF_1'(s(fun(X33,fun(X33,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X32)))) )),
  inference(superposition,[],[f663,f408])).
fof(f1436,plain,(
  ( ! [X30,X31] : (s(bool,X30) = s(bool,'const.relation.WF_1'(s(fun(X31,fun(X31,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X30)))) )),
  inference(superposition,[],[f663,f319])).
fof(f1383,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,X34) | s(bool,X36) = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X35),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f1337,f325])).
fof(f1382,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X33) = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X30),s(bool,X31),s(fun('type.realax.real',bool),X32))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1337,f332])).
fof(f1381,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X26) | s(bool,X29) = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X26),s(bool,X27),s(fun('type.realax.real',bool),X28))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1337,f331])).
fof(f1380,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X24) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1337,f408])).
fof(f1379,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X23) = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f1337,f319])).
fof(f1365,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1337,f263])).
fof(f1345,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1083,f252])).
fof(f1083,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f571,f437])).
fof(f1339,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1079,f252])).
fof(f1079,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f394,f437])).
fof(f1338,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1078,f252])).
fof(f1078,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f347,f437])).
fof(f1337,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1074,f252])).
fof(f1074,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f271,f437])).
fof(f1332,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f1061,f415])).
fof(f1061,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f569,f437])).
fof(f966,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,X28) | s(bool,X30) = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f437,f325])).
fof(f965,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X24),s(bool,X25),s(fun('type.realax.real',bool),X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f437,f332])).
fof(f964,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X23) = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X20),s(bool,X21),s(fun('type.realax.real',bool),X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f437,f331])).
fof(f963,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f437,f408])).
fof(f962,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f437,f319])).
fof(f952,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f437,f437])).
fof(f936,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X36),s(fun('type.num.num',bool),X37),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f437,f325])).
fof(f935,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X32),s(bool,X33),s(fun('type.realax.real',bool),X34)))) )),
  inference(superposition,[],[f437,f332])).
fof(f934,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X28),s(bool,X29),s(fun('type.realax.real',bool),X30)))) )),
  inference(superposition,[],[f437,f331])).
fof(f933,plain,(
  ( ! [X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26)))) )),
  inference(superposition,[],[f437,f408])).
fof(f932,plain,(
  ( ! [X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X24)))) )),
  inference(superposition,[],[f437,f319])).
fof(f893,plain,(
  ( ! [X78,X76,X74,X77,X75] : (p(s(bool,'|dollar|exists_1'(s(fun(X74,bool),X76)))) | s(X74,X78) = s(X74,sK4(X78,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(fun(X74,bool),X75),s(fun(X74,bool),X76),s(fun('type.realax.real',fun(X74,bool)),X77)),X74))) )),
  inference(superposition,[],[f327,f332])).
fof(f916,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,X64) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X63),s(bool,X64),s(fun('type.realax.real',bool),X65)))) )),
  inference(forward_demodulation,[],[f890,f319])).
fof(f890,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X64))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X63),s(bool,X64),s(fun('type.realax.real',bool),X65)))) )),
  inference(superposition,[],[f569,f332])).
fof(f914,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,X57) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X56),s(bool,X57),s(fun('type.realax.real',bool),X58)))) )),
  inference(forward_demodulation,[],[f888,f319])).
fof(f888,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X56),s(bool,X57),s(fun('type.realax.real',bool),X58)))) )),
  inference(superposition,[],[f439,f332])).
fof(f873,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X4),s(bool,X5),s(fun('type.realax.real',bool),X6)))) )),
  inference(superposition,[],[f332,f263])).
fof(f803,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,'|dollar|exists_1'(s(fun(X62,bool),X63)))) | s(X62,X66) = s(X62,sK4(X66,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(fun(X62,bool),X63),s(fun(X62,bool),X64),s(fun('type.realax.real',fun(X62,bool)),X65)),X62))) )),
  inference(superposition,[],[f327,f331])).
fof(f825,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X51),s(bool,X52),s(fun('type.realax.real',bool),X53)))) )),
  inference(forward_demodulation,[],[f800,f319])).
fof(f800,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X51),s(bool,X52),s(fun('type.realax.real',bool),X53)))) )),
  inference(superposition,[],[f569,f331])).
fof(f823,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X44),s(bool,X45),s(fun('type.realax.real',bool),X46)))) )),
  inference(forward_demodulation,[],[f798,f319])).
fof(f798,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X44))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X44),s(bool,X45),s(fun('type.realax.real',bool),X46)))) )),
  inference(superposition,[],[f439,f331])).
fof(f785,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(bool,X4),s(bool,X5),s(fun('type.realax.real',bool),X6)))) )),
  inference(superposition,[],[f331,f263])).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f699,f407])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f327,f263])).
fof(f698,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7)))) | s(X6,X9) = s(X6,sK4(X9,'const.ind_type.FCONS_3'(s(fun(X6,bool),X7),s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num','const.num.|48|_0')),X6))) )),
  inference(superposition,[],[f327,f325])).
fof(f697,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X5) = s(X3,sK4(X5,happ(s(fun(fun(X3,bool),fun(X3,bool)),'const.combin.I_0'),s(fun(X3,bool),X4)),X3))) )),
  inference(superposition,[],[f327,f408])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,'const.combin.I_1'(s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f327,f319])).
fof(f662,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f656,f252])).
fof(f656,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f571,f439])).
fof(f636,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X30) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f622,f319])).
fof(f622,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f569,f325])).
fof(f635,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f621,f319])).
fof(f621,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f439,f325])).
fof(f609,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f325,f263])).
fof(f591,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f581,f319])).
fof(f581,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(superposition,[],[f569,f408])).
fof(f569,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f559,f252])).
fof(f559,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f271,f439])).
fof(f568,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(forward_demodulation,[],[f553,f319])).
fof(f553,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X5)))) )),
  inference(superposition,[],[f439,f408])).
fof(f437,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f263,f263])).
fof(f410,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,fun('type.num.num',X0)),'const.pair.,_2'(s(X0,X1),s(fun('type.num.num',X0),X2))),s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f374,f335])).
fof(f409,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),'const.pair.pair_CASE_2'(s('type.pair.prod'('type.num.num',X1),'const.pair.,_2'(s('type.num.num',X4),s(X1,X5))),s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(backward_demodulation,[],[f374,f348])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13,X26,X27] : ! [X11,X10,X25] : s(X13,happ(s(fun(X27,X13),happ(s(fun(X26,fun(X27,X13)),X25),s(X26,X10))),s(X27,X11))) = s(X13,'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X27),'const.pair.,_2'(s(X26,X10),s(X27,X11))),s(fun(X26,fun(X27,X13)),X25)))),
  file('thm.extreal.eq_add_sub_switch','thm.pair.pair_case_thm')).
fof(f373,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X17,X18,X19] : ! [X20,X21,X22] : s(X17,happ(s(fun(X18,X17),X20),s(X18,happ(s(fun(X19,X18),X21),s(X19,X22))))) = s(X17,happ(s(fun(X19,X17),'_dst_x0x1_2'(s(fun(X18,X17),X20),s(fun(X19,X18),X21))),s(X19,X22)))),
  file('thm.extreal.eq_add_sub_switch','_dst_x0x1_ax')).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) & (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) <=> s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X10,X11,X12] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))) <=> s('type.realax.real',X11) = s('type.realax.real',X12))),
  file('thm.extreal.eq_add_sub_switch','thm.real.REAL_EQ_LADD')).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f241])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) | s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) & (s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) <=> s('type.extreal.extreal',X1) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10,X11,X12] : (s('type.extreal.extreal',X10) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))))) <=> s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.eq_sub_ladd_normal')).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal',X1) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) & (s('type.realax.real',X1) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2))) <=> s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X10,X11,X12] : (s('type.realax.real',X10) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X11),s('type.realax.real',X12))) <=> s('type.realax.real',X11) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X12))))),
  file('thm.extreal.eq_add_sub_switch','thm.real.REAL_EQ_SUB_LADD')).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f238])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X10,X11,X12] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X12))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))),s('type.realax.real',X12)))),
  file('thm.extreal.eq_add_sub_switch','thm.realax.REAL_ADD_ASSOC')).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.extreal.eq_add_sub_switch',aHL_EXT)).
fof(f348,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK5(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X26,X13] : ! [X35] : ? [X25] : ! [X8,X36,X37] : s(X26,happ(s(fun('type.ind_type.recspace'(X13),X26),X25),s('type.ind_type.recspace'(X13),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X13,X36),s(fun('type.num.num','type.ind_type.recspace'(X13)),X37))))) = s(X26,happ(s(fun(fun('type.num.num',X26),X26),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X13)),fun(fun('type.num.num',X26),X26)),happ(s(fun(X13,fun(fun('type.num.num','type.ind_type.recspace'(X13)),fun(fun('type.num.num',X26),X26))),happ(s(fun('type.num.num',fun(X13,fun(fun('type.num.num','type.ind_type.recspace'(X13)),fun(fun('type.num.num',X26),X26)))),X35),s('type.num.num',X8))),s(X13,X36))),s(fun('type.num.num','type.ind_type.recspace'(X13)),X37))),s(fun('type.num.num',X26),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X13),X26),X25),s(fun('type.num.num','type.ind_type.recspace'(X13)),X37)))))),
  file('thm.extreal.eq_add_sub_switch','thm.ind_type.CONSTR_REC')).
fof(f340,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X10,X11] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X11),s('type.realax.real',X10))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))),
  file('thm.extreal.eq_add_sub_switch','thm.real.REAL_NEG_SUB')).
fof(f339,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X10,X11] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))),
  file('thm.extreal.eq_add_sub_switch','thm.real.REAL_SUB_RNEG')).
fof(f338,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X10,X11] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X10),s('type.realax.real',X11))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X10)))),
  file('thm.extreal.eq_add_sub_switch','thm.realax.REAL_ADD_SYM')).
fof(f336,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X13] : ! [X6,X23,X24,X25] : s(X13,happ(s(fun('type.realax.real',X13),X25),s('type.realax.real',X6))) = s(X13,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s(X13,X23),s(X13,X24),s(fun('type.realax.real',X13),X25)))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_case_def|split|2')).
fof(f335,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : ! [X30,X31,X32] : s(X13,'const.ind_type.FCONS_3'(s(X13,X30),s(fun('type.num.num',X13),X31),s('type.num.num',X32))) = s(X13,happ(s(fun('type.num.num',X13),happ(s(fun(fun('type.num.num',X13),fun('type.num.num',X13)),happ(s(fun(X13,fun(fun('type.num.num',X13),fun('type.num.num',X13))),'const.ind_type.FCONS_0'),s(X13,X30))),s(fun('type.num.num',X13),X31))),s('type.num.num',X32)))),
  file('thm.extreal.eq_add_sub_switch','ARITY_0_3_const.ind_type.FCONS')).
fof(f334,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X13] : ! [X30,X31,X32] : s('type.ind_type.recspace'(X13),'const.ind_type.CONSTR_3'(s('type.num.num',X30),s(X13,X31),s(fun('type.num.num','type.ind_type.recspace'(X13)),X32))) = s('type.ind_type.recspace'(X13),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X13)),'type.ind_type.recspace'(X13)),happ(s(fun(X13,fun(fun('type.num.num','type.ind_type.recspace'(X13)),'type.ind_type.recspace'(X13))),'const.ind_type.CONSTR_1'(s('type.num.num',X30))),s(X13,X31))),s(fun('type.num.num','type.ind_type.recspace'(X13)),X32)))),
  file('thm.extreal.eq_add_sub_switch','ARITY_1_3_const.ind_type.CONSTR')).
fof(f333,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : ! [X6,X25,X28] : s(X13,happ(s(fun('type.num.num',X13),X25),s('type.num.num',X28))) = s(X13,'const.ind_type.FCONS_3'(s(X13,X6),s(fun('type.num.num',X13),X25),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28)))))),
  file('thm.extreal.eq_add_sub_switch','thm.ind_type.FCONS|split|1')).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f262,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X10] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_ainv_def|split|2')).
fof(f255,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f60])).
fof(f60,axiom,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_ainv_def|split|0')).
fof(f254,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f58])).
fof(f58,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.eq_add_sub_switch','thm.extreal.extreal_ainv_def|split|1')).
fof(f1718,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f251,f250,f249,f248,f247,f246,f252,f254,f255,f407,f261,f262,f263,f268,f267,f272,f271,f319,f408,f325,f326,f328,f327,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f341,f343,f342,f347,f344,f348,f349,f350,f351,f352,f354,f353,f358,f357,f360,f359,f364,f363,f366,f365,f373,f374,f409,f410,f394,f415,f421,f437,f439,f482,f524,f568,f569,f571,f573,f591,f609,f635,f636,f662,f663,f692,f696,f697,f698,f700,f785,f823,f825,f803,f873,f914,f916,f893,f932,f933,f934,f935,f936,f952,f962,f963,f964,f965,f966,f1332,f1337,f1338,f1339,f1345,f1365,f1379,f1380,f1381,f1382,f1383,f1436,f1437,f1438,f1439,f1440,f1441,f1602,f1603,f1604,f1605,f1606,f1625,f1619,f1620,f1621,f1622,f1623,f1638,f1653,f1664,f1713,f1714,f1715,f1717])).
fof(f1716,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f251,f250,f249,f248,f247,f246,f252,f254,f255,f407,f261,f262,f263,f268,f267,f272,f271,f319,f408,f325,f326,f328,f327,f331,f332,f333,f334,f335,f336,f337,f338,f339,f340,f341,f343,f342,f347,f344,f348,f349,f350,f351,f352,f354,f353,f358,f357,f360,f359,f364,f363,f366,f365,f373,f374,f409,f410,f394,f415,f421,f437,f439,f482,f524,f568,f569,f571,f573,f591,f609,f635,f636,f662,f663,f692,f696,f697,f698,f700,f785,f823,f825,f803,f873,f914,f916,f893,f932,f933,f934,f935,f936,f952,f962,f963,f964,f965,f966,f1332,f1337,f1338,f1339,f1345,f1365,f1379,f1380,f1381,f1382,f1383,f1436,f1437,f1438,f1439,f1440,f1441,f1602,f1603,f1604,f1605,f1606,f1625,f1619,f1620,f1621,f1622,f1623,f1638,f1653,f1664,f1713,f1714,f1715])).
fof(f1639,plain,(
  ( ! [X2,X1] : (s(bool,X2) != s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(cnf_transformation,[],[f1639_D])).
fof(f1639_D,plain,(
  ( ! [X2,X1] : (s(bool,X2) != s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1625,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1607])).
fof(f1607,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f692,f663])).
fof(f403,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | $spl0),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f405,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) | $spl2),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f250,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3)))),
  inference(cnf_transformation,[],[f192])).
fof(f404,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) | $spl1),
  inference(cnf_transformation,[],[f404_D])).
fof(f404_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f406,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) | $spl3),
  inference(cnf_transformation,[],[f406_D])).
fof(f406_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f251,plain,(
  s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK2))) != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK3),s('type.extreal.extreal',sK1))) | s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal',sK1))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK2),s('type.extreal.extreal',sK3)))),
  inference(cnf_transformation,[],[f192])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 156
% Passive clauses: 1395
% Generated clauses: 13098
% Final active clauses: 123
% Final passive clauses: 912
% Input formulas: 69
% Initial clauses: 137
% 
% Duplicate literals: 255
% Trivial inequalities: 45
% Fw subsumption resolutions: 294
% Fw demodulations: 2187
% Bw demodulations: 211
% Condensations: 445
% Global subsumptions: 5
% 
% Simple tautologies: 126
% Equational tautologies: 76
% Forward subsumptions: 7969
% Fw demodulations to eq. taut.: 252
% Bw demodulations to eq. taut.: 33
% 
% Unit resulting resolution: 1014
% Factoring: 1
% Forward superposition: 3907
% Backward superposition: 4093
% Self superposition: 174
% Equality factoring: 56
% Equality resolution: 17
% 
% Split clauses: 22
% Split components: 57
% SAT solver clauses: 1473
% SAT solver unit clauses: 214
% SAT solver binary clauses: 484
% 
% Sat splits: 22
% Sat splitting refutations: 27
% 
% Memory used [KB]: 14839
% Time elapsed: 0.516 s
% ------------------------------
% ------------------------------
% Success in time 1.85 s
