Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (102506)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 922
% Passive clauses: 9811
% Generated clauses: 25152
% Final active clauses: 891
% Final passive clauses: 7827
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 500
% Trivial inequalities: 87
% Fw subsumption resolutions: 2142
% Fw demodulations: 1076
% Bw demodulations: 11
% 
% Simple tautologies: 242
% Equational tautologies: 128
% Forward subsumptions: 11132
% Fw demodulations to eq. taut.: 15
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 12236
% Factoring: 2935
% Forward superposition: 2385
% Backward superposition: 3183
% Self superposition: 46
% Equality factoring: 26
% Equality resolution: 39
% 
% Split clauses: 418
% Split components: 837
% SAT solver clauses: 793
% SAT solver unit clauses: 24
% SAT solver binary clauses: 113
% 
% Sat splits: 418
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 783
% TWLsolver calls for satisfiability: 527
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (102521)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 572
% Passive clauses: 2685
% Generated clauses: 18075
% Final active clauses: 386
% Final passive clauses: 1228
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 265
% Trivial inequalities: 203
% Fw subsumption resolutions: 1233
% Fw demodulations: 2190
% Bw demodulations: 855
% Condensations: 422
% Global subsumptions: 22
% 
% Simple tautologies: 335
% Equational tautologies: 35
% Forward subsumptions: 7543
% Fw demodulations to eq. taut.: 27
% Bw demodulations to eq. taut.: 92
% 
% Binary resolution: 3073
% Factoring: 255
% Forward superposition: 3330
% Backward superposition: 2476
% Self superposition: 62
% Equality factoring: 28
% Equality resolution: 50
% 
% Split clauses: 280
% Split components: 576
% SAT solver clauses: 3015
% SAT solver unit clauses: 254
% SAT solver binary clauses: 119
% SAT solver learnt clauses: 47
% SAT solver learnt literals: 14
% 
% Sat splits: 280
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 2924
% TWLsolver calls for satisfiability: 1968
% 
% Memory used [KB]: 12537
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f21448,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f306,f545,f543,f305,f544,f542,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f5003,f4919,f5005,f4923,f5007,f4926,f5009,f4929,f5011,f4932,f5013,f4935,f5015,f4938,f5017,f4941,f5019,f4944,f5021,f4947,f5023,f4950,f5025,f4953,f5027,f4956,f5029,f4959,f5031,f4962,f5033,f4965,f5035,f4920,f4968,f5258,f5259,f5267,f5269,f5271,f5273,f5275,f5277,f5279,f5281,f5283,f5285,f5287,f5289,f5291,f5293,f5295,f5297,f5299,f5301,f5303,f5305,f5307,f5309,f5311,f5313,f5315,f5317,f5319,f5321,f5323,f5325,f5327,f5329,f5331,f5333,f5335,f5337,f5339,f5341,f5343,f5345,f5347,f5349,f5351,f5353,f5355,f5357,f5358,f5359,f5360,f5361,f5363,f5365,f5688,f5690,f5723,f5725,f5726,f5728,f5733,f5735,f5737,f5739,f5743,f5745,f5747,f5749,f5751,f5753,f5759,f5761,f5767,f5769,f5775,f5781,f5794,f5818,f5842,f5866,f5924,f5926,f5928,f5931,f5933,f5938,f5940,f5945,f5947,f5960,f5962,f5964,f5967,f5969,f5973,f5975,f6006,f6017,f6033,f6044,f6046,f6055,f6062,f6071,f6078,f6080,f6082,f6084,f6086,f6611,f6612,f6131,f6133,f6135,f6137,f6144,f6151,f6161,f6169,f6171,f6182,f6190,f6192,f6202,f6209,f6218,f6225,f6227,f6236,f6238,f6243,f6245,f6254,f6256,f6261,f6263,f6272,f6274,f6279,f6281,f6290,f6297,f6306,f6313,f6319,f6325,f6334,f6344,f6346,f6355,f6357,f6359,f6364,f6366,f6371,f6373,f6378,f6380,f6385,f6387,f6395,f6404,f6414,f6422,f6424,f6440,f6442,f6444,f6446,f6457,f6460,f6471,f6474,f6489,f6491,f6505,f6507,f6509,f6515,f6520,f6525,f6527,f6532,f6534,f6545,f6554,f6565,f6574,f6579,f6581,f6583,f6585,f6591,f6593,f6595,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6620,f6621,f6622,f6623,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6633,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6720,f6721,f6722,f6723,f6724,f6725,f6726,f6728,f6729,f6730,f6731,f6732,f6733,f6734,f6735,f6736,f6737,f6738,f6739,f6740,f6741,f6742,f6743,f6744,f6745,f6746,f6747,f6748,f6749,f6750,f6751,f6752,f6753,f6754,f6755,f6756,f6757,f6758,f6759,f6760,f6761,f6762,f6763,f6764,f6765,f6766,f6767,f6768,f6769,f6770,f6771,f6772,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6781,f6782,f6783,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6791,f6792,f6793,f6794,f6795,f6796,f6797,f6798,f6799,f6800,f6801,f6802,f6803,f6804,f6805,f6806,f6807,f6808,f6809,f6810,f6811,f6812,f6813,f6814,f6815,f6816,f6817,f6818,f6819,f6820,f6821,f8668,f8670,f8669,f12672,f12674,f12673,f12872,f12878,f12879,f13022,f13026,f13029,f13031,f13024,f13540,f13541,f13548,f13545,f13542,f13543,f13551,f13544,f13553,f13546,f13550,f13560,f13561,f13573,f13571,f20701,f20703,f20702,f20705,f20706,f20739,f20741,f20740,f20743,f5691,f20744,f21132,f21153,f21166,f21170,f21172,f21176,f21180,f21184,f21187,f21191,f21195,f21199,f21203,f21207,f21211,f21215,f21219,f21223,f21227,f21231,f21235,f21239,f21243,f21247,f21251,f21255,f21259,f21262,f21264,f21265,f21405])).
fof(f21405,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21404,f546])).
fof(f546,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f310,f307])).
fof(f307,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6',aHL_TRUTH)).
fof(f310,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f38])).
fof(f38,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.bool.NOT_CLAUSES|split|2')).
fof(f21404,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21403,f15725])).
fof(f15725,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f599,f12907,f781])).
fof(f781,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f775,f546])).
fof(f775,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f406,f321])).
fof(f321,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6',aHL_BOOL_CASES)).
fof(f406,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.bool.IMP_ANTISYM_AX')).
fof(f12907,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12883,f2226])).
fof(f2226,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2221,f547])).
fof(f547,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f318,f314])).
fof(f314,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.NUMERAL_DEF')).
fof(f318,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|0')).
fof(f2221,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2215])).
fof(f2215,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f413,f547])).
fof(f413,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & (((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) & ~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X16,X17] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17)))) <=> ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X17),s('type.integer.int',X16))) & ~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) | (s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_DIVIDES_MOD|48|')).
fof(f12883,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f12874,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X7,X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_INJ')).
fof(f12874,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f12873,f314])).
fof(f12873,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl3),
  inference(forward_demodulation,[],[f545,f314])).
fof(f599,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f546,f332])).
fof(f332,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.bool.EQ_CLAUSES|split|0')).
fof(f21403,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21124,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f554,f314])).
fof(f554,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f402,f314])).
fof(f402,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|2')).
fof(f21124,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(trivial_inequality_removal,[],[f21087])).
fof(f21087,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2306,f13033])).
fof(f13033,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f13032,f2400])).
fof(f2400,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2392,f416])).
fof(f2392,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2391])).
fof(f2391,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(subsumption_resolution,[],[f2390,f1641])).
fof(f1641,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f389,f420])).
fof(f420,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.integer_244_prove_ltb|split|1')).
fof(f389,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.SUC_NOT')).
fof(f2390,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7)) )),
  inference(forward_demodulation,[],[f2341,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f2341,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X7) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f422,f551])).
fof(f551,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f549,f315])).
fof(f315,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.ADD_|48|')).
fof(f549,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f400,f317])).
fof(f317,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.BIT|49|')).
fof(f422,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f13032,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12979,f307])).
fof(f12979,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f410,f12884])).
fof(f12884,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f5689,f5686])).
fof(f5686,plain,(
  ( ! [X4] : (sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(global_subsumption,[],[f306,f305,f307,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f389,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1438,f1437,f1436,f1435,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2299,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2847,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4122,f4121,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4441,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4676,f4675,f4534,f4533,f392,f4996,f393,f4997,f2295,f2219,f2306,f2226,f4155,f2879,f4173,f2223,f4109,f2878,f4091,f2871,f2868,f5368,f4995,f4994,f4454,f4418,f4415,f2791,f2788,f2717,f2714,f1469,f1468,f1467,f1466,f1465,f1464,f1463,f1462,f1461,f1460,f1459,f1458,f1457,f1456,f1455,f1454,f1453,f1452,f1451,f1412,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1404,f1403,f1402,f1401,f1400,f1399,f1398,f1397,f1396,f1395,f1394,f415,f414,f413,f412,f410,f399,f398,f4532,f4674,f546,f337,f336,f4919,f5685])).
fof(f5685,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f5676,f332])).
fof(f5676,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl4),
  inference(superposition,[],[f4919,f730])).
fof(f336,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.bool.EQ_CLAUSES|split|3')).
fof(f337,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f208])).
fof(f4674,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'T_0') != s(bool,f(s(X159,sK16(X160,X161,bool,X159)))) | s(fun(X159,bool),X160) = s(fun(X159,bool),X161) | s(bool,X162) = s(bool,g(s(X159,sK16(X160,X161,bool,X159)))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f455,f730])).
fof(f4532,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,X170) != s(bool,f(s(X167,sK16(X168,X169,bool,X167)))) | s(fun(X167,bool),X168) = s(fun(X167,bool),X169) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,g(s(X167,sK16(X168,X169,bool,X167))))) )),
  inference(superposition,[],[f455,f730])).
fof(f398,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_DIVIDES_NEG|split|0')).
fof(f399,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X16,X17] : s(bool,'const.integer.int_divides_2'(s('type.integer.int',X16),s('type.integer.int',X17))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))),s('type.integer.int',X17)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_DIVIDES_NEG|split|1')).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f414,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f415,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f1394,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int',X20))))) )),
  inference(superposition,[],[f431,f398])).
fof(f1395,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int',X23))))) )),
  inference(superposition,[],[f432,f398])).
fof(f1396,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int',X26))))) )),
  inference(superposition,[],[f447,f398])).
fof(f1397,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29))))) | p(s(bool,X30)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))))) )),
  inference(superposition,[],[f450,f398])).
fof(f1398,plain,(
  ( ! [X33,X31,X34,X32] : (~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int',X32))))) )),
  inference(superposition,[],[f459,f398])).
fof(f1399,plain,(
  ( ! [X37,X35,X38,X36] : (~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int',X36))))) )),
  inference(superposition,[],[f463,f398])).
fof(f1400,plain,(
  ( ! [X39,X41,X42,X40] : (~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int',X40))))) )),
  inference(superposition,[],[f467,f398])).
fof(f1401,plain,(
  ( ! [X45,X43,X46,X44] : (~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int',X44))))) )),
  inference(superposition,[],[f486,f398])).
fof(f1402,plain,(
  ( ! [X47,X50,X48,X49] : (~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48))))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int',X48))))) )),
  inference(superposition,[],[f495,f398])).
fof(f1403,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54) | p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int',X52))))) )),
  inference(superposition,[],[f503,f398])).
fof(f1404,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57) | ~p(s(bool,X57)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int',X56))))) )),
  inference(superposition,[],[f508,f398])).
fof(f1405,plain,(
  ( ! [X61,X59,X62,X60] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int',X60))))) )),
  inference(superposition,[],[f509,f398])).
fof(f1406,plain,(
  ( ! [X66,X64,X65,X63] : (~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65) | p(s(bool,X65)) | p(s(bool,X66)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int',X64))))) )),
  inference(superposition,[],[f510,f398])).
fof(f1407,plain,(
  ( ! [X70,X68,X69,X67] : (~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68))))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int',X68))))) )),
  inference(superposition,[],[f520,f398])).
fof(f1408,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73) | p(s(bool,X73)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int',X72))))) )),
  inference(superposition,[],[f525,f398])).
fof(f1409,plain,(
  ( ! [X78,X76,X77,X75] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int',X76))))) )),
  inference(superposition,[],[f526,f398])).
fof(f1410,plain,(
  ( ! [X80,X81,X79,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int',X80))))) )),
  inference(superposition,[],[f527,f398])).
fof(f1411,plain,(
  ( ! [X85,X83,X86,X84] : (sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int',X84))))) )),
  inference(superposition,[],[f530,f398])).
fof(f1412,plain,(
  ( ! [X90,X88,X87,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int',X88))))) )),
  inference(superposition,[],[f531,f398])).
fof(f1451,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))))) )),
  inference(superposition,[],[f431,f399])).
fof(f1452,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X24),s('type.integer.int',X25))))) )),
  inference(superposition,[],[f432,f399])).
fof(f1453,plain,(
  ( ! [X28,X29,X27] : (sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29) | p(s(bool,X29)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X27),s('type.integer.int',X28))))) )),
  inference(superposition,[],[f447,f399])).
fof(f1454,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31))) | p(s(bool,X32)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X30),s('type.integer.int',X31))))) )),
  inference(superposition,[],[f450,f399])).
fof(f1455,plain,(
  ( ! [X35,X33,X36,X34] : (~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X33),s('type.integer.int',X34))))) )),
  inference(superposition,[],[f459,f399])).
fof(f1456,plain,(
  ( ! [X39,X37,X38,X40] : (~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))))) )),
  inference(superposition,[],[f463,f399])).
fof(f1457,plain,(
  ( ! [X43,X41,X44,X42] : (~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X41),s('type.integer.int',X42))))) )),
  inference(superposition,[],[f467,f399])).
fof(f1458,plain,(
  ( ! [X47,X45,X48,X46] : (~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X45),s('type.integer.int',X46))))) )),
  inference(superposition,[],[f486,f399])).
fof(f1459,plain,(
  ( ! [X52,X50,X51,X49] : (~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X49),s('type.integer.int',X50))))) )),
  inference(superposition,[],[f495,f399])).
fof(f1460,plain,(
  ( ! [X54,X56,X55,X53] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56) | p(s(bool,X55)) | p(s(bool,X56)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X53),s('type.integer.int',X54))))) )),
  inference(superposition,[],[f503,f399])).
fof(f1461,plain,(
  ( ! [X59,X57,X60,X58] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59) | ~p(s(bool,X59)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X57),s('type.integer.int',X58))))) )),
  inference(superposition,[],[f508,f399])).
fof(f1462,plain,(
  ( ! [X61,X64,X62,X63] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X61),s('type.integer.int',X62))))) )),
  inference(superposition,[],[f509,f399])).
fof(f1463,plain,(
  ( ! [X68,X66,X67,X65] : (~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67) | p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X65),s('type.integer.int',X66))))) )),
  inference(superposition,[],[f510,f399])).
fof(f1464,plain,(
  ( ! [X70,X72,X71,X69] : (~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70))) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X69),s('type.integer.int',X70))))) )),
  inference(superposition,[],[f520,f399])).
fof(f1465,plain,(
  ( ! [X76,X74,X75,X73] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75) | p(s(bool,X75)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X73),s('type.integer.int',X74))))) )),
  inference(superposition,[],[f525,f399])).
fof(f1466,plain,(
  ( ! [X80,X78,X79,X77] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80) | ~p(s(bool,X79)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X77),s('type.integer.int',X78))))) )),
  inference(superposition,[],[f526,f399])).
fof(f1467,plain,(
  ( ! [X83,X81,X84,X82] : (~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X81),s('type.integer.int',X82))))) )),
  inference(superposition,[],[f527,f399])).
fof(f1468,plain,(
  ( ! [X88,X87,X85,X86] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X85),s('type.integer.int',X86))))) )),
  inference(superposition,[],[f530,f399])).
fof(f1469,plain,(
  ( ! [X92,X90,X91,X89] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X89),s('type.integer.int',X90))))) )),
  inference(superposition,[],[f531,f399])).
fof(f2714,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f398])).
fof(f2717,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) )),
  inference(superposition,[],[f410,f399])).
fof(f2788,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X3),s('type.integer.int',X4)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) )),
  inference(superposition,[],[f412,f398])).
fof(f2791,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int',X10)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f412,f399])).
fof(f4415,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int',X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) )),
  inference(superposition,[],[f730,f398])).
fof(f4418,plain,(
  ( ! [X39,X37,X38] : (s(bool,X39) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X37),s('type.integer.int',X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)))) )),
  inference(superposition,[],[f730,f399])).
fof(f4454,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int',X22))) | s(bool,X23) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X21),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f730,f398])).
fof(f4994,plain,(
  ( ! [X127,X125,X126] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f4663,f307])).
fof(f4663,plain,(
  ( ! [X127,X125,X126] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X126),s('type.integer.int',X125))) | s('type.integer.int',X126) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X127) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X125),s('type.integer.int',X126))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f412,f730])).
fof(f4995,plain,(
  ( ! [X130,X128,X129] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f4664,f307])).
fof(f4664,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X129),s('type.integer.int',X128))) | s('type.integer.int',X128) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X130) = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X128),s('type.integer.int',X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f410,f730])).
fof(f5368,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2708,f2868])).
fof(f2868,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2867,f547])).
fof(f2867,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(subsumption_resolution,[],[f2848,f331])).
fof(f2848,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))))) )),
  inference(superposition,[],[f414,f547])).
fof(f2871,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2850,f331])).
fof(f2850,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4)))))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f414,f331])).
fof(f4091,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f547,f2400,f414])).
fof(f2878,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f2861,f2400])).
fof(f2861,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X3))))))) )),
  inference(superposition,[],[f414,f555])).
fof(f4109,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4094,f399])).
fof(f4094,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2348,f2400,f413])).
fof(f2223,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(forward_demodulation,[],[f2213,f548])).
fof(f2213,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1060,f548,f413])).
fof(f4173,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(forward_demodulation,[],[f4158,f399])).
fof(f4158,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f2348,f2413,f413])).
fof(f2879,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(subsumption_resolution,[],[f2862,f2413])).
fof(f2862,plain,(
  ( ! [X6,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))))) )),
  inference(superposition,[],[f414,f557])).
fof(f4155,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f548,f2413,f414])).
fof(f2219,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2217])).
fof(f2217,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f413,f331])).
fof(f2295,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(trivial_inequality_removal,[],[f2293])).
fof(f2293,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X5) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5)))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f415,f331])).
fof(f4997,plain,(
  ( ! [X154,X152,X155,X153] : (p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f4673,f307])).
fof(f4673,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | s(bool,X155) = s(bool,happ(s(fun(X152,bool),X153),s(X152,sK15(X153,X152)))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f393,f730])).
fof(f393,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X12] : ! [X14] : (~! [X8] : p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))) <=> ? [X8] : ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,X8)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.bool.NOT_FORALL_THM')).
fof(f4996,plain,(
  ( ! [X151,X149,X150,X148] : (p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(subsumption_resolution,[],[f4672,f307])).
fof(f4672,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X148,bool),X149),s(X148,X150)))) | s(bool,X151) = s(bool,happ(s(fun(X148,bool),X149),s(X148,sK14(X149,X148)))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f392,f730])).
fof(f392,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f4533,plain,(
  ( ! [X173,X171,X172] : (~p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X172)))) | p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X173)))) | s(bool,'T_0') = s(bool,X172) | s(bool,'T_0') = s(bool,sK14(X171,bool))) )),
  inference(superposition,[],[f392,f730])).
fof(f4534,plain,(
  ( ! [X175,X176,X174] : (~p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X175)))) | p(s(bool,happ(s(fun(bool,bool),X174),s(bool,X176)))) | s(bool,'T_0') = s(bool,X175) | s(bool,'T_0') = s(bool,sK15(X174,bool))) )),
  inference(superposition,[],[f393,f730])).
fof(f4675,plain,(
  ( ! [X165,X163,X164] : (~p(s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X163),s(bool,X164)))) | s(bool,X165) = s(bool,sK14(X163,bool)) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f392,f730])).
fof(f4676,plain,(
  ( ! [X167,X166,X168] : (~p(s(bool,happ(s(fun(bool,bool),X166),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X166),s(bool,X167)))) | s(bool,X168) = s(bool,sK15(X166,bool)) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f393,f730])).
fof(f4677,plain,(
  ( ! [X171,X169,X172,X170] : (s(X171,f(s(bool,'T_0'))) != s(X171,g(s(bool,'T_0'))) | s(fun(bool,X171),X169) = s(fun(bool,X171),X170) | s(bool,X172) = s(bool,sK16(X169,X170,X171,bool)) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f455,f730])).
fof(f4992,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4649,f307])).
fof(f4649,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f531,f730])).
fof(f4991,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4648,f307])).
fof(f4648,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f530,f730])).
fof(f4990,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4647,f307])).
fof(f4647,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f527,f730])).
fof(f4989,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4646,f307])).
fof(f4646,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f526,f730])).
fof(f4988,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4645,f307])).
fof(f4645,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f525,f730])).
fof(f4987,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4644,f307])).
fof(f4644,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f520,f730])).
fof(f4986,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4643,f307])).
fof(f4643,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f510,f730])).
fof(f4985,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4642,f307])).
fof(f4642,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f509,f730])).
fof(f4984,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4641,f307])).
fof(f4641,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f508,f730])).
fof(f4983,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4640,f307])).
fof(f4640,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f503,f730])).
fof(f4982,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4639,f307])).
fof(f4639,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f495,f730])).
fof(f4981,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4638,f307])).
fof(f4638,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f486,f730])).
fof(f4980,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4637,f307])).
fof(f4637,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f467,f730])).
fof(f4979,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4636,f307])).
fof(f4636,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f463,f730])).
fof(f4978,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4635,f307])).
fof(f4635,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f459,f730])).
fof(f4977,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4634,f307])).
fof(f4634,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f450,f730])).
fof(f4976,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4633,f307])).
fof(f4633,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f447,f730])).
fof(f4975,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4632,f307])).
fof(f4632,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f432,f730])).
fof(f4974,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4631,f307])).
fof(f4631,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f431,f730])).
fof(f4973,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4630,f307])).
fof(f4630,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f409,f730])).
fof(f4972,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4626,f307])).
fof(f4626,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f336,f730])).
fof(f4593,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f332,f730])).
fof(f4535,plain,(
  ( ! [X177,X180,X178,X179] : (s(X179,f(s(bool,X180))) != s(X179,g(s(bool,X180))) | s(fun(bool,X179),X177) = s(fun(bool,X179),X178) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,sK16(X177,X178,X179,bool))) )),
  inference(superposition,[],[f455,f730])).
fof(f4441,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f730,f730])).
fof(f4393,plain,(
  ~sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3484,f457])).
fof(f4388,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f3479,f457])).
fof(f4380,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(forward_demodulation,[],[f4336,f1814])).
fof(f4336,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f553,f1884])).
fof(f4335,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f551,f1884])).
fof(f4334,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))),s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.num.num',X4)))) )),
  inference(superposition,[],[f401,f1884])).
fof(f4333,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1884])).
fof(f4312,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f571,f576,f3478,f460])).
fof(f4303,plain,(
  ~sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f3473,f458])).
fof(f4260,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X25))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))))) | s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f416,f1787])).
fof(f4297,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f4248,f1814])).
fof(f4248,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f553,f1787])).
fof(f4247,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f551,f1787])).
fof(f4246,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f401,f1787])).
fof(f4245,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f1787])).
fof(f4189,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1657,f418])).
fof(f4121,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1641,f331])).
fof(f4122,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1641,f418])).
fof(f4062,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1438,f458])).
fof(f4009,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1436,f2348])).
fof(f4006,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1436,f458])).
fof(f3977,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f1435,f2348])).
fof(f3949,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f1383,f458])).
fof(f3902,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1381,f2348])).
fof(f3899,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f1381,f458])).
fof(f3856,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1380,f2348])).
fof(f3771,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3693,f2348])).
fof(f3693,plain,(
  ( ! [X63] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X63))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X63) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f418,f2348])).
fof(f3657,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f2348])).
fof(f3452,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2299,f409])).
fof(f3454,plain,(
  ( ! [X0] : (sP0(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f431])).
fof(f3455,plain,(
  ( ! [X0] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2299,f432])).
fof(f3456,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f2299,f446])).
fof(f3458,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2299,f446])).
fof(f3459,plain,(
  ~sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2299,f2299,f446])).
fof(f3460,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f546,f2299,f447])).
fof(f3461,plain,(
  sP1('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2299,f450])).
fof(f3464,plain,(
  ~sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f459])).
fof(f3465,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f459])).
fof(f3467,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f571,f568,f575,f2299,f460])).
fof(f3468,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f463])).
fof(f3473,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f463])).
fof(f3474,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f463])).
fof(f3478,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f464])).
fof(f3479,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f467])).
fof(f3484,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f467])).
fof(f3485,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f467])).
fof(f3489,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f468])).
fof(f3490,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f472])).
fof(f3491,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2299,f473])).
fof(f3492,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2299,f474])).
fof(f3497,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f398,f2299,f476])).
fof(f3498,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2299,f476])).
fof(f3503,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f398,f2299,f476])).
fof(f3504,plain,(
  ( ! [X0,X1] : (sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f399,f2299,f476])).
fof(f3505,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f599,f2299,f477])).
fof(f3507,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f599,f2299,f477])).
fof(f3509,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f2299,f480])).
fof(f3510,plain,(
  ( ! [X0] : (~sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2299,f481])).
fof(f3517,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f486])).
fof(f3522,plain,(
  ~sP6('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f486])).
fof(f3523,plain,(
  ~sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f486])).
fof(f3527,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f487])).
fof(f3528,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2299,f492])).
fof(f3533,plain,(
  sP7('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f307,f2299,f492])).
fof(f3535,plain,(
  sP7('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2299,f2299,f492])).
fof(f3538,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f2299,f492])).
fof(f3539,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f2299,f307,f2299,f492])).
fof(f3540,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f2299,f2299,f492])).
fof(f3541,plain,(
  sP7('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f492])).
fof(f3542,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f495])).
fof(f3546,plain,(
  ~sP7('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f495])).
fof(f3547,plain,(
  ~sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f495])).
fof(f3550,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f546,f2299,f496])).
fof(f3552,plain,(
  ( ! [X0] : (~sP9('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2299,f497])).
fof(f3556,plain,(
  ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2299,f503])).
fof(f3559,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f504])).
fof(f3560,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2299,f505])).
fof(f3561,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f307,f2299,f508])).
fof(f3565,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f2299,f2299,f508])).
fof(f3566,plain,(
  ( ! [X0] : (sP9('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2299,f509])).
fof(f3570,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2299,f2299,f509])).
fof(f3571,plain,(
  ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f2299,f510])).
fof(f3572,plain,(
  ( ! [X0] : (~sP11('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f546,f2299,f514])).
fof(f3587,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f2299,f520])).
fof(f3592,plain,(
  ~sP10('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f307,f2299,f520])).
fof(f3593,plain,(
  ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f2299,f2299,f520])).
fof(f3597,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) )),
  inference(unit_resulting_resolution,[],[f2299,f522])).
fof(f3598,plain,(
  sP11('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f546,f546,f2299,f524])).
fof(f3601,plain,(
  ( ! [X0] : (sP11('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f307,f2299,f526])).
fof(f3605,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) )),
  inference(unit_resulting_resolution,[],[f2299,f2299,f526])).
fof(f3606,plain,(
  ~sP11('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f527])).
fof(f3609,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f2299,f527])).
fof(f3610,plain,(
  ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f2299,f546,f2299,f527])).
fof(f3611,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f583,f2299,f530])).
fof(f3613,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f584,f2299,f530])).
fof(f3614,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f580,f2299,f531])).
fof(f3615,plain,(
  ( ! [X0] : (sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f581,f2299,f531])).
fof(f3408,plain,(
  ( ! [X21,X20] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X21))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))))) | s('type.num.num',X21) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20)))))) )),
  inference(superposition,[],[f416,f1814])).
fof(f3398,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f551,f1814])).
fof(f3432,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(forward_demodulation,[],[f3396,f1884])).
fof(f3396,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f553,f1814])).
fof(f3395,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f400,f1814])).
fof(f3394,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f401,f1814])).
fof(f3337,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f416,f1572])).
fof(f3327,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f551,f1572])).
fof(f3364,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3363,f1045])).
fof(f3363,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f3325,f1884])).
fof(f3325,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1572])).
fof(f3362,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f3324,f1809])).
fof(f3324,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f400,f1572])).
fof(f3361,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f3323,f1880])).
fof(f3323,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f401,f1572])).
fof(f3140,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f559])).
fof(f3127,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f559])).
fof(f2952,plain,(
  ( ! [X39,X38,X40] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))))) | s('type.integer.int',X40) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X38),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39)))))))))) )),
  inference(superposition,[],[f418,f558])).
fof(f2939,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(superposition,[],[f331,f558])).
fof(f2847,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f1061,f547,f414])).
fof(f2708,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f410,f331])).
fof(f2605,plain,(
  ( ! [X24,X23,X22] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23))))))))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23)))))))) )),
  inference(superposition,[],[f418,f557])).
fof(f2594,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f557,f548])).
fof(f2488,plain,(
  ( ! [X24,X23,X22] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23))))))))) | s('type.integer.int',X24) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X22))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23)))))))) )),
  inference(superposition,[],[f418,f555])).
fof(f2477,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f555,f547])).
fof(f2410,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2395,f421])).
fof(f2411,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2395,f420])).
fof(f2413,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f2395,f416])).
fof(f2397,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f2392,f421])).
fof(f2398,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f2392,f420])).
fof(f2395,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(condensation,[],[f2394])).
fof(f2394,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(subsumption_resolution,[],[f2393,f1641])).
fof(f2393,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9)) )),
  inference(forward_demodulation,[],[f2342,f400])).
fof(f2342,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X9) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))))) )),
  inference(superposition,[],[f422,f553])).
fof(f2353,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f2348,f2137])).
fof(f2352,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2348,f1998])).
fof(f2351,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(backward_demodulation,[],[f2348,f1931])).
fof(f2348,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f2336,f315])).
fof(f2336,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f316,f316,f422])).
fof(f2299,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f2298,f548])).
fof(f2298,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f2290,f548])).
fof(f2290,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f548,f548,f415])).
fof(f2222,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f2214,f548])).
fof(f2214,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f1061,f548,f413])).
fof(f2139,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X14))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f419,f331])).
fof(f2137,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X10))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f418,f331])).
fof(f2128,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1060,f331])).
fof(f1998,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f548])).
fof(f1931,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f418,f547])).
fof(f1884,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f1883,f1814])).
fof(f1883,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))))) )),
  inference(forward_demodulation,[],[f1882,f1787])).
fof(f1882,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1881,f400])).
fof(f1881,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(forward_demodulation,[],[f1855,f400])).
fof(f1855,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))))) )),
  inference(superposition,[],[f401,f551])).
fof(f1880,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1844,f316])).
fof(f1844,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f401,f1045])).
fof(f1814,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1813,f1787])).
fof(f1813,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f1788,f400])).
fof(f1788,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f400,f553])).
fof(f1811,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f1787,f1569])).
fof(f1810,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(backward_demodulation,[],[f1787,f1576])).
fof(f1787,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f400,f551])).
fof(f1809,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1786,f315])).
fof(f1786,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f400,f1045])).
fof(f1664,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f421,f316])).
fof(f1663,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f421,f553])).
fof(f1662,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f421,f551])).
fof(f1661,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f421,f315])).
fof(f1659,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f421,f314])).
fof(f1657,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f389,f421])).
fof(f1658,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1052,f421])).
fof(f1655,plain,(
  ( ! [X10,X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f420,f316])).
fof(f1654,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f420,f553])).
fof(f1653,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f420,f551])).
fof(f1652,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f420,f315])).
fof(f1642,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(unit_resulting_resolution,[],[f1052,f420])).
fof(f1576,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1575,f400])).
fof(f1575,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) | s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8)))) )),
  inference(forward_demodulation,[],[f1563,f400])).
fof(f1563,plain,(
  ( ! [X8,X9] : (s('type.num.num',X9) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))))) )),
  inference(superposition,[],[f417,f553])).
fof(f1561,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f416,f553])).
fof(f1572,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1555,f1045])).
fof(f1555,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f553,f316])).
fof(f1554,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f553,f1045])).
fof(f1569,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1568,f400])).
fof(f1568,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1546,f400])).
fof(f1546,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f553,f417])).
fof(f1435,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1436,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1437,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f399,f478])).
fof(f1438,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f399,f476])).
fof(f1380,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1381,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1382,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f546,f398,f478])).
fof(f1383,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)),'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(unit_resulting_resolution,[],[f307,f398,f476])).
fof(f1316,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f569,f574,f546,f461])).
fof(f1282,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f571,f568,f575,f307,f460])).
fof(f1245,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1060,f418])).
fof(f1246,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1061,f418])).
fof(f1224,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))))))) | s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))) )),
  inference(forward_demodulation,[],[f1209,f400])).
fof(f1209,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f417,f551])).
fof(f1218,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1203,f400])).
fof(f1203,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f551,f417])).
fof(f1067,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) )),
  inference(superposition,[],[f416,f316])).
fof(f1066,plain,(
  ( ! [X6,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))) | s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f416,f551])).
fof(f1065,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X4))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f416,f315])).
fof(f1062,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f416,f314])).
fof(f1060,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f389,f416])).
fof(f1061,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1052,f416])).
fof(f1052,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f389,f1045])).
fof(f1045,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f551,f316])).
fof(f1008,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f999,f515])).
fof(f999,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f546,f307,f520])).
fof(f988,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f585,f307,f546,f519])).
fof(f957,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f949,f498])).
fof(f949,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f307,f503])).
fof(f943,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f546,f546,f502])).
fof(f915,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f906,f479])).
fof(f906,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f486])).
fof(f895,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f577,f307,f307,f482])).
fof(f891,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f886,f456])).
fof(f886,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f546,f471])).
fof(f882,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f873,f457])).
fof(f873,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f467])).
fof(f856,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f307,f463])).
fof(f812,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f599,f477])).
fof(f816,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f307,f599,f477])).
fof(f796,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f599,f475])).
fof(f798,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f599,f475])).
fof(f730,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f321,f321])).
fof(f680,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f526])).
fof(f672,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f307,f525])).
fof(f670,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f514])).
fof(f669,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f513])).
fof(f661,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f307,f509])).
fof(f657,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f307,f508])).
fof(f654,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f497])).
fof(f650,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f496])).
fof(f649,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f546,f546,f491])).
fof(f643,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f546,f490])).
fof(f629,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f546,f481])).
fof(f626,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f307,f546,f480])).
fof(f621,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f560])).
fof(f618,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f450])).
fof(f613,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f307,f447])).
fof(f610,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f307,f307,f446])).
fof(f609,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f546,f546,f445])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f523])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f522])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f546,f521])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f506])).
fof(f581,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f505])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f504])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f489])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f488])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f487])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f307,f474])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f307,f473])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f472])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f470])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f469])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f468])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f546,f466])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f546,f465])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f464])).
fof(f567,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f307,f432])).
fof(f566,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f307,f431])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f183,f200,f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) => p(s(bool,X18)))) <=> ((p(s(bool,X16)) | p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (~p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.sat.sat_27_prove_ltb')).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f304])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f199])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f302])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f299])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f299])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f157,f197,f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) | p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X16)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.sat.sat_25_prove_ltb')).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f294])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f294])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f294])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f499])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f155,f194,f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X17)) & p(s(bool,X18)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.sat.sat_23_prove_ltb')).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f288])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f288])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f193])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f286])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f284])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f153,f191,f190,f189,f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f189,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f190,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f191,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X18,X17,X16] : ((p(s(bool,X16)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X18))) & (p(s(bool,X16)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X16))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X16)))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.sat.sat_2|49|_prove_ltb')).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f281])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f281])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f188])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f189])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f278])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f190])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f275])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f191])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f272])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f272])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f455,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6',aHL_EXT)).
fof(f445,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f446,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f560,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f430,f546])).
fof(f430,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f431,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f432,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f421,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f254])).
fof(f418,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X8,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X13))) <=> s('type.integer.int',X8) = s('type.integer.int',X13))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_EQ_NEG')).
fof(f419,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f252])).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f251])).
fof(f409,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f559,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f405,f314])).
fof(f405,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|4')).
fof(f558,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f404,f314])).
fof(f404,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X8,X6] : s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|5')).
fof(f557,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f556,f314])).
fof(f556,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f403,f314])).
fof(f403,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|3')).
fof(f401,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f553,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f552,f315])).
fof(f552,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f550,f400])).
fof(f550,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(backward_demodulation,[],[f400,f319])).
fof(f333,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f331,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X16] : (~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD|48|')).
fof(f548,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f320,f314])).
fof(f320,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6)))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_MOD_REDUCE|split|1')).
fof(f319,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.BIT2')).
fof(f316,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f5689,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl41),
  inference(cnf_transformation,[],[f5689_D])).
fof(f5689_D,plain,(
  ~sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f410,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f250])).
fof(f2306,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(forward_demodulation,[],[f2297,f547])).
fof(f2297,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(trivial_inequality_removal,[],[f2291])).
fof(f2291,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int',X1))))) )),
  inference(superposition,[],[f415,f547])).
fof(f21265,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f20968,f12883])).
fof(f20968,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f555,f13033])).
fof(f21264,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f20967,f12883])).
fof(f20967,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ($spl4 | $spl41)),
  inference(superposition,[],[f13033,f555])).
fof(f21262,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21261,f12881])).
fof(f12881,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12874,f420])).
fof(f21261,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21260,f2348])).
fof(f21260,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20925,f555])).
fof(f20925,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13033,f419])).
fof(f21259,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21258,f12881])).
fof(f21258,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21257,f2348])).
fof(f21257,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21256,f547])).
fof(f21256,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20926,f555])).
fof(f20926,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1061,f13033,f2139])).
fof(f21255,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21254,f12881])).
fof(f21254,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21253,f2348])).
fof(f21253,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21252,f555])).
fof(f21252,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20927,f547])).
fof(f20927,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2400,f13033,f2139])).
fof(f21251,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21250,f12881])).
fof(f21250,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21249,f2348])).
fof(f21249,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21248,f555])).
fof(f21248,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20928,f12911])).
fof(f12911,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2348,f12883,f2353])).
fof(f20928,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12883,f13033,f2139])).
fof(f21247,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21246,f12881])).
fof(f21246,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21245,f2348])).
fof(f21245,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21244,f555])).
fof(f21244,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20929,f2128])).
fof(f20929,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1060,f13033,f2139])).
fof(f21243,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21242,f12881])).
fof(f21242,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21241,f2348])).
fof(f21241,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21240,f555])).
fof(f21240,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20930,f548])).
fof(f20930,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2413,f13033,f2139])).
fof(f21239,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21238,f12881])).
fof(f21238,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21237,f2348])).
fof(f21237,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21236,f555])).
fof(f21236,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20931,f19235])).
fof(f19235,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f19202,f13100])).
fof(f13100,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f2348,f12881,f2353])).
fof(f19202,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f12881,f2398,f5368])).
fof(f20931,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2398,f13033,f2139])).
fof(f21235,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21234,f12881])).
fof(f21234,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21233,f2348])).
fof(f21233,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21232,f555])).
fof(f21232,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20932,f13100])).
fof(f20932,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12881,f13033,f2139])).
fof(f21231,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21230,f12881])).
fof(f21230,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21229,f2348])).
fof(f21229,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21228,f555])).
fof(f21228,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20933,f4121])).
fof(f20933,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1641,f13033,f2139])).
fof(f21227,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21226,f12881])).
fof(f21226,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21225,f2348])).
fof(f21225,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21224,f555])).
fof(f21224,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20934,f19514])).
fof(f19514,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(forward_demodulation,[],[f19481,f4121])).
fof(f19481,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f1641,f2411,f5368])).
fof(f20934,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2411,f13033,f2139])).
fof(f21223,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21222,f546])).
fof(f21222,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21221,f15725])).
fof(f21221,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21220,f547])).
fof(f21220,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20935,f555])).
fof(f20935,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1061,f13033,f2295])).
fof(f21219,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21218,f546])).
fof(f21218,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21217,f15725])).
fof(f21217,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21216,f547])).
fof(f21216,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20936,f555])).
fof(f20936,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2400,f13033,f2295])).
fof(f21215,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21214,f546])).
fof(f21214,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21213,f15725])).
fof(f21213,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21212,f12911])).
fof(f21212,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20937,f555])).
fof(f20937,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12883,f13033,f2295])).
fof(f21211,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21210,f546])).
fof(f21210,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21209,f15725])).
fof(f21209,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21208,f2128])).
fof(f21208,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20938,f555])).
fof(f20938,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1060,f13033,f2295])).
fof(f21207,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21206,f546])).
fof(f21206,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21205,f15725])).
fof(f21205,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21204,f548])).
fof(f21204,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20939,f555])).
fof(f20939,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2413,f13033,f2295])).
fof(f21203,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21202,f546])).
fof(f21202,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21201,f15725])).
fof(f21201,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21200,f19235])).
fof(f21200,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20940,f555])).
fof(f20940,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2398,f13033,f2295])).
fof(f21199,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21198,f546])).
fof(f21198,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21197,f15725])).
fof(f21197,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21196,f13100])).
fof(f21196,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20941,f555])).
fof(f20941,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12881,f13033,f2295])).
fof(f21195,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21194,f546])).
fof(f21194,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21193,f15725])).
fof(f21193,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21192,f4121])).
fof(f21192,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20942,f555])).
fof(f20942,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f1641,f13033,f2295])).
fof(f21191,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21190,f546])).
fof(f21190,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21189,f15725])).
fof(f21189,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21188,f19514])).
fof(f21188,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20943,f555])).
fof(f20943,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2411,f13033,f2295])).
fof(f21187,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21186,f546])).
fof(f21186,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21185,f15725])).
fof(f21185,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20944,f555])).
fof(f20944,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13033,f2306])).
fof(f21184,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21183,f546])).
fof(f21183,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21182,f15725])).
fof(f21182,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21181,f547])).
fof(f21181,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20945,f555])).
fof(f20945,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f547,f13033,f415])).
fof(f21180,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21179,f546])).
fof(f21179,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21178,f15725])).
fof(f21178,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21177,f548])).
fof(f21177,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20946,f555])).
fof(f20946,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))) ) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f548,f13033,f415])).
fof(f21176,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21175,f546])).
fof(f21175,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21174,f15725])).
fof(f21174,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21173,f12911])).
fof(f21173,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20947,f555])).
fof(f20947,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12911,f13033,f415])).
fof(f21172,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21171,f21138])).
fof(f21138,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21137,f555])).
fof(f21137,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20963,f398])).
fof(f20963,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f12881,f13033,f413])).
fof(f21171,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20948,f555])).
fof(f20948,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13033,f13033,f415])).
fof(f21170,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21169,f546])).
fof(f21169,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21168,f15725])).
fof(f21168,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21167,f13100])).
fof(f21167,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f20949,f555])).
fof(f20949,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl3 | $spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13100,f13033,f415])).
fof(f21166,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21165,f546])).
fof(f21165,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl41)),
  inference(forward_demodulation,[],[f21164,f15725])).
fof(f21164,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21163,f555])).
fof(f21163,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20950,f399])).
fof(f20950,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f2348,f13033,f415])).
fof(f21153,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21152,f21138])).
fof(f21152,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20954,f555])).
fof(f20954,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13033,f13033,f415])).
fof(f21132,plain,(
  $false | ($spl3 | $spl4 | $spl41)),
  inference(subsumption_resolution,[],[f21131,f12881])).
fof(f21131,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f21130,f2348])).
fof(f21130,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(forward_demodulation,[],[f20966,f555])).
fof(f20966,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | ($spl4 | $spl41)),
  inference(unit_resulting_resolution,[],[f13033,f419])).
fof(f20744,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl78),
  inference(cnf_transformation,[],[f20744_D])).
fof(f20744_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5691,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl42),
  inference(cnf_transformation,[],[f5691_D])).
fof(f5691_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f20743,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20728,f332])).
fof(f20728,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4977,f321])).
fof(f20740,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl76),
  inference(cnf_transformation,[],[f20740_D])).
fof(f20740_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f20741,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20724,f332])).
fof(f20724,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4977,f730])).
fof(f20739,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20723,f332])).
fof(f20723,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4977,f730])).
fof(f20706,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f20706_D])).
fof(f20706_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f20705,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f20690,f332])).
fof(f20690,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4976,f321])).
fof(f20702,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl72),
  inference(cnf_transformation,[],[f20702_D])).
fof(f20702_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f20703,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f20686,f332])).
fof(f20686,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f4976,f730])).
fof(f20701,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f20685,f332])).
fof(f20685,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f4976,f730])).
fof(f13571,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f13571_D])).
fof(f13571_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f13573,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f13573_D])).
fof(f13573_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f13561,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2426,f730])).
fof(f2426,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK16(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK16(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f455,f321])).
fof(f13560,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2426,f730])).
fof(f13550,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl62),
  inference(cnf_transformation,[],[f13550_D])).
fof(f13550_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13546,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) | $spl58),
  inference(cnf_transformation,[],[f13546_D])).
fof(f13546_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13553,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f13553_D])).
fof(f13553_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f13544,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2424,f321])).
fof(f2424,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,X5,bool))) )),
  inference(superposition,[],[f455,f321])).
fof(f13551,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl64),
  inference(cnf_transformation,[],[f13551_D])).
fof(f13551_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13543,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2424,f730])).
fof(f13542,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2424,f730])).
fof(f13545,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f13545_D])).
fof(f13545_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13548,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f13548_D])).
fof(f13548_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f13541,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2424,f730])).
fof(f13540,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2424,f730])).
fof(f13024,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | $spl50),
  inference(cnf_transformation,[],[f13024_D])).
fof(f13024_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13031,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl54),
  inference(cnf_transformation,[],[f13031_D])).
fof(f13031_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13029,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12978,f307])).
fof(f12978,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f412,f12884])).
fof(f13026,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | $spl52),
  inference(cnf_transformation,[],[f13026_D])).
fof(f13026_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13022,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f12976,f307])).
fof(f12976,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))) | ($spl4 | $spl41)),
  inference(superposition,[],[f2791,f12884])).
fof(f12879,plain,(
  $false | ($spl0 | $spl40)),
  inference(subsumption_resolution,[],[f7021,f12877])).
fof(f12877,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f12876,f314])).
fof(f12876,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f12875,f314])).
fof(f12875,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl0),
  inference(forward_demodulation,[],[f542,f398])).
fof(f7021,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl40),
  inference(unit_resulting_resolution,[],[f307,f5690,f446])).
fof(f12878,plain,(
  $false | ($spl0 | $spl44)),
  inference(subsumption_resolution,[],[f12877,f7023])).
fof(f7023,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl44),
  inference(unit_resulting_resolution,[],[f307,f5725,f446])).
fof(f12872,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f12871,f546])).
fof(f12871,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f12833,f606])).
fof(f606,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f337])).
fof(f565,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f564,f314])).
fof(f564,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f563,f314])).
fof(f563,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))) | $spl1),
  inference(forward_demodulation,[],[f543,f398])).
fof(f12833,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(trivial_inequality_removal,[],[f12832])).
fof(f12832,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | $spl2),
  inference(superposition,[],[f2878,f562])).
fof(f562,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f561,f314])).
fof(f561,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f544,f314])).
fof(f12673,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f12673_D])).
fof(f12673_D,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12674,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(subsumption_resolution,[],[f12599,f332])).
fof(f12599,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s('type.integer.int',X15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))))) )),
  inference(superposition,[],[f2871,f730])).
fof(f12672,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13)))))) )),
  inference(subsumption_resolution,[],[f12598,f332])).
fof(f12598,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X13))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2871,f730])).
fof(f8669,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl46),
  inference(cnf_transformation,[],[f8669_D])).
fof(f8669_D,plain,(
  ( ! [X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f8670,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8659,f332])).
fof(f8659,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f2868,f730])).
fof(f8668,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f8658,f332])).
fof(f8658,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2868,f730])).
fof(f6821,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5729,f5741])).
fof(f5741,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f5691,f409])).
fof(f5729,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f336])).
fof(f6820,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5732,f307])).
fof(f5732,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f599,f5691,f409])).
fof(f6819,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5734,f307])).
fof(f5734,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f600,f5691,f409])).
fof(f600,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f332])).
fof(f6818,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5738,f5691])).
fof(f5738,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f600,f5691,f409])).
fof(f6817,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5742,f307])).
fof(f5742,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f599,f5691,f409])).
fof(f6816,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5744,f307])).
fof(f5744,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f600,f5691,f409])).
fof(f6815,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5748,f5691])).
fof(f5748,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f600,f5691,f409])).
fof(f6814,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5755,f307])).
fof(f5755,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f618,f5691,f446])).
fof(f6813,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5756,f5691])).
fof(f5756,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3461,f5691,f446])).
fof(f6812,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5762,f307])).
fof(f5762,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f613,f5691,f446])).
fof(f6811,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5765,f5691])).
fof(f5765,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3460,f5691,f446])).
fof(f6810,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5782,f307])).
fof(f5782,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f798,f546,f5691,f459])).
fof(f6809,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5783,f5691])).
fof(f5783,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f803,f546,f5691,f459])).
fof(f803,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f606,f476])).
fof(f6808,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5784,f5691])).
fof(f5784,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1723,f546,f5691,f459])).
fof(f1723,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1672,f476])).
fof(f1672,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f1520])).
fof(f1520,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) ) | $spl1),
  inference(forward_demodulation,[],[f1485,f398])).
fof(f1485,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X5)))) ) | $spl1),
  inference(superposition,[],[f1415,f399])).
fof(f1415,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3)))) ) | $spl1),
  inference(forward_demodulation,[],[f1384,f606])).
fof(f1384,plain,(
  ( ! [X2,X3] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int',X3))) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3)))))) ) | $spl1),
  inference(superposition,[],[f398,f1038])).
fof(f1038,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f1035,f546])).
fof(f1035,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f781,f606])).
fof(f6807,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5785,f5691])).
fof(f5785,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2232,f546,f5691,f459])).
fof(f2232,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2186,f476])).
fof(f2186,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1672,f2065])).
fof(f2065,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int',X9))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(forward_demodulation,[],[f2050,f398])).
fof(f2050,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X8))),s('type.integer.int',X9)))) ) | $spl1),
  inference(superposition,[],[f1523,f399])).
fof(f1523,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f1493,f307])).
fof(f1493,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))))) | s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int',X10),s('type.integer.int',X11)))) ) | $spl1),
  inference(superposition,[],[f336,f1415])).
fof(f6806,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5786,f5691])).
fof(f5786,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3260,f546,f5691,f459])).
fof(f3260,plain,(
  sP5('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2229,f476])).
fof(f2229,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f2186,f2065])).
fof(f6805,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5788,f5691])).
fof(f5788,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3493,f546,f5691,f459])).
fof(f3493,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f606,f2299,f476])).
fof(f6804,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5820,f307])).
fof(f5820,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f803,f546,f5691,f459])).
fof(f6803,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5821,f307])).
fof(f5821,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1723,f546,f5691,f459])).
fof(f6802,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5822,f307])).
fof(f5822,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2232,f546,f5691,f459])).
fof(f6801,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5823,f307])).
fof(f5823,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3260,f546,f5691,f459])).
fof(f6800,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5824,f307])).
fof(f5824,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1316,f546,f5691,f459])).
fof(f6799,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5825,f5691])).
fof(f5825,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3493,f546,f5691,f459])).
fof(f6798,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5867,f6085])).
fof(f6085,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f474])).
fof(f5867,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f568,f5691,f460])).
fof(f6797,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5868,f6085])).
fof(f5868,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f568,f5691,f460])).
fof(f817,plain,(
  ~sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f600,f477])).
fof(f6796,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5869,f6085])).
fof(f5869,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f568,f5691,f460])).
fof(f6795,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5870,f6085])).
fof(f5870,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f568,f5691,f460])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f469])).
fof(f6794,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5871,f6085])).
fof(f5871,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3506,f588,f568,f5691,f460])).
fof(f3506,plain,(
  ~sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f2299,f477])).
fof(f6793,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5872,f6085])).
fof(f5872,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f571,f3478,f5691,f460])).
fof(f6792,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5873,f6085])).
fof(f5873,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f572,f3478,f5691,f460])).
fof(f6791,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5874,f6085])).
fof(f5874,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f569,f5691,f460])).
fof(f6790,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5875,f6085])).
fof(f5875,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f572,f569,f5691,f460])).
fof(f6789,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5881,f6085])).
fof(f5881,plain,(
  ~sP2('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f569,f5691,f460])).
fof(f6788,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5882,f6085])).
fof(f5882,plain,(
  ~sP2('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f571,f569,f5691,f460])).
fof(f6787,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5885,f6085])).
fof(f5885,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f586,f5691,f460])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f465])).
fof(f6786,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5889,f6085])).
fof(f5889,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3506,f588,f586,f5691,f460])).
fof(f6785,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5890,f6085])).
fof(f5890,plain,(
  ~sP2('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f586,f5691,f460])).
fof(f6784,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5909,f6045])).
fof(f6045,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f464])).
fof(f5909,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f574,f5691,f460])).
fof(f6783,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5910,f6045])).
fof(f5910,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f574,f5691,f460])).
fof(f6782,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5912,f6045])).
fof(f5912,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f571,f574,f5691,f460])).
fof(f6781,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5913,f6045])).
fof(f5913,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f574,f5691,f460])).
fof(f6780,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5914,f6045])).
fof(f5914,plain,(
  ~sP4('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f572,f574,f5691,f460])).
fof(f6779,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5915,f6045])).
fof(f5915,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f574,f5691,f460])).
fof(f6778,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5921,f6079])).
fof(f6079,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f468])).
fof(f5921,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f568,f574,f5691,f460])).
fof(f6777,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5922,f6079])).
fof(f5922,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f568,f574,f5691,f460])).
fof(f6776,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5929,f6079])).
fof(f5929,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f3478,f574,f5691,f460])).
fof(f6775,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5935,f6079])).
fof(f5935,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f569,f574,f5691,f460])).
fof(f6774,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5936,f6079])).
fof(f5936,plain,(
  ~sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f569,f574,f5691,f460])).
fof(f6773,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5943,f6079])).
fof(f5943,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f586,f574,f5691,f460])).
fof(f6772,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5950,f6045])).
fof(f5950,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f575,f5691,f460])).
fof(f6771,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5951,f6045])).
fof(f5951,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f575,f5691,f460])).
fof(f6770,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5952,f6045])).
fof(f5952,plain,(
  ~sP4('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f575,f5691,f460])).
fof(f6769,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5953,f6045])).
fof(f5953,plain,(
  ~sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f575,f5691,f460])).
fof(f6768,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5954,f6045])).
fof(f5954,plain,(
  ~sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3506,f588,f575,f5691,f460])).
fof(f6767,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5955,f6079])).
fof(f5955,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f568,f575,f5691,f460])).
fof(f6766,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5956,f6079])).
fof(f5956,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f568,f575,f5691,f460])).
fof(f6765,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5958,f6079])).
fof(f5958,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3506,f568,f575,f5691,f460])).
fof(f6764,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5961,f6130])).
fof(f6130,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f599,f5691,f477])).
fof(f5961,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f572,f568,f575,f5691,f460])).
fof(f6763,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5963,f6132])).
fof(f6132,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f600,f5691,f477])).
fof(f5963,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f568,f575,f5691,f460])).
fof(f6762,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5965,f6079])).
fof(f5965,plain,(
  ~sP3('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f816,f569,f575,f5691,f460])).
fof(f6761,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5966,f6130])).
fof(f5966,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f572,f569,f575,f5691,f460])).
fof(f6760,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5968,f816])).
fof(f5968,plain,(
  sP5('T_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f571,f569,f575,f5691,f460])).
fof(f6759,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5970,f6079])).
fof(f5970,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f586,f575,f5691,f460])).
fof(f6758,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5971,f6079])).
fof(f5971,plain,(
  ~sP3('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3506,f586,f575,f5691,f460])).
fof(f6757,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5972,f6132])).
fof(f5972,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f586,f575,f5691,f460])).
fof(f6756,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5974,f817])).
fof(f5974,plain,(
  sP5('T_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f571,f586,f575,f5691,f460])).
fof(f6755,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5979,f6130])).
fof(f5979,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f572,f587,f575,f5691,f460])).
fof(f587,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f466])).
fof(f6754,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5980,f6132])).
fof(f5980,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f587,f575,f5691,f460])).
fof(f6753,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5995,f5691])).
fof(f5995,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f586,f546,f5691,f463])).
fof(f6752,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5996,f5691])).
fof(f5996,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1913,f546,f5691,f463])).
fof(f1913,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f1725,f458])).
fof(f1725,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f1672,f476])).
fof(f6751,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5997,f5691])).
fof(f5997,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2454,f546,f5691,f463])).
fof(f2454,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f2234,f458])).
fof(f2234,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2186,f476])).
fof(f6750,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f5998,f5691])).
fof(f5998,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3437,f546,f5691,f463])).
fof(f3437,plain,(
  sP4('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f3262,f458])).
fof(f3262,plain,(
  sP5('T_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f307,f2229,f476])).
fof(f6749,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f5999,f5691])).
fof(f5999,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f570,f546,f5691,f463])).
fof(f6748,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6000,f5691])).
fof(f6000,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5691,f463])).
fof(f6747,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6005,f6045])).
fof(f6005,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f463])).
fof(f6746,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6016,f6045])).
fof(f6016,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f463])).
fof(f6745,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6018,f5691])).
fof(f6018,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f569,f546,f5691,f463])).
fof(f6744,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6019,f5691])).
fof(f6019,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f586,f546,f5691,f463])).
fof(f6743,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6020,f307])).
fof(f6020,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1913,f546,f5691,f463])).
fof(f6742,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6021,f307])).
fof(f6021,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2454,f546,f5691,f463])).
fof(f6741,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6022,f307])).
fof(f6022,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3437,f546,f5691,f463])).
fof(f6740,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6024,f5691])).
fof(f6024,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f587,f546,f5691,f463])).
fof(f6739,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6032,f6045])).
fof(f6032,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f463])).
fof(f6738,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6043,f6045])).
fof(f6043,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f463])).
fof(f6737,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6048,f5691])).
fof(f6048,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5691,f467])).
fof(f6736,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6049,f5691])).
fof(f6049,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f5691,f467])).
fof(f6735,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6050,f5691])).
fof(f6050,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5691,f467])).
fof(f589,plain,(
  ( ! [X0,X1] : (sP3('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f470])).
fof(f6734,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6054,f6079])).
fof(f6054,plain,(
  ( ! [X0] : (~sP3(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f467])).
fof(f6733,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6061,f6079])).
fof(f6061,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f467])).
fof(f6732,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6063,f5691])).
fof(f6063,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f5691,f467])).
fof(f6731,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6064,f5691])).
fof(f6064,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f588,f546,f5691,f467])).
fof(f6730,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6066,f5691])).
fof(f6066,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f589,f546,f5691,f467])).
fof(f6729,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6070,f6079])).
fof(f6070,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f467])).
fof(f6728,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6077,f6079])).
fof(f6077,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f467])).
fof(f6726,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6094,f606])).
fof(f6094,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f826,f5691,f476])).
fof(f826,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f606,f478])).
fof(f6725,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6095,f1672])).
fof(f6095,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1724,f5691,f476])).
fof(f1724,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1672,f478])).
fof(f6724,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6096,f2186])).
fof(f6096,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2233,f5691,f476])).
fof(f2233,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2186,f478])).
fof(f6723,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6097,f2229])).
fof(f6097,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3261,f5691,f476])).
fof(f3261,plain,(
  ~sP5('F_0','F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2229,f478])).
fof(f6722,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6098,f606])).
fof(f6098,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f825,f5691,f476])).
fof(f825,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f606,f478])).
fof(f6721,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6099,f1672])).
fof(f6099,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1722,f5691,f476])).
fof(f1722,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f1672,f478])).
fof(f6720,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6100,f2186])).
fof(f6100,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f2231,f5691,f476])).
fof(f2231,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2186,f478])).
fof(f6719,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6101,f2229])).
fof(f6101,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3259,f5691,f476])).
fof(f3259,plain,(
  ~sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f2229,f478])).
fof(f6718,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6102,f398])).
fof(f6102,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1380,f5691,f476])).
fof(f6717,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6103,f398])).
fof(f6103,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1382,f5691,f476])).
fof(f6716,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6104,f399])).
fof(f6104,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1435,f5691,f476])).
fof(f6715,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6105,f399])).
fof(f6105,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_divides_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f1437,f5691,f476])).
fof(f6714,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6108,f599])).
fof(f6108,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f796,f5691,f477])).
fof(f6713,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6109,f599])).
fof(f6109,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f798,f5691,f477])).
fof(f6712,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6110,f600])).
fof(f6110,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f799,f5691,f477])).
fof(f799,plain,(
  sP5('F_0','T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f600,f475])).
fof(f6711,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6111,f600])).
fof(f6111,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f797,f5691,f477])).
fof(f797,plain,(
  sP5('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f546,f600,f475])).
fof(f6710,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6127,f599])).
fof(f6127,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1274,f5691,f477])).
fof(f1274,plain,(
  sP5('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f573,f587,f576,f307,f460])).
fof(f6709,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6143,f6280])).
fof(f6280,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f5691,f5691,f492])).
fof(f6143,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f480])).
fof(f6708,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6150,f6280])).
fof(f6150,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f481])).
fof(f6707,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6153,f5691])).
fof(f6153,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f579,f629,f5691,f482])).
fof(f6706,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6154,f5691])).
fof(f6154,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f577,f629,f5691,f482])).
fof(f6705,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6156,f5691])).
fof(f6156,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f629,f5691,f482])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f488])).
fof(f6704,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6160,f6226])).
fof(f6226,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f487])).
fof(f6160,plain,(
  ~sP6('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f629,f307,f5691,f482])).
fof(f6703,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6168,f6226])).
fof(f6168,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f482])).
fof(f6702,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6170,f6226])).
fof(f6170,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f482])).
fof(f6701,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6173,f5691])).
fof(f6173,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f626,f5691,f482])).
fof(f6700,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6174,f5691])).
fof(f6174,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f577,f626,f5691,f482])).
fof(f6699,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6176,f5691])).
fof(f6176,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f626,f5691,f482])).
fof(f591,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f489])).
fof(f6698,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6181,f6226])).
fof(f6181,plain,(
  ~sP6('T_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f5691,f482])).
fof(f6697,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6189,f6226])).
fof(f6189,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f482])).
fof(f6696,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6191,f6226])).
fof(f6191,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f482])).
fof(f6695,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6195,f5691])).
fof(f6195,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5691,f486])).
fof(f6694,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6196,f5691])).
fof(f6196,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f5691,f486])).
fof(f6693,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6197,f5691])).
fof(f6197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f546,f5691,f486])).
fof(f6692,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6201,f6226])).
fof(f6201,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f486])).
fof(f6691,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6208,f6226])).
fof(f6208,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f486])).
fof(f6690,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6210,f5691])).
fof(f6210,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f5691,f486])).
fof(f6689,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6211,f5691])).
fof(f6211,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f590,f546,f5691,f486])).
fof(f6688,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6213,f5691])).
fof(f6213,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f591,f546,f5691,f486])).
fof(f6687,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6217,f6226])).
fof(f6217,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f307,f5691,f486])).
fof(f6686,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6224,f6226])).
fof(f6224,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f5691,f486])).
fof(f6685,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6229,f5691])).
fof(f6229,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f307,f5691,f492])).
fof(f6684,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6231,f5691])).
fof(f6231,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f492])).
fof(f6683,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6232,f307])).
fof(f6232,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f5691,f5691,f492])).
fof(f6682,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6233,f307])).
fof(f6233,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f307,f5691,f492])).
fof(f6681,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6239,f307])).
fof(f6239,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f492])).
fof(f6680,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6240,f307])).
fof(f6240,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f492])).
fof(f6679,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6246,f5691])).
fof(f6246,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f5691,f492])).
fof(f6678,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6248,f5691])).
fof(f6248,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f492])).
fof(f6677,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6250,f307])).
fof(f6250,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f5691,f5691,f492])).
fof(f6676,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6251,f307])).
fof(f6251,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f5691,f492])).
fof(f6675,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6257,f307])).
fof(f6257,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f492])).
fof(f6674,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6258,f307])).
fof(f6258,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f492])).
fof(f6673,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6265,f307])).
fof(f6265,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f307,f5691,f492])).
fof(f6672,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6266,f5691])).
fof(f6266,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f492])).
fof(f6671,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6268,f307])).
fof(f6268,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f5691,f5691,f492])).
fof(f6670,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6270,f307])).
fof(f6270,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f307,f5691,f492])).
fof(f6669,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6276,f5691])).
fof(f6276,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f5691,f492])).
fof(f6668,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6277,f307])).
fof(f6277,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f915,f5691,f5691,f492])).
fof(f6667,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6283,f5691])).
fof(f6283,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f708,f546,f5691,f495])).
fof(f708,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f590,f546,f532])).
fof(f6666,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6284,f5691])).
fof(f6284,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f649,f546,f5691,f495])).
fof(f6665,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6285,f5691])).
fof(f6285,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f710,f546,f5691,f495])).
fof(f710,plain,(
  ( ! [X0] : (sP7('F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f591,f546,f532])).
fof(f6664,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6298,f5691])).
fof(f6298,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f643,f546,f5691,f495])).
fof(f6663,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6299,f5691])).
fof(f6299,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f708,f546,f5691,f495])).
fof(f6662,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6301,f5691])).
fof(f6301,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f710,f546,f5691,f495])).
fof(f6661,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6318,f6594])).
fof(f6594,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f582,f5691,f531])).
fof(f6318,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f496])).
fof(f6660,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6324,f6594])).
fof(f6324,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f497])).
fof(f6659,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6333,f6358])).
fof(f6358,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f505])).
fof(f6333,plain,(
  ~sP8('F_0','T_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f654,f546,f5691,f500])).
fof(f6658,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6341,f6358])).
fof(f6341,plain,(
  ~sP8('F_0','F_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f650,f546,f5691,f501])).
fof(f6657,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6354,f6358])).
fof(f6354,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f546,f5691,f503])).
fof(f6656,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6362,f307])).
fof(f6362,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f957,f5691,f508])).
fof(f6655,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6367,f307])).
fof(f6367,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f508])).
fof(f6654,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6368,f5691])).
fof(f6368,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f654,f5691,f508])).
fof(f6653,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6376,f307])).
fof(f6376,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f957,f5691,f509])).
fof(f6652,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6381,f5691])).
fof(f6381,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f509])).
fof(f6651,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6382,f307])).
fof(f6382,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f654,f5691,f509])).
fof(f6650,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6403,f6584])).
fof(f6584,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f585,f5691,f530])).
fof(f6403,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691,f514])).
fof(f6649,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6406,f5691])).
fof(f6406,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f585,f670,f5691,f517])).
fof(f6648,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6407,f5691])).
fof(f6407,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f584,f670,f5691,f517])).
fof(f6647,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6409,f5691])).
fof(f6409,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f593,f670,f5691,f517])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f521])).
fof(f6646,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6413,f6508])).
fof(f6508,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f522])).
fof(f6413,plain,(
  ~sP10('F_0','T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f670,f307,f5691,f517])).
fof(f6645,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6421,f6508])).
fof(f6421,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f5691,f517])).
fof(f6644,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6423,f6508])).
fof(f6423,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5691,f5691,f517])).
fof(f6643,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6436,f6508])).
fof(f6436,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f5691,f517])).
fof(f6642,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6437,f6508])).
fof(f6437,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5691,f5691,f517])).
fof(f6641,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6454,f6508])).
fof(f6454,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f669,f546,f5691,f518])).
fof(f6640,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6468,f6508])).
fof(f6468,plain,(
  ~sP10('F_0','F_0','F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f669,f546,f5691,f519])).
fof(f6639,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6485,f5691])).
fof(f6485,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f593,f546,f5691,f520])).
fof(f6638,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6486,f5691])).
fof(f6486,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f585,f546,f5691,f520])).
fof(f6637,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6487,f5691])).
fof(f6487,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f594,f546,f5691,f520])).
fof(f594,plain,(
  ( ! [X0,X1] : (sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f565,f523])).
fof(f6636,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6488,f6508])).
fof(f6488,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5691,f520])).
fof(f6635,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6490,f6508])).
fof(f6490,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f546,f5691,f520])).
fof(f6634,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6500,f5691])).
fof(f6500,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f583,f546,f5691,f520])).
fof(f6633,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6501,f5691])).
fof(f6501,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f593,f546,f5691,f520])).
fof(f6632,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6503,f5691])).
fof(f6503,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f594,f546,f5691,f520])).
fof(f6631,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6504,f6508])).
fof(f6504,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f307,f546,f5691,f520])).
fof(f6630,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6506,f6508])).
fof(f6506,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f546,f5691,f520])).
fof(f6629,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6523,f307])).
fof(f6523,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f1008,f5691,f526])).
fof(f6628,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6528,f5691])).
fof(f6528,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f526])).
fof(f6627,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6529,f307])).
fof(f6529,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5691,f526])).
fof(f6626,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6535,f307])).
fof(f6535,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f672,f546,f5691,f527])).
fof(f6625,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6536,f307])).
fof(f6536,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f692,f546,f5691,f527])).
fof(f692,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f593,f307,f530])).
fof(f6624,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6538,f5691])).
fof(f6538,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f989,f546,f5691,f527])).
fof(f989,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f594,f307,f546,f519])).
fof(f6623,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6539,f5691])).
fof(f6539,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3586,f546,f5691,f527])).
fof(f3586,plain,(
  sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'F_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(unit_resulting_resolution,[],[f546,f594,f2299,f519])).
fof(f6622,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6540,f5691])).
fof(f6540,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3612,f546,f5691,f527])).
fof(f3612,plain,(
  ( ! [X0] : (sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f593,f2299,f530])).
fof(f6621,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6556,f5691])).
fof(f6556,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f692,f546,f5691,f527])).
fof(f6620,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6557,f307])).
fof(f6557,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(unit_resulting_resolution,[],[f988,f546,f5691,f527])).
fof(f6619,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6558,f307])).
fof(f6558,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f989,f546,f5691,f527])).
fof(f6618,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6559,f5691])).
fof(f6559,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3586,f546,f5691,f527])).
fof(f6617,plain,(
  $false | ($spl1 | $spl42)),
  inference(subsumption_resolution,[],[f6560,f5691])).
fof(f6560,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f3612,f546,f5691,f527])).
fof(f6616,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6575,f583])).
fof(f6575,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f530])).
fof(f6615,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6576,f584])).
fof(f6576,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f670,f5691,f530])).
fof(f6614,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6587,f580])).
fof(f6587,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f531])).
fof(f6613,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f6588,f581])).
fof(f6588,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl42),
  inference(unit_resulting_resolution,[],[f654,f5691,f531])).
fof(f6595,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f582,f5691,f531])).
fof(f6593,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f654,f581,f5691,f531])).
fof(f6591,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f580,f5691,f531])).
fof(f6585,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f585,f5691,f530])).
fof(f6583,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f5691,f530])).
fof(f6581,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f670,f593,f5691,f530])).
fof(f6579,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f5691,f530])).
fof(f6574,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f672,f546,f5691,f5691,f527])).
fof(f6565,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f988,f546,f307,f5691,f527])).
fof(f6554,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f672,f546,f5691,f5691,f527])).
fof(f6545,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f672,f546,f307,f5691,f527])).
fof(f6534,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f5691,f526])).
fof(f6532,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f307,f5691,f526])).
fof(f6527,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f5691,f5691,f526])).
fof(f6525,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f1008,f307,f5691,f526])).
fof(f6520,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f546,f5691,f525])).
fof(f6515,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f546,f546,f5691,f524])).
fof(f6509,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f999,f5691,f522])).
fof(f6507,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f583,f5691,f546,f5691,f520])).
fof(f6505,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f583,f307,f546,f5691,f520])).
fof(f6491,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f583,f5691,f546,f5691,f520])).
fof(f6489,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f585,f307,f546,f5691,f520])).
fof(f6474,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f585,f546,f5691,f519])).
fof(f6471,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f546,f5691,f519])).
fof(f6460,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f585,f546,f5691,f518])).
fof(f6457,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f546,f5691,f518])).
fof(f6446,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f585,f5691,f5691,f517])).
fof(f6444,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f670,f584,f5691,f5691,f517])).
fof(f6442,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f670,f593,f5691,f5691,f517])).
fof(f6440,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f669,f583,f5691,f5691,f517])).
fof(f6424,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f585,f670,f5691,f5691,f517])).
fof(f6422,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f585,f669,f5691,f5691,f517])).
fof(f6414,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f585,f670,f307,f5691,f517])).
fof(f6404,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f988,f546,f5691,f514])).
fof(f6395,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f943,f546,f546,f5691,f510])).
fof(f6387,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f5691,f509])).
fof(f6385,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f654,f307,f5691,f509])).
fof(f6380,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f5691,f509])).
fof(f6378,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f957,f307,f5691,f509])).
fof(f6373,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f5691,f508])).
fof(f6371,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f307,f5691,f508])).
fof(f6366,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f5691,f5691,f508])).
fof(f6364,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f957,f307,f5691,f508])).
fof(f6359,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f949,f5691,f505])).
fof(f6357,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f949,f5691,f504])).
fof(f6355,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f582,f546,f546,f5691,f503])).
fof(f6346,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f582,f546,f5691,f501])).
fof(f6344,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f650,f580,f546,f5691,f501])).
fof(f6334,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f582,f654,f546,f5691,f500])).
fof(f6325,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f943,f546,f5691,f497])).
fof(f6319,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f943,f546,f5691,f496])).
fof(f6313,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f643,f546,f5691,f5691,f495])).
fof(f6306,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f643,f546,f307,f5691,f495])).
fof(f6297,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f643,f546,f5691,f5691,f495])).
fof(f6290,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f546,f307,f5691,f495])).
fof(f6281,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f5691,f492])).
fof(f6279,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f629,f307,f5691,f5691,f492])).
fof(f6274,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f307,f5691,f492])).
fof(f6272,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f915,f307,f307,f5691,f492])).
fof(f6263,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f5691,f492])).
fof(f6261,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f5691,f5691,f492])).
fof(f6256,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f307,f5691,f492])).
fof(f6254,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f307,f5691,f492])).
fof(f6245,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f5691,f5691,f5691,f492])).
fof(f6243,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f626,f307,f5691,f5691,f492])).
fof(f6238,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f629,f5691,f307,f5691,f492])).
fof(f6236,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f629,f307,f307,f5691,f492])).
fof(f6227,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f906,f5691,f487])).
fof(f6225,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f5691,f5691,f486])).
fof(f6218,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f307,f5691,f486])).
fof(f6209,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f546,f5691,f5691,f486])).
fof(f6202,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f546,f307,f5691,f486])).
fof(f6192,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f629,f5691,f5691,f482])).
fof(f6190,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f626,f5691,f5691,f482])).
fof(f6182,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f626,f307,f5691,f482])).
fof(f6171,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f629,f5691,f5691,f482])).
fof(f6169,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f578,f626,f5691,f5691,f482])).
fof(f6161,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f579,f629,f307,f5691,f482])).
fof(f6151,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f649,f546,f5691,f481])).
fof(f6144,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f643,f546,f5691,f480])).
fof(f6137,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f797,f600,f5691,f477])).
fof(f6135,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f796,f599,f5691,f477])).
fof(f6133,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f799,f600,f5691,f477])).
fof(f6131,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f798,f599,f5691,f477])).
fof(f6612,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6093])).
fof(f6093,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(unit_resulting_resolution,[],[f891,f5691,f476])).
fof(f6611,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f6092])).
fof(f6092,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f882,f5691,f476])).
fof(f6086,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f886,f5691,f474])).
fof(f6084,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f886,f5691,f473])).
fof(f6082,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f886,f5691,f472])).
fof(f6080,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f873,f5691,f468])).
fof(f6078,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f5691,f5691,f467])).
fof(f6071,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f307,f5691,f467])).
fof(f6062,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f572,f546,f5691,f5691,f467])).
fof(f6055,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f573,f546,f307,f5691,f467])).
fof(f6046,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f856,f5691,f464])).
fof(f6044,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f569,f546,f5691,f5691,f463])).
fof(f6033,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f1913,f546,f307,f5691,f463])).
fof(f6017,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f569,f546,f5691,f5691,f463])).
fof(f6006,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f570,f546,f307,f5691,f463])).
fof(f5975,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f586,f575,f5691,f460])).
fof(f5973,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f586,f575,f5691,f460])).
fof(f5969,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f569,f575,f5691,f460])).
fof(f5967,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f569,f575,f5691,f460])).
fof(f5964,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f568,f575,f5691,f460])).
fof(f5962,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f568,f575,f5691,f460])).
fof(f5960,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f568,f575,f5691,f460])).
fof(f5947,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f571,f586,f574,f5691,f460])).
fof(f5945,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f586,f574,f5691,f460])).
fof(f5940,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f569,f574,f5691,f460])).
fof(f5938,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f569,f574,f5691,f460])).
fof(f5933,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f572,f3478,f574,f5691,f460])).
fof(f5931,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f3464,f571,f3478,f574,f5691,f460])).
fof(f5928,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f817,f588,f568,f574,f5691,f460])).
fof(f5926,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f572,f568,f574,f5691,f460])).
fof(f5924,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f816,f571,f568,f574,f5691,f460])).
fof(f5866,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f798,f546,f5691,f5691,f459])).
fof(f5842,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f803,f546,f307,f5691,f459])).
fof(f5818,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f798,f546,f5691,f5691,f459])).
fof(f5794,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f798,f546,f307,f5691,f459])).
fof(f5781,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f609,f546,f5691,f450])).
fof(f5775,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f609,f546,f5691,f447])).
fof(f5769,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f613,f5691,f5691,f446])).
fof(f5767,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f613,f307,f5691,f446])).
fof(f5761,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f613,f5691,f5691,f446])).
fof(f5759,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f618,f307,f5691,f446])).
fof(f5753,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f621,f5691,f432])).
fof(f5751,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f621,f5691,f431])).
fof(f5749,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5691,f600,f5691,f409])).
fof(f5747,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f599,f5691,f409])).
fof(f5745,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f600,f5691,f409])).
fof(f5743,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f599,f5691,f409])).
fof(f5739,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f5691,f600,f5691,f409])).
fof(f5737,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f5691,f599,f5691,f409])).
fof(f5735,plain,(
  $false | ($spl1 | $spl42)),
  inference(unit_resulting_resolution,[],[f307,f600,f5691,f409])).
fof(f5733,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f307,f599,f5691,f409])).
fof(f5728,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f599,f5691,f333])).
fof(f5726,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f546,f5691])).
fof(f5725,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | $spl44),
  inference(cnf_transformation,[],[f5725_D])).
fof(f5725_D,plain,(
  sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5723,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0')) ) | ($spl1 | $spl8)),
  inference(subsumption_resolution,[],[f5718,f332])).
fof(f5718,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))),'T_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl8)),
  inference(superposition,[],[f4923,f1038])).
fof(f5690,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | $spl40),
  inference(cnf_transformation,[],[f5690_D])).
fof(f5690_D,plain,(
  sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5688,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13)))))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f5683,f332])).
fof(f5683,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK13))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f4919,f1038])).
fof(f5365,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5354,f5356,f5362,f5364])).
fof(f5364,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5244,f307])).
fof(f5244,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'T_0'))) ) | $spl6),
  inference(condensation,[],[f5231])).
fof(f5231,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2219])).
fof(f5362,plain,(
  ( ! [X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5229,f307])).
fof(f5229,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2226])).
fof(f5356,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5217,f307])).
fof(f5217,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X4))),s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2708])).
fof(f5354,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5212,f307])).
fof(f5212,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2791])).
fof(f5352,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5211,f307])).
fof(f5211,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2788])).
fof(f5350,plain,(
  ( ! [X10,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5210,f307])).
fof(f5210,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2717])).
fof(f5348,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5209,f307])).
fof(f5209,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))),s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f2714])).
fof(f5346,plain,(
  ( ! [X92,X90,X91,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(subsumption_resolution,[],[f5208,f307])).
fof(f5208,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X89))),s('type.integer.int',X90)),X91,X92)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1469])).
fof(f5344,plain,(
  ( ! [X88,X87,X85,X86] : (~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(subsumption_resolution,[],[f5207,f307])).
fof(f5207,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X85))),s('type.integer.int',X86)),X87,X88)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1468])).
fof(f5342,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(subsumption_resolution,[],[f5206,f307])).
fof(f5206,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X81))),s('type.integer.int',X82)),X84,X83)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1467])).
fof(f5340,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(subsumption_resolution,[],[f5205,f307])).
fof(f5205,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X79)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X77))),s('type.integer.int',X78)),X79,X80)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1466])).
fof(f5338,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(subsumption_resolution,[],[f5204,f307])).
fof(f5204,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | sP11('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X73))),s('type.integer.int',X74)),X76,X75)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1465])).
fof(f5336,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5203,f307])).
fof(f5203,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP10(X72,X71,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X69))),s('type.integer.int',X70)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1464])).
fof(f5334,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(subsumption_resolution,[],[f5202,f307])).
fof(f5202,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X65))),s('type.integer.int',X66)),X68,X67)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1463])).
fof(f5332,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(subsumption_resolution,[],[f5201,f307])).
fof(f5201,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X61))),s('type.integer.int',X62)),X63,X64)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1462])).
fof(f5330,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(subsumption_resolution,[],[f5200,f307])).
fof(f5200,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | sP9('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X58)),X60,X59)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1461])).
fof(f5328,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(subsumption_resolution,[],[f5199,f307])).
fof(f5199,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X53))),s('type.integer.int',X54)),X55,X56)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1460])).
fof(f5326,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5198,f307])).
fof(f5198,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X49))),s('type.integer.int',X50)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1459])).
fof(f5324,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(subsumption_resolution,[],[f5197,f307])).
fof(f5197,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP6(X47,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X45))),s('type.integer.int',X46)),X48)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1458])).
fof(f5322,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(subsumption_resolution,[],[f5196,f307])).
fof(f5196,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP3(X43,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X41))),s('type.integer.int',X42)),X44)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1457])).
fof(f5320,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(subsumption_resolution,[],[f5195,f307])).
fof(f5195,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP4(X39,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X37))),s('type.integer.int',X38)),X40)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1456])).
fof(f5318,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(subsumption_resolution,[],[f5194,f307])).
fof(f5194,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33))),s('type.integer.int',X34)),X36)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1455])).
fof(f5316,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5193,f307])).
fof(f5193,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30))),s('type.integer.int',X31)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1454])).
fof(f5314,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(subsumption_resolution,[],[f5192,f307])).
fof(f5192,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X27))),s('type.integer.int',X28)),X29)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1453])).
fof(f5312,plain,(
  ( ! [X26,X24,X25] : (sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(subsumption_resolution,[],[f5191,f307])).
fof(f5191,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X24))),s('type.integer.int',X25)),X26)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1452])).
fof(f5310,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(subsumption_resolution,[],[f5190,f307])).
fof(f5190,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | sP0(X23,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X21))),s('type.integer.int',X22)))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1451])).
fof(f5308,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(subsumption_resolution,[],[f5189,f307])).
fof(f5189,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90) | sP9('const.integer.int_divides_2'(s('type.integer.int',X87),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X88)))),X89,X90)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1412])).
fof(f5306,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(subsumption_resolution,[],[f5188,f307])).
fof(f5188,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86) | sP11('const.integer.int_divides_2'(s('type.integer.int',X83),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X84)))),X85,X86)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1411])).
fof(f5304,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(subsumption_resolution,[],[f5187,f307])).
fof(f5187,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP11('const.integer.int_divides_2'(s('type.integer.int',X79),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X80)))),X82,X81)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1410])).
fof(f5302,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(subsumption_resolution,[],[f5186,f307])).
fof(f5186,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X75),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X76)))),X77,X78)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1409])).
fof(f5300,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(subsumption_resolution,[],[f5185,f307])).
fof(f5185,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | sP11('const.integer.int_divides_2'(s('type.integer.int',X71),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X72)))),X74,X73)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1408])).
fof(f5298,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5184,f307])).
fof(f5184,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP10(X70,X69,'const.integer.int_divides_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X68)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1407])).
fof(f5296,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(subsumption_resolution,[],[f5183,f307])).
fof(f5183,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9('const.integer.int_divides_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X64)))),X66,X65)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1406])).
fof(f5294,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(subsumption_resolution,[],[f5182,f307])).
fof(f5182,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X60)))),X61,X62)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1405])).
fof(f5292,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(subsumption_resolution,[],[f5181,f307])).
fof(f5181,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_divides_2'(s('type.integer.int',X55),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X56)))),X58,X57)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1404])).
fof(f5290,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(subsumption_resolution,[],[f5180,f307])).
fof(f5180,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP8('const.integer.int_divides_2'(s('type.integer.int',X51),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X52)))),X53,X54)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1403])).
fof(f5288,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5179,f307])).
fof(f5179,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,X49,'const.integer.int_divides_2'(s('type.integer.int',X47),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X48)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1402])).
fof(f5286,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(subsumption_resolution,[],[f5178,f307])).
fof(f5178,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'const.integer.int_divides_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X44)))),X46)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1401])).
fof(f5284,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(subsumption_resolution,[],[f5177,f307])).
fof(f5177,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP3(X41,'const.integer.int_divides_2'(s('type.integer.int',X39),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40)))),X42)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1400])).
fof(f5282,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(subsumption_resolution,[],[f5176,f307])).
fof(f5176,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP4(X37,'const.integer.int_divides_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X36)))),X38)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1399])).
fof(f5280,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(subsumption_resolution,[],[f5175,f307])).
fof(f5175,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5(X33,'const.integer.int_divides_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))),X34)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1398])).
fof(f5278,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5174,f307])).
fof(f5174,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_divides_2'(s('type.integer.int',X28),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1397])).
fof(f5276,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(subsumption_resolution,[],[f5173,f307])).
fof(f5173,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_divides_2'(s('type.integer.int',X25),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X26)))),X27)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1396])).
fof(f5274,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(subsumption_resolution,[],[f5172,f307])).
fof(f5172,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_divides_2'(s('type.integer.int',X22),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X23)))),X24)) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1395])).
fof(f5272,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5171,f307])).
fof(f5171,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_divides_2'(s('type.integer.int',X19),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f1394])).
fof(f5270,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5168,f307])).
fof(f5168,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f413])).
fof(f5268,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5167,f307])).
fof(f5167,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f412])).
fof(f5266,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f5166,f307])).
fof(f5166,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f410])).
fof(f5257,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f5039,f4920])).
fof(f5039,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f4920,f336])).
fof(f5363,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5354,f5356,f5362])).
fof(f5361,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5226,f307])).
fof(f5226,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4920,f4173])).
fof(f5360,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5225,f307])).
fof(f5225,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4920,f2223])).
fof(f5359,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5224,f307])).
fof(f5224,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4920,f4109])).
fof(f5358,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5223,f307])).
fof(f5223,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4920,f2222])).
fof(f5357,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5354,f5356])).
fof(f5355,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352,f5354])).
fof(f5353,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350,f5352])).
fof(f5351,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348,f5350])).
fof(f5349,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346,f5348])).
fof(f5347,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344,f5346])).
fof(f5345,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342,f5344])).
fof(f5343,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340,f5342])).
fof(f5341,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338,f5340])).
fof(f5339,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336,f5338])).
fof(f5337,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334,f5336])).
fof(f5335,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332,f5334])).
fof(f5333,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330,f5332])).
fof(f5331,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328,f5330])).
fof(f5329,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326,f5328])).
fof(f5327,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324,f5326])).
fof(f5325,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322,f5324])).
fof(f5323,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320,f5322])).
fof(f5321,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318,f5320])).
fof(f5319,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316,f5318])).
fof(f5317,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314,f5316])).
fof(f5315,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312,f5314])).
fof(f5313,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310,f5312])).
fof(f5311,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308,f5310])).
fof(f5309,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306,f5308])).
fof(f5307,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304,f5306])).
fof(f5305,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302,f5304])).
fof(f5303,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300,f5302])).
fof(f5301,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298,f5300])).
fof(f5299,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296,f5298])).
fof(f5297,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294,f5296])).
fof(f5295,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292,f5294])).
fof(f5293,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290,f5292])).
fof(f5291,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288,f5290])).
fof(f5289,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286,f5288])).
fof(f5287,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284,f5286])).
fof(f5285,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282,f5284])).
fof(f5283,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280,f5282])).
fof(f5281,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278,f5280])).
fof(f5279,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276,f5278])).
fof(f5277,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274,f5276])).
fof(f5275,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272,f5274])).
fof(f5273,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270,f5272])).
fof(f5271,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268,f5270])).
fof(f5269,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266,f5268])).
fof(f5267,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257,f5266])).
fof(f5259,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5041,f307])).
fof(f5041,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f4920,f546])).
fof(f5258,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f306,f305,f307,f546,f314,f315,f316,f317,f547,f319,f548,f321,f331,f333,f332,f337,f336,f389,f393,f392,f398,f399,f400,f551,f553,f401,f555,f557,f558,f559,f409,f406,f415,f414,f413,f412,f410,f417,f416,f419,f418,f422,f421,f420,f432,f431,f560,f450,f447,f446,f445,f455,f462,f461,f460,f459,f458,f457,f456,f466,f465,f464,f463,f470,f469,f468,f467,f474,f473,f472,f471,f478,f477,f476,f475,f532,f482,f481,f480,f479,f489,f488,f487,f486,f495,f492,f491,f490,f502,f501,f500,f531,f498,f497,f496,f506,f505,f504,f503,f510,f509,f508,f507,f519,f518,f517,f530,f515,f514,f513,f523,f522,f521,f520,f527,f526,f525,f524,f566,f567,f568,f569,f570,f571,f572,f573,f574,f575,f576,f577,f578,f579,f580,f581,f582,f583,f584,f585,f599,f609,f610,f613,f618,f621,f626,f629,f643,f649,f650,f654,f657,f661,f669,f670,f672,f680,f730,f781,f798,f796,f816,f812,f856,f873,f882,f886,f891,f895,f906,f915,f943,f949,f957,f988,f999,f1008,f1045,f1052,f1061,f1060,f1062,f1065,f1066,f1067,f1218,f1224,f1246,f1245,f1282,f1316,f1383,f1382,f1381,f1380,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1410,f1411,f1412,f1438,f1437,f1436,f1435,f1451,f1452,f1453,f1454,f1455,f1456,f1457,f1458,f1459,f1460,f1461,f1462,f1463,f1464,f1465,f1466,f1467,f1468,f1469,f1569,f1554,f1572,f1561,f1576,f1642,f1641,f1652,f1653,f1654,f1655,f1658,f1657,f1659,f1661,f1662,f1663,f1664,f1809,f1787,f1810,f1811,f1814,f1880,f1884,f1931,f1998,f2128,f2137,f2139,f2222,f2223,f2226,f2219,f2299,f2306,f2295,f2348,f2351,f2352,f2353,f2392,f2395,f2400,f2398,f2397,f2413,f2411,f2410,f2477,f2488,f2594,f2605,f2708,f2714,f2717,f2788,f2791,f2847,f2868,f2871,f2878,f2879,f2939,f2952,f3127,f3140,f3361,f3362,f3364,f3327,f3337,f3394,f3395,f3432,f3398,f3408,f3615,f3614,f3613,f3611,f3610,f3609,f3606,f3605,f3601,f3598,f3597,f3593,f3592,f3587,f3572,f3571,f3570,f3566,f3565,f3561,f3560,f3559,f3556,f3552,f3550,f3547,f3546,f3542,f3541,f3540,f3539,f3538,f3535,f3533,f3528,f3527,f3523,f3522,f3517,f3510,f3509,f3507,f3505,f3504,f3503,f3498,f3497,f3492,f3491,f3490,f3489,f3485,f3484,f3479,f3478,f3474,f3473,f3468,f3467,f3465,f3464,f3461,f3460,f3459,f3458,f3456,f3455,f3454,f3452,f3657,f3771,f3856,f3899,f3902,f3949,f3977,f4006,f4009,f4062,f4109,f4091,f4122,f4121,f4173,f4155,f4189,f4245,f4246,f4247,f4297,f4260,f4303,f4312,f4333,f4334,f4335,f4380,f4388,f4393,f4415,f4418,f4441,f4454,f4532,f4533,f4534,f4535,f4593,f4918,f4922,f4925,f4928,f4931,f4934,f4937,f4940,f4943,f4946,f4949,f4952,f4955,f4958,f4961,f4964,f4967,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4994,f4995,f4996,f4997,f4674,f4675,f4676,f4677,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4920,f5257])).
fof(f4968,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f4968_D])).
fof(f4968_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4920,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f4920_D])).
fof(f4920_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5035,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5034])).
fof(f5034,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4705,f781])).
fof(f4705,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f730])).
fof(f4965,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f4965_D])).
fof(f4965_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5033,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5032])).
fof(f5032,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4704,f781])).
fof(f4704,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f730])).
fof(f4962,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4962_D])).
fof(f4962_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5031,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5030])).
fof(f5030,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4703,f781])).
fof(f4703,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f730])).
fof(f4959,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f4959_D])).
fof(f4959_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5029,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5028])).
fof(f5028,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4702,f781])).
fof(f4702,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f730])).
fof(f4956,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f4956_D])).
fof(f4956_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5027,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f5026])).
fof(f5026,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4701,f781])).
fof(f4701,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f730])).
fof(f4953,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4953_D])).
fof(f4953_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5025,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f5024])).
fof(f5024,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4700,f781])).
fof(f4700,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f730])).
fof(f4950,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f4950_D])).
fof(f4950_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5023,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5022])).
fof(f5022,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4699,f781])).
fof(f4699,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f730])).
fof(f4947,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f4947_D])).
fof(f4947_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5021,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5020])).
fof(f5020,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4698,f781])).
fof(f4698,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f730])).
fof(f4944,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f4944_D])).
fof(f4944_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5019,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f5018])).
fof(f5018,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4697,f781])).
fof(f4697,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f730])).
fof(f4941,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f4941_D])).
fof(f4941_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f5017,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f5016])).
fof(f5016,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4696,f781])).
fof(f4696,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f730])).
fof(f4938,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f4938_D])).
fof(f4938_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5015,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f5014])).
fof(f5014,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4695,f781])).
fof(f4695,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f730])).
fof(f4935,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f4935_D])).
fof(f4935_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5013,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f5012])).
fof(f5012,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4694,f781])).
fof(f4694,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f730])).
fof(f4932,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f4932_D])).
fof(f4932_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5011,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f5010])).
fof(f5010,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4693,f781])).
fof(f4693,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f730])).
fof(f4929,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f4929_D])).
fof(f4929_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5009,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f5008])).
fof(f5008,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4692,f781])).
fof(f4692,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f730])).
fof(f4926,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f4926_D])).
fof(f4926_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5007,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f5006])).
fof(f5006,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4691,f781])).
fof(f4691,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f730])).
fof(f4923,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl8),
  inference(cnf_transformation,[],[f4923_D])).
fof(f4923_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5005,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f5004])).
fof(f5004,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4690,f781])).
fof(f4690,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f730])).
fof(f4919,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl4),
  inference(cnf_transformation,[],[f4919_D])).
fof(f4919_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5003,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f5002])).
fof(f5002,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4689,f781])).
fof(f4689,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f730])).
fof(f4967,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4966])).
fof(f4966,plain,(
  ( ! [X94,X92,X95,X93] : (~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4620,f781])).
fof(f4620,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP8('T_0',X93,X94) | sP9('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f531,f730])).
fof(f4964,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4963])).
fof(f4963,plain,(
  ( ! [X90,X88,X91,X89] : (~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f4619,f781])).
fof(f4619,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP10('T_0',X89,X90) | sP11('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f530,f730])).
fof(f4961,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4960])).
fof(f4960,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f4618,f781])).
fof(f4618,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP11('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f527,f730])).
fof(f4958,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4957])).
fof(f4957,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4617,f781])).
fof(f4617,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP11('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f526,f730])).
fof(f4955,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4954])).
fof(f4954,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f4616,f781])).
fof(f4616,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP11('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f525,f730])).
fof(f4952,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4951])).
fof(f4951,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f4615,f781])).
fof(f4615,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP10(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f520,f730])).
fof(f4949,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4948])).
fof(f4948,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f4614,f781])).
fof(f4614,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f510,f730])).
fof(f4946,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4945])).
fof(f4945,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4613,f781])).
fof(f4613,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f509,f730])).
fof(f4943,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4942])).
fof(f4942,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f4612,f781])).
fof(f4612,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP9('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f508,f730])).
fof(f4940,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4939])).
fof(f4939,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f4611,f781])).
fof(f4611,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP8('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f503,f730])).
fof(f4937,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4936])).
fof(f4936,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f4610,f781])).
fof(f4610,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP7(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f495,f730])).
fof(f4934,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4933])).
fof(f4933,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f4609,f781])).
fof(f4609,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP6(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f486,f730])).
fof(f4931,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4930])).
fof(f4930,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f4608,f781])).
fof(f4608,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP3(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f467,f730])).
fof(f4928,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4927])).
fof(f4927,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4607,f781])).
fof(f4607,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f463,f730])).
fof(f4925,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4924])).
fof(f4924,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f4606,f781])).
fof(f4606,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f459,f730])).
fof(f4922,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4921])).
fof(f4921,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4605,f781])).
fof(f4605,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f450,f730])).
fof(f4918,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4917])).
fof(f4917,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f4604,f781])).
fof(f4604,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f447,f730])).
fof(f542,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | $spl0),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f544,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl2),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f305,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))))),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f202])).
fof(f202,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) & (~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))))))))),
  file('thm.integer.INT_DIVIDES_REDUCE|split|6','thm.integer.INT_DIVIDES_REDUCE|split|6')).
fof(f543,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) | $spl1),
  inference(cnf_transformation,[],[f543_D])).
fof(f543_D,plain,(
  ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13)))))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f545,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | $spl3),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f306,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))) | ~p(s(bool,'const.integer.int_divides_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK12))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK13))))))))))),
  inference(cnf_transformation,[],[f203])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 787
% Passive clauses: 2415
% Generated clauses: 21296
% Final active clauses: 542
% Final passive clauses: 1174
% Input formulas: 69
% Initial clauses: 225
% 
% Duplicate literals: 196
% Trivial inequalities: 136
% Fw subsumption resolutions: 841
% Fw demodulations: 5545
% Bw demodulations: 245
% Condensations: 219
% Global subsumptions: 51
% 
% Simple tautologies: 113
% Equational tautologies: 71
% Forward subsumptions: 11078
% Fw demodulations to eq. taut.: 141
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 5384
% Factoring: 1
% Forward superposition: 4695
% Backward superposition: 3356
% Self superposition: 48
% Equality factoring: 28
% Equality resolution: 26
% 
% Split clauses: 57
% Split components: 119
% SAT solver clauses: 3015
% SAT solver unit clauses: 714
% SAT solver binary clauses: 851
% 
% Sat splits: 57
% Sat splitting refutations: 433
% 
% Memory used [KB]: 13560
% Time elapsed: 0.726 s
% ------------------------------
% ------------------------------
% Success in time 2.043 s
