Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.transc
% (86999)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 225
% Passive clauses: 2835
% Generated clauses: 28680
% Final active clauses: 206
% Final passive clauses: 2418
% Input formulas: 29
% Initial clauses: 58
% 
% Duplicate literals: 894
% Trivial inequalities: 389
% Fw subsumption resolutions: 4656
% Fw demodulations: 5907
% Bw demodulations: 17
% 
% Simple tautologies: 403
% Equational tautologies: 22
% Forward subsumptions: 11971
% 
% Binary resolution: 1028
% Factoring: 1
% Forward superposition: 9910
% Backward superposition: 5409
% Self superposition: 111
% Equality factoring: 47
% Equality resolution: 30
% 
% Split clauses: 694
% Split components: 1403
% SAT solver clauses: 1760
% SAT solver unit clauses: 22
% SAT solver binary clauses: 137
% 
% Sat splits: 694
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1759
% TWLsolver calls for satisfiability: 235
% 
% Memory used [KB]: 13304
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% (87018)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 123
% Passive clauses: 598
% Generated clauses: 11278
% Final active clauses: 109
% Final passive clauses: 447
% Input formulas: 29
% Initial clauses: 58
% 
% Duplicate literals: 1047
% Trivial inequalities: 143
% Fw subsumption resolutions: 1065
% Fw demodulations: 943
% Bw demodulations: 38
% Condensations: 1724
% Global subsumptions: 8
% 
% Simple tautologies: 275
% Equational tautologies: 59
% Forward subsumptions: 3753
% 
% Binary resolution: 209
% Forward superposition: 2068
% Backward superposition: 3366
% Self superposition: 166
% Equality factoring: 160
% Equality resolution: 17
% 
% Split clauses: 257
% Split components: 533
% SAT solver clauses: 1299
% SAT solver unit clauses: 72
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 16
% 
% Sat splits: 257
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1267
% TWLsolver calls for satisfiability: 249
% 
% Memory used [KB]: 8059
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f8315,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f684,f685,f701,f730,f731,f734,f2773,f2774,f2780,f2776,f2781,f2832,f2834,f2910,f2911,f2915,f2919,f2923,f2927,f2932,f2934,f2936,f2940,f2942,f2944,f2946,f2948,f2953,f2957,f2962,f2967,f2969,f2971,f2973,f2979,f2980,f2988,f3008,f3014,f3016,f4963,f4967,f4965,f5995,f5996,f6002,f686,f6003,f6145,f6174,f6172,f6313,f6366,f6367,f6371,f6377,f6379,f6383,f6387,f6391,f6402,f6405,f6407,f6410,f6430,f6432,f6440,f6448,f6469,f6473,f6477,f6480,f6483,f6489,f6494,f6497,f6501,f6504,f6510,f6515,f6518,f6541,f6545,f6549,f6554,f6561,f6565,f6567,f6570,f6574,f7643,f7647,f7645,f8257,f8262,f8293,f8294,f8298,f8299,f8305,f8258,f8309,f8260,f8310])).
fof(f8310,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl34),
  inference(cnf_transformation,[],[f8310_D])).
fof(f8310_D,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f8260,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | $spl28),
  inference(cnf_transformation,[],[f8260_D])).
fof(f8260_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8309,plain,(
  ( ! [X25] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) ) | $spl19),
  inference(forward_demodulation,[],[f8308,f8128])).
fof(f8128,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(subsumption_resolution,[],[f8092,f109])).
fof(f109,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.EXP_TOTAL',aHL_TRUTH)).
fof(f8092,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(superposition,[],[f116,f8008])).
fof(f8008,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl19),
  inference(unit_resulting_resolution,[],[f7636,f283])).
fof(f283,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f280,f111])).
fof(f111,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.EXP_TOTAL',aHL_FALSITY)).
fof(f280,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f162,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.transc.EXP_TOTAL',aHL_BOOL_CASES)).
fof(f162,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X6)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))),
  file('thm.transc.EXP_TOTAL','thm.real.REAL_LT_IMP_LE')).
fof(f7636,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(subsumption_resolution,[],[f7635,f109])).
fof(f7635,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(forward_demodulation,[],[f7634,f173])).
fof(f173,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))),
  inference(unit_resulting_resolution,[],[f107,f120])).
fof(f120,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.transc.EXP_TOTAL','thm.bool.EQ_CLAUSES|split|0')).
fof(f107,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) & ! [X1] : s('type.realax.real',sK0) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f65])).
fof(f65,plain,(
  ? [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) & ! [X1] : s('type.realax.real',X0) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => ? [X1] : s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => ? [X7] : s('type.realax.real',X6) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => ? [X7] : s('type.realax.real',X6) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7))))),
  file('thm.transc.EXP_TOTAL','thm.transc.EXP_TOTAL')).
fof(f7634,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(forward_demodulation,[],[f7633,f7195])).
fof(f7195,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | $spl19),
  inference(unit_resulting_resolution,[],[f6171,f118])).
fof(f118,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real',X0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real',X0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7] : (~s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real',X7) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))))))),
  file('thm.transc.EXP_TOTAL','thm.real.REAL_INVINV')).
fof(f6171,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(cnf_transformation,[],[f6171_D])).
fof(f6171_D,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f7633,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(subsumption_resolution,[],[f7632,f109])).
fof(f7632,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(forward_demodulation,[],[f7631,f929])).
fof(f929,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f919,f120])).
fof(f919,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(subsumption_resolution,[],[f910,f735])).
fof(f735,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f116,f115,f114,f117,f118,f120,f119,f124,f123,f161,f162,f163,f164,f171,f173,f188,f191,f211,f213,f212,f248,f259,f283,f312,f517,f370,f532,f577,f565,f568,f569,f570,f571,f572,f573,f670,f672,f684,f687,f689,f706,f700])).
fof(f700,plain,(
  ( ! [X10] : (~p(s(bool,'F_0')) | s('type.realax.real',X10) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X10)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X10)))) )),
  inference(superposition,[],[f116,f113])).
fof(f706,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X2) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X2)))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f696,f109])).
fof(f696,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X2) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X2)))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f116,f188])).
fof(f689,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(subsumption_resolution,[],[f688,f111])).
fof(f688,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(forward_demodulation,[],[f676,f212])).
fof(f676,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X5)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f248,f517])).
fof(f687,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f675,f111])).
fof(f675,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f248,f212])).
fof(f672,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X3))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f248,f118])).
fof(f670,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f248,f118])).
fof(f573,plain,(
  ( ! [X12,X13] : (s('type.realax.real',X12) != s('type.realax.real',X13) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12)))))))) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f163,f118])).
fof(f572,plain,(
  ( ! [X10,X11] : (s('type.realax.real',X10) != s('type.realax.real',X11) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X11)))) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f163,f118])).
fof(f571,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8)))) | s('type.realax.real',X8) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f162,f118])).
fof(f570,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f162,f118])).
fof(f569,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X5)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X4)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f161,f118])).
fof(f568,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f161,f118])).
fof(f565,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f118,f112])).
fof(f577,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f576,f112])).
fof(f576,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f563,f112])).
fof(f563,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f118,f112])).
fof(f532,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f380,f109])).
fof(f380,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f191,f188])).
fof(f370,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f109,f188])).
fof(f517,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f322,f171])).
fof(f322,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f188,f212])).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f188,f188])).
fof(f259,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f254,f109])).
fof(f254,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f191,f173])).
fof(f248,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f245,f111])).
fof(f245,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f161,f113])).
fof(f212,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f211,f124])).
fof(f213,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f211,f119])).
fof(f211,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(equality_resolution,[],[f163])).
fof(f191,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f123,f113])).
fof(f188,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f113,f113])).
fof(f171,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f111,f119])).
fof(f164,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.transc.EXP_TOTAL',aHL_EXT)).
fof(f163,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => ~s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7,X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X6)))) => ~s('type.realax.real',X6) = s('type.realax.real',X7))),
  file('thm.transc.EXP_TOTAL','thm.real.REAL_LT_IMP_NE')).
fof(f161,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))),
  file('thm.transc.EXP_TOTAL','thm.real.REAL_LET_TOTAL')).
fof(f123,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.transc.EXP_TOTAL','thm.bool.EQ_CLAUSES|split|3')).
fof(f124,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f119,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f117,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7)))))))),
  file('thm.transc.EXP_TOTAL','thm.real.REAL_INV_LT|49|')).
fof(f114,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X0)))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X0))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) & s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f66])).
fof(f66,plain,(
  ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | ? [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) & s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) => ? [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) & s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X6)))) => ? [X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) & s('type.realax.real',X6) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))),
  file('thm.transc.EXP_TOTAL','thm.transc.EXP_TOTAL_LEMMA')).
fof(f115,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f112,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))),
  file('thm.transc.EXP_TOTAL','thm.transc.EXP_NEG')).
fof(f108,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f910,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(superposition,[],[f161,f758])).
fof(f758,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))),
  inference(forward_demodulation,[],[f750,f212])).
fof(f750,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))) )),
  inference(unit_resulting_resolution,[],[f213,f692,f370])).
fof(f692,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))))),
  inference(unit_resulting_resolution,[],[f108,f116])).
fof(f7631,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(forward_demodulation,[],[f7612,f7195])).
fof(f7612,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(superposition,[],[f117,f7195])).
fof(f116,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f8308,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8307,f112])).
fof(f8307,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8249,f112])).
fof(f8249,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(superposition,[],[f6074,f8128])).
fof(f6074,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f5955,f118])).
fof(f5955,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f5870,f173])).
fof(f5870,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))) )),
  inference(condensation,[],[f5869])).
fof(f5869,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37))) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X36),s('type.realax.real',X37)))) )),
  inference(duplicate_literal_removal,[],[f5725])).
fof(f5725,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',X37))) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X36),s('type.realax.real',X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f259,f1705])).
fof(f1705,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,X8) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1694,f109])).
fof(f1694,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,X8) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f283,f188])).
fof(f8258,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X0)))) ) | $spl26),
  inference(cnf_transformation,[],[f8258_D])).
fof(f8258_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X0)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8305,plain,(
  ( ! [X23] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X23))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X23)))) ) | $spl19),
  inference(forward_demodulation,[],[f8304,f8128])).
fof(f8304,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X23))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X23))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8303,f112])).
fof(f8303,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X23))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X23))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8247,f112])).
fof(f8247,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X23))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X23))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(superposition,[],[f6067,f8128])).
fof(f6067,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))),s('type.realax.real',X7))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f5955,f118])).
fof(f8299,plain,(
  ( ! [X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) | $spl32),
  inference(cnf_transformation,[],[f8299_D])).
fof(f8299_D,plain,(
  ( ! [X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f8298,plain,(
  ( ! [X16] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) | $spl19),
  inference(forward_demodulation,[],[f8297,f8128])).
fof(f8297,plain,(
  ( ! [X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8296,f112])).
fof(f8296,plain,(
  ( ! [X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8240,f112])).
fof(f8240,plain,(
  ( ! [X16] : (s('type.realax.real',X16) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(superposition,[],[f573,f8128])).
fof(f8294,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) | $spl30),
  inference(cnf_transformation,[],[f8294_D])).
fof(f8294_D,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f8293,plain,(
  ( ! [X14] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))) ) | $spl19),
  inference(forward_demodulation,[],[f8292,f8128])).
fof(f8292,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8291,f112])).
fof(f8291,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(forward_demodulation,[],[f8238,f112])).
fof(f8238,plain,(
  ( ! [X14] : (s('type.realax.real',X14) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X14)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(superposition,[],[f572,f8128])).
fof(f8262,plain,(
  $false | $spl19),
  inference(subsumption_resolution,[],[f8261,f108])).
fof(f8261,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(forward_demodulation,[],[f8213,f7195])).
fof(f8213,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))) | $spl19),
  inference(superposition,[],[f112,f8128])).
fof(f8257,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X0)))) ) | $spl19),
  inference(forward_demodulation,[],[f8212,f8128])).
fof(f8212,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0)))))))))),s('type.realax.real',X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real',sK0))))))) ) | $spl19),
  inference(superposition,[],[f6101,f8128])).
fof(f6101,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6065,f112])).
fof(f6065,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f5955,f565])).
fof(f7645,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl23),
  inference(cnf_transformation,[],[f7645_D])).
fof(f7645_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f7647,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl25),
  inference(cnf_transformation,[],[f7647_D])).
fof(f7647_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f7643,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(subsumption_resolution,[],[f7642,f735])).
fof(f7642,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(forward_demodulation,[],[f7617,f974])).
fof(f974,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))),
  inference(forward_demodulation,[],[f966,f758])).
fof(f966,plain,(
  s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))),
  inference(unit_resulting_resolution,[],[f751,f959,f370])).
fof(f959,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))))),
  inference(subsumption_resolution,[],[f950,f735])).
fof(f950,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))))),
  inference(superposition,[],[f162,f758])).
fof(f751,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))),
  inference(unit_resulting_resolution,[],[f692,f119])).
fof(f7617,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))) | $spl19),
  inference(superposition,[],[f117,f7195])).
fof(f6574,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553,f6560,f6564,f6566,f6569,f6573])).
fof(f6573,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6572,f109])).
fof(f6572,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6571,f6314])).
fof(f6571,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6287,f108])).
fof(f6287,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5526])).
fof(f6569,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6568,f109])).
fof(f6568,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6286,f6314])).
fof(f6286,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5511])).
fof(f6566,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6285,f108])).
fof(f6285,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5494])).
fof(f6564,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6563,f108])).
fof(f6563,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6562,f6172])).
fof(f6562,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6284,f108])).
fof(f6284,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5490])).
fof(f6560,plain,(
  ( ! [X9] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6559,f109])).
fof(f6559,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6558,f6314])).
fof(f6558,plain,(
  ( ! [X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X9)))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6283,f108])).
fof(f6283,plain,(
  ( ! [X9] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X9))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X9)))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5447])).
fof(f6553,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6552,f6172])).
fof(f6552,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6280,f108])).
fof(f6280,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5434])).
fof(f6548,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6547,f108])).
fof(f6547,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6546,f6172])).
fof(f6546,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6274,f108])).
fof(f6274,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5109])).
fof(f6544,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6543,f6172])).
fof(f6543,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6273,f108])).
fof(f6273,plain,(
  ( ! [X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f5047])).
fof(f6514,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6513,f6172])).
fof(f6513,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6512,f109])).
fof(f6512,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f6511,f6314])).
fof(f6511,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6238,f108])).
fof(f6238,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f2003])).
fof(f6509,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6508,f108])).
fof(f6508,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6507,f6172])).
fof(f6507,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6506,f109])).
fof(f6506,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6505,f6314])).
fof(f6505,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6237,f108])).
fof(f6237,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f2001])).
fof(f6503,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f6502,f109])).
fof(f6502,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X0)) ) | $spl18),
  inference(forward_demodulation,[],[f6236,f6314])).
fof(f6236,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1993])).
fof(f6500,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5)) ) | $spl18),
  inference(forward_demodulation,[],[f6499,f6172])).
fof(f6499,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6498,f109])).
fof(f6498,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f6235,f6314])).
fof(f6235,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1989])).
fof(f6493,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real',X3) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6492,f6172])).
fof(f6492,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6491,f109])).
fof(f6491,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f6490,f6314])).
fof(f6490,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6227,f108])).
fof(f6227,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1893])).
fof(f6488,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6487,f108])).
fof(f6487,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6486,f6172])).
fof(f6486,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6485,f109])).
fof(f6485,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(forward_demodulation,[],[f6484,f6314])).
fof(f6484,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6226,f108])).
fof(f6226,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1891])).
fof(f6482,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f6481,f109])).
fof(f6481,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X0)) ) | $spl18),
  inference(forward_demodulation,[],[f6225,f6314])).
fof(f6225,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1883])).
fof(f6479,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6478,f109])).
fof(f6478,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f6224,f6314])).
fof(f6224,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1882])).
fof(f6476,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real',X4) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5)) ) | $spl18),
  inference(forward_demodulation,[],[f6475,f6172])).
fof(f6475,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6474,f109])).
fof(f6474,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f6223,f6314])).
fof(f6223,plain,(
  ( ! [X4,X5] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))))),s('type.realax.real',X5)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1879])).
fof(f6472,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6471,f109])).
fof(f6471,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl18),
  inference(forward_demodulation,[],[f6470,f6314])).
fof(f6470,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20)))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6217,f108])).
fof(f6217,plain,(
  ( ! [X21,X20] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20)))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1780])).
fof(f6468,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6467,f109])).
fof(f6467,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl18),
  inference(forward_demodulation,[],[f6466,f6314])).
fof(f6466,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X18))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6216,f108])).
fof(f6216,plain,(
  ( ! [X19,X18] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X18))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1779])).
fof(f6447,plain,(
  ( ! [X7] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6446,f109])).
fof(f6446,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6445,f6314])).
fof(f6445,plain,(
  ( ! [X7] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6210,f108])).
fof(f6210,plain,(
  ( ! [X7] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1771])).
fof(f6439,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6438,f108])).
fof(f6438,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6437,f6172])).
fof(f6437,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(subsumption_resolution,[],[f6208,f108])).
fof(f6208,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1769])).
fof(f6431,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6205,f108])).
fof(f6205,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1742])).
fof(f6429,plain,(
  ( ! [X2] : (s('type.realax.real',sK0) = s('type.realax.real',X2) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6428,f6172])).
fof(f6428,plain,(
  ( ! [X2] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(subsumption_resolution,[],[f6204,f108])).
fof(f6204,plain,(
  ( ! [X2] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1741])).
fof(f6409,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6408,f109])).
fof(f6408,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X1) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) ) | $spl18),
  inference(forward_demodulation,[],[f6190,f6314])).
fof(f6190,plain,(
  ( ! [X1] : (s('type.realax.real',sK0) = s('type.realax.real',X1) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f694])).
fof(f6406,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6187,f108])).
fof(f6187,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f577])).
fof(f6404,plain,(
  ( ! [X12,X13] : (s('type.realax.real',sK0) = s('type.realax.real',X12) | s('type.realax.real',X12) != s('type.realax.real',X13)) ) | $spl18),
  inference(subsumption_resolution,[],[f6403,f109])).
fof(f6403,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X12) | s('type.realax.real',X12) != s('type.realax.real',X13)) ) | $spl18),
  inference(forward_demodulation,[],[f6186,f6314])).
fof(f6186,plain,(
  ( ! [X12,X13] : (s('type.realax.real',sK0) = s('type.realax.real',X12) | s('type.realax.real',X12) != s('type.realax.real',X13) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X12))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f573])).
fof(f6401,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real',X10) != s('type.realax.real',X11)) ) | $spl18),
  inference(subsumption_resolution,[],[f6400,f109])).
fof(f6400,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real',X10) != s('type.realax.real',X11)) ) | $spl18),
  inference(forward_demodulation,[],[f6185,f6314])).
fof(f6185,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK0) = s('type.realax.real',X10) | s('type.realax.real',X10) != s('type.realax.real',X11) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10))))),s('type.realax.real',X11))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f572])).
fof(f6390,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6180,f108])).
fof(f6180,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f565])).
fof(f6386,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6349,f109])).
fof(f6349,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6314,f738])).
fof(f6382,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6345,f109])).
fof(f6345,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl18),
  inference(backward_demodulation,[],[f6314,f116])).
fof(f6370,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl18),
  inference(subsumption_resolution,[],[f6329,f109])).
fof(f6329,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl18),
  inference(backward_demodulation,[],[f6314,f163])).
fof(f6365,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(subsumption_resolution,[],[f6315,f6314])).
fof(f6315,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl18),
  inference(backward_demodulation,[],[f6314,f123])).
fof(f6314,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f6179,f517])).
fof(f6179,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f259])).
fof(f6177,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real',X0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f118])).
fof(f6176,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f117])).
fof(f6175,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK0),s('type.realax.real',sK1(X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f114])).
fof(f6122,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6067,f118])).
fof(f6167,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6120,f112])).
fof(f6120,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f6067,f565])).
fof(f6115,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))),s('type.realax.real',X7))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6067,f118])).
fof(f6153,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6152,f112])).
fof(f6152,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6151,f112])).
fof(f6151,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6150,f112])).
fof(f6150,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6113,f112])).
fof(f6113,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f6067,f565])).
fof(f6106,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f6072,f112])).
fof(f6072,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f5955,f565])).
fof(f6057,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X5)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6032,f109])).
fof(f6032,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X5)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X5),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X6) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1315,f188])).
fof(f6053,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X3)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f6029,f109])).
fof(f6029,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X3)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f1315,f689])).
fof(f6052,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X1)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f6028,f109])).
fof(f6028,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X1)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1315,f1705])).
fof(f6051,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(subsumption_resolution,[],[f6027,f109])).
fof(f6027,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(superposition,[],[f1315,f687])).
fof(f6024,plain,(
  ( ! [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X4)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X4)))) )),
  inference(superposition,[],[f1315,f1742])).
fof(f6049,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(forward_demodulation,[],[f6048,f1190])).
fof(f6048,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(subsumption_resolution,[],[f6047,f109])).
fof(f6047,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(forward_demodulation,[],[f6023,f687])).
fof(f6023,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(superposition,[],[f1315,f1190])).
fof(f6021,plain,(
  ( ! [X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1315,f118])).
fof(f6041,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f6019,f112])).
fof(f6019,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f1315,f565])).
fof(f6040,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) )),
  inference(forward_demodulation,[],[f6018,f112])).
fof(f6018,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f1315,f112])).
fof(f5843,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(trivial_inequality_removal,[],[f5839])).
fof(f5839,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(equality_factoring,[],[f1705])).
fof(f6005,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X0)))) )),
  inference(condensation,[],[f6004])).
fof(f6004,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X46)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X45),s('type.realax.real',X44)))) | s('type.realax.real',X44) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))),s('type.realax.real',X45)))) )),
  inference(subsumption_resolution,[],[f5811,f119])).
fof(f5811,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X46)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X45),s('type.realax.real',X44)))) | s('type.realax.real',X44) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))))),s('type.realax.real',X45))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f569,f1705])).
fof(f5998,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))) )),
  inference(condensation,[],[f5997])).
fof(f5997,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X24)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X22)))) | s('type.realax.real',X23) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X23)))))))) )),
  inference(subsumption_resolution,[],[f5799,f119])).
fof(f5799,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X24)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X22)))) | s('type.realax.real',X23) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X23))))))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f568,f1705])).
fof(f5790,plain,(
  ( ! [X280,X277,X279,X276,X278] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X279),s('type.realax.real',X280))) != s(bool,f(s(X276,sK2(X277,X278,bool,X276)))) | s(fun(X276,bool),X277) = s(fun(X276,bool),X278) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X279),s('type.realax.real',X280))) | s(bool,'T_0') = s(bool,g(s(X276,sK2(X277,X278,bool,X276))))) )),
  inference(superposition,[],[f164,f1705])).
fof(f5789,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X273,f(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X274),s('type.realax.real',X275))))) != s(X273,g(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X274),s('type.realax.real',X275))))) | s(fun(bool,X273),X271) = s(fun(bool,X273),X272) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X274),s('type.realax.real',X275))) | s(bool,'T_0') = s(bool,sK2(X271,X272,X273,bool))) )),
  inference(superposition,[],[f164,f1705])).
fof(f5866,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(condensation,[],[f5865])).
fof(f5865,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X54) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',X53))) | s(bool,'T_0') = s(bool,X51) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X52),s('type.realax.real',X53)))) )),
  inference(duplicate_literal_removal,[],[f5729])).
fof(f5729,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X54) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',X53))) | s(bool,'T_0') = s(bool,X51) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X52),s('type.realax.real',X53))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f532,f1705])).
fof(f5948,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X68) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real',X67))) | s(bool,'T_0') = s(bool,X68) | s(bool,X69) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real',X67))) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f5699,f109])).
fof(f5699,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | s(bool,X68) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real',X67))) | s(bool,'T_0') = s(bool,X68) | s(bool,X69) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real',X67))) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f191,f1705])).
fof(f5947,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30))))))))) | s(bool,X31) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))))))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f5688,f109])).
fof(f5688,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30))))))))) | s(bool,X31) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))))))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f1146,f1705])).
fof(f5946,plain,(
  ( ! [X28,X29] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28)))))))) | s(bool,X29) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5687,f109])).
fof(f5687,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28)))))))) | s(bool,X29) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f738,f1705])).
fof(f5945,plain,(
  ( ! [X26,X27] : (s('type.realax.real',X26) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X26)))) | s(bool,X27) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f5686,f109])).
fof(f5686,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | s('type.realax.real',X26) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X26)))) | s(bool,X27) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f116,f1705])).
fof(f5944,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X24))))) | s(bool,X25) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f5685,f109])).
fof(f5685,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X24))))) | s(bool,X25) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f114,f1705])).
fof(f5655,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real',X9))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))))) | s(bool,'T_0') = s(bool,X10) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1705,f118])).
fof(f5937,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))))) | s(bool,'T_0') = s(bool,X4) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f5653,f112])).
fof(f5653,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s(bool,'T_0') = s(bool,X4) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f1705,f565])).
fof(f5648,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X11))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))),s('type.realax.real',X11))) | s(bool,'T_0') = s(bool,X10) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1705,f118])).
fof(f5932,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))),s('type.realax.real',X5))) | s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,X4) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f5646,f112])).
fof(f5646,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,X4) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f1705,f565])).
fof(f5896,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(condensation,[],[f5895])).
fof(f5895,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X32),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) )),
  inference(duplicate_literal_removal,[],[f5623])).
fof(f5623,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X32),s('type.realax.real',X33))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real',X31))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1705,f1705])).
fof(f5908,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(condensation,[],[f5907])).
fof(f5907,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f5615])).
fof(f5615,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1705,f113])).
fof(f5602,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X9))) | s(bool,X11) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X10),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))))) | s(bool,'T_0') = s(bool,X11) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1705,f118])).
fof(f5924,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))))) | s(bool,'T_0') = s(bool,X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f5600,f112])).
fof(f5600,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))))) | s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s(bool,'T_0') = s(bool,X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f1705,f565])).
fof(f5595,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10))) | s(bool,X11) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,X11) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1705,f118])).
fof(f5919,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f5593,f112])).
fof(f5593,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))),s('type.realax.real',X4))) | s(bool,X5) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f1705,f565])).
fof(f5531,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f5518,f109])).
fof(f5518,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f738,f188])).
fof(f5530,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f5516,f109])).
fof(f5516,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f738,f689])).
fof(f5511,plain,(
  ( ! [X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f738,f118])).
fof(f5526,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f5509,f112])).
fof(f5509,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f738,f565])).
fof(f5447,plain,(
  ( ! [X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X9)))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X9)))) )),
  inference(superposition,[],[f116,f1742])).
fof(f5446,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))) )),
  inference(superposition,[],[f115,f1742])).
fof(f5445,plain,(
  ( ! [X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f114,f1742])).
fof(f5494,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f5441,f112])).
fof(f5441,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f112,f1742])).
fof(f5434,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X3))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1742,f118])).
fof(f5490,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f5432,f112])).
fof(f5432,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f1742,f565])).
fof(f5298,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f5285,f109])).
fof(f5285,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1146,f188])).
fof(f5297,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f5283,f109])).
fof(f5283,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f1146,f689])).
fof(f5280,plain,(
  ( ! [X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1146,f118])).
fof(f5293,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f5278,f112])).
fof(f5278,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f1146,f565])).
fof(f5047,plain,(
  ( ! [X3] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f577,f118])).
fof(f5109,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f5045,f112])).
fof(f5045,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f577,f565])).
fof(f4837,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4812,f109])).
fof(f4812,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f191,f4446])).
fof(f4724,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15))))),s('type.realax.real',X15))) | s(bool,'T_0') = s(bool,X16) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f4691,f109])).
fof(f4691,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15))))),s('type.realax.real',X15))) | s(bool,'T_0') = s(bool,X16) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f191,f2594])).
fof(f4718,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f4667,f112])).
fof(f4667,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f2594,f112])).
fof(f4662,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15))))))) | s(bool,'T_0') = s(bool,X16) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f4624,f109])).
fof(f4624,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s(bool,X16) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15))))))) | s(bool,'T_0') = s(bool,X16) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f191,f2132])).
fof(f4655,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f4602,f112])).
fof(f4602,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f2132,f112])).
fof(f4653,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) )),
  inference(forward_demodulation,[],[f4652,f112])).
fof(f4652,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f4651,f112])).
fof(f4651,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f4598,f112])).
fof(f4598,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f2132,f112])).
fof(f4446,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(unit_resulting_resolution,[],[f2738,f120])).
fof(f694,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f116,f118])).
fof(f736,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X4)))) | s(bool,X5) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f697,f109])).
fof(f697,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',X4) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X4)))) | s(bool,X5) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f116,f188])).
fof(f738,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(forward_demodulation,[],[f693,f112])).
fof(f693,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(superposition,[],[f116,f112])).
fof(f762,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f753,f109])).
fof(f753,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f692,f188])).
fof(f967,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0)))),
  inference(unit_resulting_resolution,[],[f959,f119])).
fof(f979,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f969,f109])).
fof(f969,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f959,f188])).
fof(f984,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f962,f120])).
fof(f1039,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1026,f109])).
fof(f1026,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f191,f929])).
fof(f1145,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1132,f109])).
fof(f1132,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f191,f984])).
fof(f1146,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))))))) )),
  inference(superposition,[],[f114,f112])).
fof(f1147,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f114,f118])).
fof(f1156,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X2))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1150,f109])).
fof(f1150,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X2))))) | s(bool,X3) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f114,f188])).
fof(f1190,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(subsumption_resolution,[],[f1166,f109])).
fof(f1166,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(superposition,[],[f116,f687])).
fof(f1306,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f115,f118])).
fof(f1314,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) )),
  inference(forward_demodulation,[],[f1305,f112])).
fof(f1305,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(superposition,[],[f115,f112])).
fof(f1315,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(subsumption_resolution,[],[f1307,f735])).
fof(f1307,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f115,f113])).
fof(f1502,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f117,f118])).
fof(f1503,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))) )),
  inference(superposition,[],[f117,f112])).
fof(f1504,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f117,f118])).
fof(f1515,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) )),
  inference(forward_demodulation,[],[f1514,f112])).
fof(f1514,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(forward_demodulation,[],[f1513,f112])).
fof(f1513,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(forward_demodulation,[],[f1501,f112])).
fof(f1501,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) )),
  inference(superposition,[],[f117,f112])).
fof(f1516,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) )),
  inference(subsumption_resolution,[],[f1505,f735])).
fof(f1505,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f117,f113])).
fof(f1744,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f117,f565])).
fof(f1748,plain,(
  ( ! [X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X5)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X5)))) )),
  inference(superposition,[],[f114,f565])).
fof(f1770,plain,(
  ( ! [X6] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X6)))) )),
  inference(forward_demodulation,[],[f1749,f112])).
fof(f1749,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X6)))) )),
  inference(superposition,[],[f115,f565])).
fof(f1771,plain,(
  ( ! [X7] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) )),
  inference(forward_demodulation,[],[f1750,f112])).
fof(f1750,plain,(
  ( ! [X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f116,f565])).
fof(f1774,plain,(
  ( ! [X8] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))) )),
  inference(forward_demodulation,[],[f1773,f112])).
fof(f1773,plain,(
  ( ! [X8] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))) )),
  inference(forward_demodulation,[],[f1772,f112])).
fof(f1772,plain,(
  ( ! [X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))) )),
  inference(forward_demodulation,[],[f1751,f112])).
fof(f1751,plain,(
  ( ! [X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X8)))) )),
  inference(superposition,[],[f117,f565])).
fof(f1984,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X8))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f1945,f109])).
fof(f1945,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X8))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f114,f689])).
fof(f1985,plain,(
  ( ! [X9] : (s('type.realax.real',X9) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X9)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(subsumption_resolution,[],[f1946,f109])).
fof(f1946,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s('type.realax.real',X9) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X9)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f116,f689])).
fof(f2255,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2254])).
fof(f2254,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(duplicate_literal_removal,[],[f2236])).
fof(f2236,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f532,f979])).
fof(f2257,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2256])).
fof(f2256,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(duplicate_literal_removal,[],[f2235])).
fof(f2235,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f532,f762])).
fof(f2689,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(superposition,[],[f112,f1190])).
fof(f2741,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(subsumption_resolution,[],[f2740,f109])).
fof(f2740,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(forward_demodulation,[],[f2739,f687])).
fof(f2739,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))),
  inference(forward_demodulation,[],[f2694,f1190])).
fof(f2694,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(superposition,[],[f115,f1190])).
fof(f2744,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(forward_demodulation,[],[f2743,f1190])).
fof(f2743,plain,(
  s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(subsumption_resolution,[],[f2742,f109])).
fof(f2742,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(forward_demodulation,[],[f2695,f687])).
fof(f2695,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(superposition,[],[f116,f1190])).
fof(f1088,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) )),
  inference(forward_demodulation,[],[f1076,f212])).
fof(f1076,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X3))))) != s(X2,g(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X3))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) )),
  inference(superposition,[],[f164,f517])).
fof(f1089,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) )),
  inference(forward_demodulation,[],[f1082,f212])).
fof(f1082,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X3))) != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f164,f517])).
fof(f1085,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) != s(bool,f(s(X12,sK2(X13,X14,bool,X12)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,g(s(X12,sK2(X13,X14,bool,X12)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f164,f188])).
fof(f1083,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s(X4,sK2(X5,X6,bool,X4)))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,g(s(X4,sK2(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f164,f188])).
fof(f1684,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))),s('type.realax.real',X3))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f283,f118])).
fof(f1686,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f283,f118])).
fof(f1986,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X33),s('type.realax.real',X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X33)))) )),
  inference(subsumption_resolution,[],[f1957,f109])).
fof(f1957,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X35) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X33),s('type.realax.real',X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X33)))) )),
  inference(superposition,[],[f191,f689])).
fof(f1780,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20)))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) )),
  inference(forward_demodulation,[],[f1758,f112])).
fof(f1758,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) )),
  inference(superposition,[],[f163,f565])).
fof(f1778,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X16)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X16)))) )),
  inference(forward_demodulation,[],[f1756,f112])).
fof(f1756,plain,(
  ( ! [X17,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X16)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X16)))) )),
  inference(superposition,[],[f162,f565])).
fof(f1775,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X10))),s('type.realax.real',X11)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X10)))) )),
  inference(forward_demodulation,[],[f1753,f112])).
fof(f1753,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X10))),s('type.realax.real',X11)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f161,f565])).
fof(f2001,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f2000,f112])).
fof(f2000,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1999,f112])).
fof(f1999,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1998,f112])).
fof(f1998,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1988,f112])).
fof(f1988,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f573,f565])).
fof(f1784,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30)))) )),
  inference(forward_demodulation,[],[f1764,f112])).
fof(f1764,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X30)))) )),
  inference(superposition,[],[f283,f565])).
fof(f1781,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24))))))),s('type.realax.real',X25))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X24)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X24)))) )),
  inference(forward_demodulation,[],[f1761,f112])).
fof(f1761,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X24)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24))))))),s('type.realax.real',X25))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X24)))) )),
  inference(superposition,[],[f248,f565])).
fof(f2071,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f568,f118])).
fof(f2102,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) )),
  inference(subsumption_resolution,[],[f2081,f735])).
fof(f2081,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) )),
  inference(superposition,[],[f568,f113])).
fof(f2599,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X5)))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f571,f118])).
fof(f2627,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2610,f109])).
fof(f2610,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f571,f188])).
fof(f1989,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f573,f118])).
fof(f1993,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f573])).
fof(f1779,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X18))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) )),
  inference(forward_demodulation,[],[f1757,f112])).
fof(f1757,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X18))))))),s('type.realax.real',X19)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) )),
  inference(superposition,[],[f163,f565])).
fof(f1777,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))))))),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X14))),s('type.realax.real',X15)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X14)))) )),
  inference(forward_demodulation,[],[f1755,f112])).
fof(f1755,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X14))),s('type.realax.real',X15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))))))),s('type.realax.real',X15)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X14)))) )),
  inference(superposition,[],[f162,f565])).
fof(f1776,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))))))),s('type.realax.real',X13)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X12)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X12)))) )),
  inference(forward_demodulation,[],[f1754,f112])).
fof(f1754,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X12)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))))))),s('type.realax.real',X13)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X12)))) )),
  inference(superposition,[],[f161,f565])).
fof(f1891,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1890,f112])).
fof(f1890,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1889,f112])).
fof(f1889,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1888,f112])).
fof(f1888,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1878,f112])).
fof(f1878,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))))),s('type.realax.real',X3)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f572,f565])).
fof(f1783,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))))),s('type.realax.real',X29))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28))),s('type.realax.real',X29)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28)))) )),
  inference(forward_demodulation,[],[f1763,f112])).
fof(f1763,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28))),s('type.realax.real',X29)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X28))))))),s('type.realax.real',X29))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X28)))) )),
  inference(superposition,[],[f283,f565])).
fof(f1782,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X26))),s('type.realax.real',X27)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X26)))) )),
  inference(forward_demodulation,[],[f1762,f112])).
fof(f1762,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X26))),s('type.realax.real',X27)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X26)))) )),
  inference(superposition,[],[f248,f565])).
fof(f1910,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f1894,f112])).
fof(f1894,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f1883,f112])).
fof(f2023,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f2007,f112])).
fof(f2007,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f1993,f112])).
fof(f2003,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1991,f112])).
fof(f1991,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f573,f565])).
fof(f2619,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f2598,f112])).
fof(f2598,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f571,f565])).
fof(f2091,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f2070,f112])).
fof(f2070,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f568,f565])).
fof(f1883,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f572])).
fof(f2325,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X5)))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f569,f118])).
fof(f2349,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real',X4)))) )),
  inference(subsumption_resolution,[],[f2328,f735])).
fof(f2328,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f569,f113])).
fof(f2471,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f570,f118])).
fof(f2492,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2475,f109])).
fof(f2475,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X7) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f570,f188])).
fof(f1893,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1881,f112])).
fof(f1881,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f572,f565])).
fof(f2491,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f2470,f112])).
fof(f2470,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f570,f565])).
fof(f2345,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f2324,f112])).
fof(f2324,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3)))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f569,f565])).
fof(f1882,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4)))))))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f572,f118])).
fof(f1879,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))))),s('type.realax.real',X5)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f572,f118])).
fof(f1927,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f689,f118])).
fof(f1930,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))),s('type.realax.real',X5))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f689,f118])).
fof(f1983,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1929,f112])).
fof(f1929,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f689,f565])).
fof(f2099,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2079,f735])).
fof(f2079,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f568,f212])).
fof(f2738,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))),
  inference(subsumption_resolution,[],[f2737,f109])).
fof(f2737,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))),
  inference(forward_demodulation,[],[f2693,f687])).
fof(f2693,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))),
  inference(superposition,[],[f114,f1190])).
fof(f1981,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f1926,f112])).
fof(f1926,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f689,f565])).
fof(f2127,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f2110,f112])).
fof(f2110,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f2099,f112])).
fof(f2586,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f2569,f112])).
fof(f2569,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f2346,f112])).
fof(f2346,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2326,f735])).
fof(f2326,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f569,f212])).
fof(f2594,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2577,f735])).
fof(f2577,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f2346,f113])).
fof(f2259,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2258])).
fof(f2258,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',X20))) | s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(duplicate_literal_removal,[],[f2234])).
fof(f2234,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',X20))) | s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f532,f517])).
fof(f2265,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2264])).
fof(f2264,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2229])).
fof(f2229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f532,f113])).
fof(f2132,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f2115,f735])).
fof(f2115,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) )),
  inference(superposition,[],[f2099,f113])).
fof(f2126,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))))) )),
  inference(forward_demodulation,[],[f2125,f112])).
fof(f2125,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f2124,f112])).
fof(f2124,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f2107,f112])).
fof(f2107,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f2099,f112])).
fof(f2068,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f2049,f109])).
fof(f2049,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f191,f1858])).
fof(f2028,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f2012,f109])).
fof(f2012,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1993,f188])).
fof(f2022,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))))) )),
  inference(forward_demodulation,[],[f2021,f112])).
fof(f2021,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f2020,f112])).
fof(f2020,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f2004,f112])).
fof(f2004,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f1993,f112])).
fof(f1918,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1902,f109])).
fof(f1902,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,X2) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1883,f188])).
fof(f1858,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(unit_resulting_resolution,[],[f1189,f120])).
fof(f1742,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f565,f112])).
fof(f1741,plain,(
  ( ! [X2] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f565,f118])).
fof(f1769,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f1740,f112])).
fof(f1740,plain,(
  ( ! [X1] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f565,f565])).
fof(f1738,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1724,f109])).
fof(f1724,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f191,f1704])).
fof(f1704,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))),
  inference(subsumption_resolution,[],[f1687,f109])).
fof(f1687,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0)))),
  inference(superposition,[],[f283,f173])).
fof(f1191,plain,(
  ( ! [X14,X13] : (s(bool,X14) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f1176,f109])).
fof(f1176,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,X14) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f191,f687])).
fof(f1189,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(subsumption_resolution,[],[f1165,f109])).
fof(f1165,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(superposition,[],[f114,f687])).
fof(f1079,plain,(
  ( ! [X14,X12,X15,X13] : (s(X14,f(s(bool,X15))) != s(X14,g(s(bool,X15))) | s(fun(bool,X14),X12) = s(fun(bool,X14),X13) | s(bool,'T_0') = s(bool,sK2(X12,X13,X14,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f164,f188])).
fof(f1077,plain,(
  ( ! [X6,X4,X7,X5] : (s(X6,f(s(bool,'T_0'))) != s(X6,g(s(bool,'T_0'))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | s(bool,X7) = s(bool,sK2(X4,X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f164,f188])).
fof(f962,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(unit_resulting_resolution,[],[f959,f161])).
fof(f924,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f919,f163])).
fof(f6570,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553,f6560,f6564,f6566,f6569])).
fof(f6567,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553,f6560,f6564,f6566])).
fof(f6565,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553,f6560,f6564])).
fof(f6561,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553,f6560])).
fof(f6554,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548,f6553])).
fof(f6549,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544,f6548])).
fof(f6545,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514,f6544])).
fof(f6541,plain,(
  $false | ($spl11 | $spl18)),
  inference(subsumption_resolution,[],[f6270,f6314])).
fof(f6270,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | ($spl11 | $spl18)),
  inference(backward_demodulation,[],[f6172,f5023])).
fof(f5023,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl11),
  inference(unit_resulting_resolution,[],[f4965,f119])).
fof(f6518,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f6517,f109])).
fof(f6517,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f6516,f6314])).
fof(f6516,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6239,f108])).
fof(f6239,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f2023])).
fof(f6515,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509,f6514])).
fof(f6510,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503,f6509])).
fof(f6504,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500,f6503])).
fof(f6501,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493,f6500])).
fof(f6497,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f6496,f109])).
fof(f6496,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f6495,f6314])).
fof(f6495,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f6228,f108])).
fof(f6228,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))) ) | $spl18),
  inference(backward_demodulation,[],[f6172,f1910])).
fof(f6494,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488,f6493])).
fof(f6489,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482,f6488])).
fof(f6483,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479,f6482])).
fof(f6480,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476,f6479])).
fof(f6477,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472,f6476])).
fof(f6473,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468,f6472])).
fof(f6469,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447,f6468])).
fof(f6448,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439,f6447])).
fof(f6440,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431,f6439])).
fof(f6432,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429,f6431])).
fof(f6430,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409,f6429])).
fof(f6410,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406,f6409])).
fof(f6407,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404,f6406])).
fof(f6405,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401,f6404])).
fof(f6402,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390,f6401])).
fof(f6391,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386,f6390])).
fof(f6387,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382,f6386])).
fof(f6383,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370,f6382])).
fof(f6379,plain,(
  $false | ($spl4 | $spl18)),
  inference(global_subsumption,[],[f6378,f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370])).
fof(f6378,plain,(
  ( ! [X19] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | ($spl4 | $spl18)),
  inference(subsumption_resolution,[],[f6343,f109])).
fof(f6343,plain,(
  ( ! [X19] : (~p(s(bool,'T_0')) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | ($spl4 | $spl18)),
  inference(backward_demodulation,[],[f6314,f2774])).
fof(f6377,plain,(
  $false | ($spl8 | $spl18)),
  inference(global_subsumption,[],[f6376,f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370])).
fof(f6376,plain,(
  ( ! [X21] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f6340,f109])).
fof(f6340,plain,(
  ( ! [X21] : (~p(s(bool,'T_0')) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | ($spl8 | $spl18)),
  inference(backward_demodulation,[],[f6314,f2781])).
fof(f6371,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365,f6370])).
fof(f6367,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f6318,f109])).
fof(f6318,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(backward_demodulation,[],[f6314,f735])).
fof(f6366,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f692,f919,f924,f959,f962,f1077,f1079,f1189,f1191,f1704,f1738,f1769,f1741,f1742,f1858,f1918,f2022,f2028,f2068,f2126,f2132,f2265,f2259,f2594,f2773,f2780,f2346,f2586,f2127,f1981,f2738,f2099,f1983,f1930,f1927,f1879,f1882,f2345,f2491,f1893,f2492,f2471,f2349,f2325,f572,f570,f569,f1883,f2091,f2619,f2003,f2023,f1910,f1782,f1783,f1891,f1776,f1777,f1779,f1993,f1989,f2627,f2599,f2102,f2071,f573,f571,f568,f1781,f1784,f2001,f1775,f1778,f1780,f1986,f1705,f1686,f1684,f689,f672,f670,f283,f248,f163,f162,f161,f1083,f1085,f1089,f1088,f735,f212,f124,f123,f2744,f2741,f2689,f2257,f2255,f1985,f1984,f1774,f1771,f1770,f1748,f1744,f1516,f1515,f1504,f1503,f1502,f1315,f1314,f1306,f1190,f1156,f1147,f1146,f1145,f1039,f984,f979,f974,f967,f929,f762,f758,f751,f738,f736,f694,f117,f116,f115,f114,f4446,f4653,f4655,f4662,f4718,f4724,f4837,f5109,f5047,f5293,f5280,f5297,f5298,f5490,f5434,f5494,f5445,f5446,f5447,f5526,f5511,f5530,f5531,f5919,f5595,f5924,f5602,f5908,f5896,f5932,f5648,f5937,f5655,f5944,f5945,f5946,f5947,f5948,f5955,f5866,f5789,f5790,f5995,f5998,f6002,f6005,f5843,f6040,f6041,f6021,f6049,f6024,f6051,f6052,f6053,f6057,f6101,f6067,f6106,f6074,f6153,f6115,f6167,f6122,f6145,f6172,f6175,f6176,f6177,f6314,f6365])).
fof(f6313,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f6178,f213])).
fof(f6178,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK0),s('type.realax.real',sK0))) | $spl18),
  inference(backward_demodulation,[],[f6172,f173])).
fof(f6172,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl18),
  inference(cnf_transformation,[],[f6172_D])).
fof(f6172_D,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6174,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl20),
  inference(cnf_transformation,[],[f6174_D])).
fof(f6174_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6145,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(trivial_inequality_removal,[],[f6133])).
fof(f6133,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f6067,f929])).
fof(f6003,plain,(
  ( ! [X35] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35))))) ) | $spl16),
  inference(cnf_transformation,[],[f6003_D])).
fof(f6003_D,plain,(
  ( ! [X35] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f686,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl2),
  inference(cnf_transformation,[],[f686_D])).
fof(f686_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f6002,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35)))))) )),
  inference(subsumption_resolution,[],[f5806,f119])).
fof(f5806,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X35)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f117,f1705])).
fof(f5996,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) ) | $spl14),
  inference(cnf_transformation,[],[f5996_D])).
fof(f5996_D,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5995,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) )),
  inference(subsumption_resolution,[],[f5796,f119])).
fof(f5796,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f248,f1705])).
fof(f4965,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl11),
  inference(cnf_transformation,[],[f4965_D])).
fof(f4965_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f4967,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl13),
  inference(cnf_transformation,[],[f4967_D])).
fof(f4967_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f4963,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl7),
  inference(subsumption_resolution,[],[f4962,f735])).
fof(f4962,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl7),
  inference(forward_demodulation,[],[f4893,f212])).
fof(f4893,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl7),
  inference(superposition,[],[f117,f4115])).
fof(f4115,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f2775,f118])).
fof(f2775,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl7),
  inference(cnf_transformation,[],[f2775_D])).
fof(f2775_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3016,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972,f2978,f2809,f3007,f2827,f2829])).
fof(f2829,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) | $spl6),
  inference(backward_demodulation,[],[f2776,f2744])).
fof(f2827,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | $spl6),
  inference(backward_demodulation,[],[f2776,f2689])).
fof(f3007,plain,(
  ( ! [X7] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3006,f109])).
fof(f3006,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) ) | $spl6),
  inference(forward_demodulation,[],[f2821,f2835])).
fof(f2821,plain,(
  ( ! [X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X7)))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f1771])).
fof(f2809,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl6),
  inference(backward_demodulation,[],[f2776,f1190])).
fof(f2978,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2898,f109])).
fof(f2898,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2795])).
fof(f2972,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2895,f109])).
fof(f2895,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2791])).
fof(f2970,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2894,f109])).
fof(f2894,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2793])).
fof(f2968,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2892,f109])).
fof(f2892,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1879])).
fof(f2966,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2891,f109])).
fof(f2891,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1882])).
fof(f2961,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2888,f109])).
fof(f2888,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1893])).
fof(f2956,plain,(
  ( ! [X10,X11] : (s('type.realax.real',X10) != s('type.realax.real',X11) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2884,f109])).
fof(f2884,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.realax.real',X10) != s('type.realax.real',X11) | s('type.realax.real',X10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f572])).
fof(f2952,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2881,f109])).
fof(f2881,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1883])).
fof(f2947,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2878,f109])).
fof(f2878,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real',X3) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2003])).
fof(f2945,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2877,f109])).
fof(f2877,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2023])).
fof(f2943,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2876,f109])).
fof(f2876,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2022])).
fof(f2941,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2875,f109])).
fof(f2875,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1910])).
fof(f2939,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2872,f109])).
fof(f2872,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1891])).
fof(f2935,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2869,f109])).
fof(f2869,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s('type.realax.real',X19) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X18)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1779])).
fof(f2933,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2868,f109])).
fof(f2868,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1993])).
fof(f2931,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2867,f109])).
fof(f2867,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',X4) != s('type.realax.real',X5) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))))) | s('type.realax.real',X4) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1989])).
fof(f2926,plain,(
  ( ! [X12,X13] : (s('type.realax.real',X12) != s('type.realax.real',X13) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(subsumption_resolution,[],[f2863,f109])).
fof(f2863,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s('type.realax.real',X12) != s('type.realax.real',X13) | s('type.realax.real',X12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f573])).
fof(f2922,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2858,f109])).
fof(f2858,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) | s('type.realax.real',X3) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X2)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f2001])).
fof(f2918,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl6),
  inference(subsumption_resolution,[],[f2855,f109])).
fof(f2855,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s('type.realax.real',X21) != s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X20)))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f1780])).
fof(f2914,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f2846,f109])).
fof(f2846,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f163])).
fof(f2909,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f2836,f2835])).
fof(f2836,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f2835,f123])).
fof(f2835,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f2798,f1738])).
fof(f2798,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f762])).
fof(f2795,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))))) | s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f738])).
fof(f2794,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))) | s('type.realax.real',X4) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X4)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f736])).
fof(f2793,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f694])).
fof(f2831,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl6),
  inference(forward_demodulation,[],[f2792,f2776])).
fof(f2792,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f117])).
fof(f2791,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f116])).
fof(f2830,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2790,f2776])).
fof(f2790,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1(X0)),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0))))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f115])).
fof(f2789,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1(X0)))))) ) | $spl6),
  inference(backward_demodulation,[],[f2776,f114])).
fof(f3014,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972,f2978,f2809,f3007,f2827])).
fof(f3008,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972,f2978,f2809,f3007])).
fof(f2988,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972,f2978,f2809])).
fof(f2980,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2800,f2835])).
fof(f2800,plain,(
  s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | $spl6),
  inference(backward_demodulation,[],[f2776,f967])).
fof(f2979,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972,f2978])).
fof(f2973,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970,f2972])).
fof(f2971,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968,f2970])).
fof(f2969,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966,f2968])).
fof(f2967,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961,f2966])).
fof(f2962,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956,f2961])).
fof(f2957,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952,f2956])).
fof(f2953,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947,f2952])).
fof(f2948,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945,f2947])).
fof(f2946,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943,f2945])).
fof(f2944,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941,f2943])).
fof(f2942,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939,f2941])).
fof(f2940,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935,f2939])).
fof(f2936,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933,f2935])).
fof(f2934,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931,f2933])).
fof(f2932,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926,f2931])).
fof(f2927,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922,f2926])).
fof(f2923,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918,f2922])).
fof(f2919,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914,f2918])).
fof(f2915,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909,f2914])).
fof(f2911,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2839,f109])).
fof(f2839,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f2835,f735])).
fof(f2910,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f118,f120,f119,f164,f171,f173,f188,f191,f211,f213,f259,f312,f517,f370,f532,f577,f565,f684,f687,f706,f700,f735,f736,f116,f115,f114,f569,f570,f572,f568,f571,f573,f689,f672,f670,f283,f248,f163,f162,f161,f117,f124,f123,f212,f694,f738,f692,f751,f758,f762,f919,f929,f924,f959,f967,f974,f962,f979,f984,f1039,f1088,f1077,f1079,f1089,f1083,f1085,f1145,f1146,f1147,f1156,f1189,f1190,f1191,f1314,f1306,f1315,f1515,f1502,f1503,f1504,f1516,f1684,f1686,f1704,f1705,f1738,f1769,f1741,f1742,f1744,f1748,f1770,f1771,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1858,f1891,f1879,f1893,f1882,f1883,f1910,f1918,f1981,f1927,f1983,f1930,f1984,f1985,f1986,f2001,f1989,f2003,f1993,f2022,f2023,f2028,f2068,f2091,f2071,f2099,f2102,f2126,f2127,f2132,f2265,f2259,f2257,f2255,f2345,f2325,f2346,f2349,f2491,f2471,f2492,f2586,f2594,f2619,f2599,f2627,f2689,f2738,f2741,f2744,f2773,f2780,f2776,f2789,f2830,f2791,f2831,f2793,f2794,f2795,f2835,f2909])).
fof(f2834,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2833,f171])).
fof(f2833,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl6),
  inference(forward_demodulation,[],[f2797,f1704])).
fof(f2797,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | $spl6),
  inference(backward_demodulation,[],[f2776,f758])).
fof(f2832,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f2796,f1704])).
fof(f2796,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK0))) | $spl6),
  inference(backward_demodulation,[],[f2776,f751])).
fof(f2781,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(cnf_transformation,[],[f2781_D])).
fof(f2781_D,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2776,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl6),
  inference(cnf_transformation,[],[f2776_D])).
fof(f2776_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2780,plain,(
  ( ! [X21] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f2779,f1190])).
fof(f2779,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f2778,f112])).
fof(f2778,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f2722,f112])).
fof(f2722,plain,(
  ( ! [X21] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f573,f1190])).
fof(f2774,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl4),
  inference(cnf_transformation,[],[f2774_D])).
fof(f2774_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2773,plain,(
  ( ! [X19] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f2772,f1190])).
fof(f2772,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f2771,f112])).
fof(f2771,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f2720,f112])).
fof(f2720,plain,(
  ( ! [X19] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),s('type.realax.real',X19)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f572,f1190])).
fof(f734,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f733,f732])).
fof(f732,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f715,f109])).
fof(f715,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl2),
  inference(backward_demodulation,[],[f707,f163])).
fof(f707,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | $spl2),
  inference(global_subsumption,[],[f108,f107,f109,f111,f112,f113,f116,f115,f114,f117,f118,f120,f119,f124,f123,f161,f162,f163,f164,f171,f173,f188,f191,f211,f213,f212,f248,f259,f283,f312,f517,f370,f532,f577,f565,f568,f569,f570,f571,f572,f573,f670,f672,f684,f687,f689,f686,f703,f704,f705,f706])).
fof(f705,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) ) | $spl2),
  inference(subsumption_resolution,[],[f695,f686])).
fof(f695,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f116,f517])).
fof(f704,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(subsumption_resolution,[],[f694,f686])).
fof(f703,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f702,f112])).
fof(f702,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1('const.realax.inv_1'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f693,f686])).
fof(f733,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f728,f109])).
fof(f728,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK1(X0))))) ) | $spl2),
  inference(backward_demodulation,[],[f707,f116])).
fof(f731,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f711,f109])).
fof(f711,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f707,f111])).
fof(f730,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f729,f686])).
fof(f729,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f709,f707])).
fof(f709,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(backward_demodulation,[],[f707,f123])).
fof(f701,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f692,f686])).
fof(f685,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl0),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f684,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(subsumption_resolution,[],[f674,f119])).
fof(f674,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f248,f517])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 177
% Passive clauses: 594
% Generated clauses: 8488
% Final active clauses: 69
% Final passive clauses: 177
% Input formulas: 29
% Initial clauses: 58
% 
% Duplicate literals: 356
% Trivial inequalities: 104
% Fw subsumption resolutions: 564
% Fw demodulations: 1586
% Bw demodulations: 324
% Condensations: 460
% Global subsumptions: 61
% 
% Simple tautologies: 158
% Equational tautologies: 174
% Forward subsumptions: 4344
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 938
% Forward superposition: 2068
% Backward superposition: 1523
% Self superposition: 47
% Equality factoring: 40
% Equality resolution: 30
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 756
% SAT solver unit clauses: 126
% SAT solver binary clauses: 229
% 
% Sat splits: 13
% Sat splitting refutations: 73
% 
% Memory used [KB]: 3837
% Time elapsed: 0.243 s
% ------------------------------
% ------------------------------
% Success in time 1.559 s
