Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (47519)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 152
% Passive clauses: 3671
% Generated clauses: 20289
% Final active clauses: 131
% Final passive clauses: 3192
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 694
% Trivial inequalities: 280
% Fw subsumption resolutions: 3277
% Fw demodulations: 3479
% Bw demodulations: 17
% 
% Simple tautologies: 429
% Equational tautologies: 45
% Forward subsumptions: 7601
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 1198
% Factoring: 7
% Forward superposition: 6100
% Backward superposition: 5875
% Self superposition: 172
% Equality factoring: 102
% Equality resolution: 35
% 
% Split clauses: 490
% Split components: 1009
% SAT solver clauses: 1065
% SAT solver unit clauses: 23
% SAT solver binary clauses: 169
% 
% Sat splits: 490
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1059
% TWLsolver calls for satisfiability: 216
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (47671)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 100
% Passive clauses: 952
% Generated clauses: 8635
% Final active clauses: 92
% Final passive clauses: 833
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 491
% Trivial inequalities: 102
% Fw subsumption resolutions: 849
% Fw demodulations: 477
% Bw demodulations: 21
% Condensations: 652
% Global subsumptions: 17
% 
% Simple tautologies: 198
% Equational tautologies: 36
% Forward subsumptions: 3240
% 
% Binary resolution: 632
% Factoring: 3
% Forward superposition: 2319
% Backward superposition: 1784
% Self superposition: 80
% Equality factoring: 68
% Equality resolution: 17
% 
% Split clauses: 133
% Split components: 277
% SAT solver clauses: 1694
% SAT solver unit clauses: 112
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 34
% 
% Sat splits: 133
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1626
% TWLsolver calls for satisfiability: 460
% 
% Memory used [KB]: 7675
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (47735)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 465
% Passive clauses: 3132
% Generated clauses: 157402
% Final active clauses: 145
% Final passive clauses: 603
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 7196
% Trivial inequalities: 671
% Fw subsumption resolutions: 14661
% Fw demodulations: 21813
% Bw demodulations: 1217
% Condensations: 24589
% Global subsumptions: 281
% 
% Simple tautologies: 1012
% Equational tautologies: 3445
% Forward subsumptions: 77806
% Fw demodulations to eq. taut.: 57
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 15520
% Factoring: 11
% Forward superposition: 32610
% Backward superposition: 32333
% Self superposition: 1023
% Equality factoring: 544
% Equality resolution: 72
% 
% Split clauses: 288
% Split components: 581
% SAT solver clauses: 4510
% SAT solver unit clauses: 53
% SAT solver binary clauses: 487
% 
% Sat splits: 288
% Sat splitting refutations: 519
% 
% Memory used [KB]: 79316
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (47904)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1879
% Passive clauses: 36280
% Generated clauses: 137863
% Final active clauses: 209
% Final passive clauses: 19108
% Input formulas: 37
% Initial clauses: 100
% 
% Selected by SInE selection: 37
% SInE iterations: 1
% 
% Duplicate literals: 2936
% Trivial inequalities: 2899
% Fw subsumption resolutions: 19902
% Fw demodulations: 14153
% Bw demodulations: 1003
% 
% Simple tautologies: 4565
% Equational tautologies: 21
% Forward subsumptions: 53145
% Bw demodulations to eq. taut.: 14
% 
% Binary resolution: 17619
% Factoring: 3
% Forward superposition: 45642
% Backward superposition: 27821
% Self superposition: 66
% Equality factoring: 41
% Equality resolution: 402
% 
% Split clauses: 5584
% Split components: 11850
% SAT solver clauses: 6188
% SAT solver unit clauses: 105
% SAT solver binary clauses: 530
% 
% Sat splits: 5584
% Sat splitting refutations: 575
% 
% Memory used [KB]: 44775
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (48177)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 83
% Passive clauses: 1934
% Generated clauses: 11035
% Final active clauses: 60
% Final passive clauses: 1194
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 460
% Trivial inequalities: 326
% Fw subsumption resolutions: 1568
% Bw demodulations: 572
% Global subsumptions: 16
% 
% Simple tautologies: 526
% Equational tautologies: 61
% Forward subsumptions: 5977
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 159
% Forward superposition: 3679
% Backward superposition: 3295
% Self superposition: 112
% Equality factoring: 62
% Equality resolution: 14
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 2042
% SAT solver unit clauses: 56
% SAT solver binary clauses: 90
% 
% Sat splits: 12
% Sat splitting refutations: 78
% 
% Memory used [KB]: 6780
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (48210)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 687
% Passive clauses: 3909
% Generated clauses: 29437
% Final active clauses: 681
% Final passive clauses: 3133
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 1265
% Trivial inequalities: 1114
% Fw subsumption resolutions: 1745
% Bw subsumption resolutions: 4
% Fw demodulations: 306
% Bw demodulations: 8
% Condensations: 1279
% 
% Simple tautologies: 1775
% Forward subsumptions: 17979
% Backward subsumptions: 134
% 
% Binary resolution: 7820
% Forward superposition: 10926
% Backward superposition: 5030
% Equality resolution: 9
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 3913
% SAT solver unit clauses: 9
% SAT solver binary clauses: 120
% 
% Sat splits: 2
% 
% Memory used [KB]: 20084
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f4651,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f155,f259,f263,f265,f262,f267,f261,f269,f260,f271,f272,f276,f274,f277,f281,f279,f284,f291,f302,f301,f300,f299,f298,f297,f296,f306,f304,f308,f309,f311,f312,f314,f416,f307,f528,f527,f526,f525,f524,f523,f522,f532,f530,f534,f535,f537,f640,f533,f768,f767,f766,f765,f764,f763,f762,f772,f770,f774,f775,f777,f809,f915,f914,f913,f912,f911,f910,f909,f919,f917,f921,f922,f924,f960,f1039,f1188,f773,f1100,f1190,f1101,f1345,f920,f1703,f1362,f1364,f1366,f1368,f1370,f1372,f1374,f1376,f1378,f1380,f1382,f1384,f1386,f1713,f1714,f1391,f1393,f1399,f1401,f1403,f1407,f1411,f1715,f1716,f1425,f1429,f1717,f1718,f1719,f1720,f1721,f1722,f1723,f1724,f1725,f1726,f1727,f1728,f1729,f1730,f1731,f1732,f1733,f1734,f1735,f1736,f1737,f1738,f1739,f1740,f1741,f1742,f1743,f1744,f1745,f1746,f1747,f1748,f1749,f1750,f1751,f1752,f1977,f1986,f1987,f1988,f1989,f1990,f1991,f1992,f1993,f1994,f1995,f1996,f1997,f1998,f1999,f2000,f2001,f2002,f2003,f2004,f2005,f1254,f1040,f1253,f2009,f2010,f285,f2013,f292,f2037,f2036,f2035,f2034,f2033,f2032,f2031,f2041,f2039,f286,f2043,f2052,f2051,f2050,f2049,f2048,f2047,f2046,f2056,f2054,f293,f2919,f2605,f2607,f2609,f2611,f2613,f2615,f2617,f2619,f2621,f2938,f2939,f2631,f2635,f2639,f2940,f2941,f2653,f2657,f2942,f2943,f2944,f2945,f2946,f2947,f2948,f2949,f2950,f2951,f2952,f2953,f2954,f2955,f2956,f2957,f2958,f2959,f2960,f2961,f2962,f2963,f2964,f2965,f2966,f2967,f2968,f2969,f2970,f2971,f2972,f2973,f2974,f2975,f2976,f2977,f3194,f3203,f3204,f3205,f3206,f3207,f3208,f3209,f3210,f3211,f3212,f3213,f3214,f3215,f3216,f3217,f3218,f3219,f3220,f3221,f3222,f3223,f3224,f3226,f295,f3229,f288,f3264,f886,f1189,f3409,f1041,f2008,f3586,f3584,f3583,f3582,f3575,f3573,f3572,f3571,f3570,f3569,f3568,f3567,f3566,f3565,f3564,f3563,f3561,f3559,f3558,f3557,f3555,f3553,f3545,f3544,f3543,f3542,f3610,f3539,f3538,f3620,f3593,f3594,f3669,f3668,f3667,f3660,f3658,f3657,f3656,f3655,f3654,f3653,f3652,f3651,f3650,f3649,f3648,f3647,f3645,f3643,f3642,f3641,f3639,f3637,f3629,f3628,f3682,f3627,f3626,f3670,f3671,f3691,f3690,f3743,f3741,f3740,f3739,f3732,f3730,f3729,f3728,f3727,f3726,f3725,f3724,f3723,f3722,f3721,f3720,f3718,f3716,f3715,f3714,f3712,f3710,f3702,f3701,f3700,f3699,f3766,f3696,f3695,f3775,f3750,f3751,f3773,f3826,f3824,f3823,f3822,f3815,f3813,f3812,f3811,f3810,f3809,f3808,f3807,f3806,f3805,f3804,f3803,f3801,f3799,f3798,f3797,f3795,f3793,f3785,f3784,f3783,f3782,f3849,f3779,f3778,f3858,f3833,f3834,f3856,f3911,f3909,f3908,f3907,f3900,f3898,f3897,f3896,f3895,f3894,f3893,f3892,f3891,f3890,f3889,f3888,f3886,f3884,f3883,f3882,f3880,f3878,f3870,f3869,f3868,f3867,f3936,f3864,f3863,f3945,f3920,f3921,f3943,f4087,f4089,f4003,f883,f3952,f3953,f3954,f3618,f3955,f3956,f3957,f3958,f884,f4004,f885,f4141,f4139,f4135,f4133,f4125,f4124,f4119,f4118,f4200,f4177,f4178,f3681,f4256,f4254,f4253,f4252,f4251,f4243,f4241,f4240,f4239,f4238,f4237,f4236,f4235,f4234,f4233,f4232,f4231,f4230,f4228,f4226,f4225,f4224,f4222,f4220,f4212,f4211,f4210,f4209,f4283,f4206,f4205,f4293,f4295,f4296,f1192,f4266,f4267,f4291,f4628,f4627,f4626,f4625,f4617,f4615,f4614,f4613,f4612,f4611,f4610,f4609,f4608,f4607,f4606,f4605,f4604,f4603,f4601,f4599,f4598,f4597,f4595,f4593,f4585,f4584,f4290,f4583,f4582,f4649,f888,f4630,f4631,f4650,f4282])).
fof(f4282,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl102),
  inference(cnf_transformation,[],[f4282_D])).
fof(f4282_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4650,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(global_subsumption,[],[f155,f1190,f4628,f4617,f4615,f4614,f4601,f4599,f4598,f4597,f4595,f4593,f4585,f4584,f4583,f4582,f4581,f4580,f4579,f4630,f4631,f4638,f4640])).
fof(f4640,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(superposition,[],[f1190,f162])).
fof(f162,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1',aHL_BOOL_CASES)).
fof(f4638,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl62),
  inference(resolution,[],[f1190,f233])).
fof(f233,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f4579,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f228])).
fof(f228,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X15,X11,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X15),s(fun(X6,bool),X11)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X15)))) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.pred_set.IN_UNION')).
fof(f4580,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f229])).
fof(f229,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f4581,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X15,X11,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X15),s(fun(X6,bool),X11)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X15)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.pred_set.IN_DIFF')).
fof(f4631,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(resolution,[],[f1190,f168])).
fof(f168,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f4630,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(resolution,[],[f1190,f166])).
fof(f166,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f888,plain,(
  p(s(bool,'F_0')) | $spl50),
  inference(cnf_transformation,[],[f888_D])).
fof(f888_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4649,plain,(
  p(s(bool,'F_0')) | ($spl62 | $spl74 | $spl78)),
  inference(duplicate_literal_removal,[],[f4648])).
fof(f4648,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | ($spl62 | $spl74 | $spl78)),
  inference(forward_demodulation,[],[f4647,f3681])).
fof(f4647,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl62 | $spl74)),
  inference(forward_demodulation,[],[f4629,f3618])).
fof(f4629,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl62),
  inference(resolution,[],[f1190,f227])).
fof(f227,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f4582,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f166])).
fof(f4583,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f168])).
fof(f4290,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl105),
  inference(cnf_transformation,[],[f4290_D])).
fof(f4290_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f4584,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f169])).
fof(f169,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f4585,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1190,f171])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f4593,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f156,f1190,f233])).
fof(f156,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1',aHL_TRUTH)).
fof(f4595,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f156,f1190,f233])).
fof(f4597,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f1190,f234])).
fof(f234,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9,X10] : (s(bool,X9) = s(bool,X10) <=> ((p(s(bool,X9)) => p(s(bool,X10))) & (p(s(bool,X10)) => p(s(bool,X9)))))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.EQ_IMP_THM')).
fof(f161,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1',aHL_FALSITY)).
fof(f4598,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f1190,f235])).
fof(f235,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f149])).
fof(f4599,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f156,f1190,f239])).
fof(f239,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f149])).
fof(f4601,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f156,f1190,f239])).
fof(f4603,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f1190,f248])).
fof(f248,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7,X13,X8,X14] : ((s(bool,X7) = s(bool,X13) & (p(s(bool,X13)) => s(bool,X8) = s(bool,X14))) => ((p(s(bool,X7)) => p(s(bool,X8))) <=> (p(s(bool,X13)) => p(s(bool,X14)))))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CONG')).
fof(f275,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4604,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f1190,f248])).
fof(f280,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f280_D])).
fof(f280_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4605,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f923,f1190,f248])).
fof(f923,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl56),
  inference(cnf_transformation,[],[f923_D])).
fof(f923_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4606,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl20 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f287,f1190,f248])).
fof(f287,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl20),
  inference(cnf_transformation,[],[f287_D])).
fof(f287_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4607,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f1190,f250])).
fof(f250,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f4608,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f1190,f250])).
fof(f4609,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl18 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f286,f1190,f250])).
fof(f4610,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f1190,f252])).
fof(f252,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f4611,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f1190,f252])).
fof(f4612,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f923,f1190,f252])).
fof(f4613,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl20 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f161,f287,f1190,f252])).
fof(f4614,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f161,f156,f1190,f253])).
fof(f253,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f154])).
fof(f4615,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f161,f1190,f1190,f253])).
fof(f4617,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f161,f1190,f1190,f253])).
fof(f4625,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f1190,f254])).
fof(f254,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f4626,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f1190,f254])).
fof(f4627,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl18 | $spl62)),
  inference(unit_resulting_resolution,[],[f161,f156,f286,f1190,f254])).
fof(f4628,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f161,f161,f1190,f257])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f251])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f4291,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl104),
  inference(cnf_transformation,[],[f4291_D])).
fof(f4291_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4267,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(resolution,[],[f1189,f172])).
fof(f172,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f4266,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(resolution,[],[f1189,f170])).
fof(f170,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f1192,plain,(
  ~p(s(bool,'T_0')) | $spl65),
  inference(cnf_transformation,[],[f1192_D])).
fof(f1192_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f4296,plain,(
  ~p(s(bool,'T_0')) | ($spl63 | $spl72)),
  inference(forward_demodulation,[],[f4264,f3609])).
fof(f3609,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl72),
  inference(cnf_transformation,[],[f3609_D])).
fof(f3609_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4264,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl63),
  inference(resolution,[],[f1189,f228])).
fof(f4295,plain,(
  ~p(s(bool,'T_0')) | ($spl63 | $spl72)),
  inference(forward_demodulation,[],[f4203,f3609])).
fof(f4203,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f228])).
fof(f4293,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl107),
  inference(cnf_transformation,[],[f4293_D])).
fof(f4293_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f4205,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f1189,f227])).
fof(f4206,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f1189,f227])).
fof(f4283,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl103),
  inference(cnf_transformation,[],[f4283_D])).
fof(f4283_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f4209,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f165])).
fof(f165,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f105])).
fof(f4210,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f167])).
fof(f167,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f106])).
fof(f4211,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f170])).
fof(f4212,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1189,f172])).
fof(f4220,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f1189,f230])).
fof(f230,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f95])).
fof(f4222,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f1189,f230])).
fof(f4224,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f156,f1189,f234])).
fof(f4225,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f156,f1189,f235])).
fof(f4226,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f1189,f236])).
fof(f236,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f149])).
fof(f4228,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f1189,f236])).
fof(f4230,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1189,f248])).
fof(f4231,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1189,f248])).
fof(f4232,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl44 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1189,f248])).
fof(f776,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl44),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4233,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1189,f248])).
fof(f4234,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1189,f250])).
fof(f4235,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1189,f250])).
fof(f4236,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl42 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f774,f1189,f250])).
fof(f4237,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1189,f252])).
fof(f4238,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1189,f252])).
fof(f4239,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl44 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1189,f252])).
fof(f4240,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1189,f252])).
fof(f4241,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f1189,f253])).
fof(f4243,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f156,f1189,f156,f1189,f253])).
fof(f4251,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1189,f254])).
fof(f4252,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1189,f254])).
fof(f4253,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl42 | $spl63)),
  inference(unit_resulting_resolution,[],[f156,f156,f774,f1189,f254])).
fof(f4254,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f161,f156,f1189,f257])).
fof(f4256,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f156,f1189,f1189,f257])).
fof(f3681,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl78),
  inference(cnf_transformation,[],[f3681_D])).
fof(f3681_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4178,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(resolution,[],[f885,f172])).
fof(f4177,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(resolution,[],[f885,f170])).
fof(f4200,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl101),
  inference(cnf_transformation,[],[f4200_D])).
fof(f4200_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f4118,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl49),
  inference(unit_resulting_resolution,[],[f885,f885,f227])).
fof(f4119,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl49),
  inference(unit_resulting_resolution,[],[f885,f885,f227])).
fof(f4124,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f885,f170])).
fof(f4125,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f885,f172])).
fof(f4133,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f161,f885,f230])).
fof(f4135,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f161,f885,f230])).
fof(f4139,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f161,f885,f236])).
fof(f4141,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl49),
  inference(unit_resulting_resolution,[],[f161,f885,f236])).
fof(f885,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl49),
  inference(cnf_transformation,[],[f885_D])).
fof(f885_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f4004,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl42),
  inference(superposition,[],[f225,f774])).
fof(f884,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f3958,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f156,f3610,f233])).
fof(f3957,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f156,f3610,f239])).
fof(f3956,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f156,f3610,f233])).
fof(f3955,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f156,f3610,f239])).
fof(f3618,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl74),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3954,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl73),
  inference(unit_resulting_resolution,[],[f3610,f162])).
fof(f3953,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f3610,f166])).
fof(f3952,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(unit_resulting_resolution,[],[f3610,f168])).
fof(f883,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl46),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4003,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(superposition,[],[f224,f774])).
fof(f224,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f4089,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl80)),
  inference(duplicate_literal_removal,[],[f4088])).
fof(f4088,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl80)),
  inference(forward_demodulation,[],[f4004,f3690])).
fof(f4087,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl74)),
  inference(forward_demodulation,[],[f4003,f3618])).
fof(f3943,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl96),
  inference(cnf_transformation,[],[f3943_D])).
fof(f3943_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3921,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(resolution,[],[f2008,f172])).
fof(f3920,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(resolution,[],[f2008,f170])).
fof(f3945,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl99),
  inference(cnf_transformation,[],[f3945_D])).
fof(f3945_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f3863,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f2008,f227])).
fof(f3864,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f2008,f227])).
fof(f3936,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl95),
  inference(cnf_transformation,[],[f3936_D])).
fof(f3936_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f3867,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f165])).
fof(f3868,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f167])).
fof(f3869,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f170])).
fof(f3870,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f2008,f172])).
fof(f3878,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f2008,f230])).
fof(f3880,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f2008,f230])).
fof(f3882,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f156,f2008,f234])).
fof(f3883,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f156,f2008,f235])).
fof(f3884,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f2008,f236])).
fof(f3886,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f2008,f236])).
fof(f3888,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f2008,f248])).
fof(f3889,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f2008,f248])).
fof(f3890,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl44 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f2008,f248])).
fof(f3891,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f2008,f248])).
fof(f3892,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f2008,f250])).
fof(f3893,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f2008,f250])).
fof(f3894,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl12 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f2008,f252])).
fof(f3895,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl16 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f2008,f252])).
fof(f3896,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl44 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f2008,f252])).
fof(f3897,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl56 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f2008,f252])).
fof(f3898,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f2008,f253])).
fof(f3900,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f156,f2008,f156,f2008,f253])).
fof(f3907,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl10 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f2008,f254])).
fof(f3908,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl14 | $spl67)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f2008,f254])).
fof(f3909,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f161,f156,f2008,f257])).
fof(f3911,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl67),
  inference(unit_resulting_resolution,[],[f156,f2008,f2008,f257])).
fof(f3856,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl90),
  inference(cnf_transformation,[],[f3856_D])).
fof(f3856_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3834,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(resolution,[],[f1040,f172])).
fof(f3833,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(resolution,[],[f1040,f170])).
fof(f3858,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl93),
  inference(cnf_transformation,[],[f3858_D])).
fof(f3858_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f3778,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f1040,f227])).
fof(f3779,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f1040,f227])).
fof(f3849,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl89),
  inference(cnf_transformation,[],[f3849_D])).
fof(f3849_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f3782,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f165])).
fof(f3783,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f167])).
fof(f3784,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f170])).
fof(f3785,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f1040,f172])).
fof(f3793,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f1040,f230])).
fof(f3795,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f1040,f230])).
fof(f3797,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f156,f1040,f234])).
fof(f3798,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f156,f1040,f235])).
fof(f3799,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f1040,f236])).
fof(f3801,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f1040,f236])).
fof(f3803,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1040,f248])).
fof(f3804,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1040,f248])).
fof(f3805,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl44 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1040,f248])).
fof(f3806,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl56 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1040,f248])).
fof(f3807,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1040,f250])).
fof(f3808,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1040,f250])).
fof(f3809,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1040,f252])).
fof(f3810,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1040,f252])).
fof(f3811,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl44 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1040,f252])).
fof(f3812,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl56 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1040,f252])).
fof(f3813,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f1040,f253])).
fof(f3815,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f156,f1040,f156,f1040,f253])).
fof(f3822,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1040,f254])).
fof(f3823,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl61)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1040,f254])).
fof(f3824,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f161,f156,f1040,f257])).
fof(f3826,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl61),
  inference(unit_resulting_resolution,[],[f156,f1040,f1040,f257])).
fof(f3773,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl84),
  inference(cnf_transformation,[],[f3773_D])).
fof(f3773_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3751,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(resolution,[],[f1039,f172])).
fof(f3750,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(resolution,[],[f1039,f170])).
fof(f3775,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl87),
  inference(cnf_transformation,[],[f3775_D])).
fof(f3775_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f3695,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f1039,f227])).
fof(f3696,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f1039,f227])).
fof(f3766,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl83),
  inference(cnf_transformation,[],[f3766_D])).
fof(f3766_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f3699,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f165])).
fof(f3700,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f167])).
fof(f3701,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f170])).
fof(f3702,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f1039,f172])).
fof(f3710,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f1039,f230])).
fof(f3712,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f1039,f230])).
fof(f3714,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f156,f1039,f234])).
fof(f3715,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f156,f1039,f235])).
fof(f3716,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f1039,f236])).
fof(f3718,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f1039,f236])).
fof(f3720,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1039,f248])).
fof(f3721,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl16 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1039,f248])).
fof(f3722,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl44 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1039,f248])).
fof(f3723,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl56 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1039,f248])).
fof(f3724,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl10 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1039,f250])).
fof(f3725,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl14 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1039,f250])).
fof(f3726,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1039,f252])).
fof(f3727,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl16 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1039,f252])).
fof(f3728,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl44 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f1039,f252])).
fof(f3729,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl56 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f1039,f252])).
fof(f3730,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f1039,f253])).
fof(f3732,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f156,f1039,f156,f1039,f253])).
fof(f3739,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl10 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1039,f254])).
fof(f3740,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl14 | $spl59)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1039,f254])).
fof(f3741,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f161,f156,f1039,f257])).
fof(f3743,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f156,f1039,f1039,f257])).
fof(f3690,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl80),
  inference(cnf_transformation,[],[f3690_D])).
fof(f3690_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3691,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(global_subsumption,[],[f155,f886,f3669,f3660,f3658,f3657,f3645,f3643,f3642,f3641,f3639,f3637,f3629,f3628,f3627,f3626,f3625,f3624,f3623,f3670,f3671,f3678,f3680])).
fof(f3680,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(superposition,[],[f886,f162])).
fof(f3678,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) ) | $spl48),
  inference(resolution,[],[f886,f233])).
fof(f3623,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f228])).
fof(f3624,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))))))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f229])).
fof(f3625,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))))))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f225])).
fof(f3671,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(resolution,[],[f886,f168])).
fof(f3670,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(resolution,[],[f886,f166])).
fof(f3626,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f166])).
fof(f3627,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f168])).
fof(f3682,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl79),
  inference(cnf_transformation,[],[f3682_D])).
fof(f3682_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f3628,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f169])).
fof(f3629,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f886,f171])).
fof(f3637,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f156,f886,f233])).
fof(f3639,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f156,f886,f233])).
fof(f3641,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f886,f234])).
fof(f3642,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f886,f235])).
fof(f3643,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f156,f886,f239])).
fof(f3645,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f156,f886,f239])).
fof(f3647,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f886,f248])).
fof(f3648,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f886,f248])).
fof(f3649,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl44 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f776,f886,f248])).
fof(f3650,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f161,f161,f923,f886,f248])).
fof(f3651,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f886,f250])).
fof(f3652,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f886,f250])).
fof(f3653,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f886,f252])).
fof(f3654,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl16 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f886,f252])).
fof(f3655,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl44 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f161,f776,f886,f252])).
fof(f3656,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f161,f161,f923,f886,f252])).
fof(f3657,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f161,f156,f886,f253])).
fof(f3658,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f161,f886,f886,f253])).
fof(f3660,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f161,f886,f886,f253])).
fof(f3667,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl10 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f886,f254])).
fof(f3668,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl14 | $spl48)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f886,f254])).
fof(f3669,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl48),
  inference(unit_resulting_resolution,[],[f161,f161,f886,f257])).
fof(f3594,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(resolution,[],[f884,f172])).
fof(f3593,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(resolution,[],[f884,f170])).
fof(f3620,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl77),
  inference(cnf_transformation,[],[f3620_D])).
fof(f3620_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f3538,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f884,f227])).
fof(f3539,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f884,f227])).
fof(f3610,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl73),
  inference(cnf_transformation,[],[f3610_D])).
fof(f3610_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f3542,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f165])).
fof(f3543,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f167])).
fof(f3544,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f170])).
fof(f3545,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f884,f172])).
fof(f3553,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f884,f230])).
fof(f3555,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f884,f230])).
fof(f3557,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f156,f884,f234])).
fof(f3558,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f156,f884,f235])).
fof(f3559,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f884,f236])).
fof(f3561,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f884,f236])).
fof(f3563,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f884,f248])).
fof(f3564,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl16 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f884,f248])).
fof(f3565,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl44 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f884,f248])).
fof(f3566,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl47 | $spl56)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f884,f248])).
fof(f3567,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl10 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f884,f250])).
fof(f3568,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl14 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f884,f250])).
fof(f3569,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f884,f252])).
fof(f3570,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl16 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f884,f252])).
fof(f3571,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl44 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f161,f776,f884,f252])).
fof(f3572,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl47 | $spl56)),
  inference(unit_resulting_resolution,[],[f156,f161,f923,f884,f252])).
fof(f3573,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f884,f253])).
fof(f3575,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f156,f884,f156,f884,f253])).
fof(f3582,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl10 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f884,f254])).
fof(f3583,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl14 | $spl47)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f884,f254])).
fof(f3584,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f161,f156,f884,f257])).
fof(f3586,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f156,f884,f884,f257])).
fof(f2008,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl67),
  inference(cnf_transformation,[],[f2008_D])).
fof(f2008_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1041,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(cnf_transformation,[],[f1041_D])).
fof(f1041_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3409,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl56),
  inference(superposition,[],[f226,f923])).
fof(f226,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f1189,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl63),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f886,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl48),
  inference(cnf_transformation,[],[f886_D])).
fof(f886_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3264,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl44),
  inference(superposition,[],[f226,f776])).
fof(f288,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl21),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f3229,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f776,f267])).
fof(f295,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(cnf_transformation,[],[f295_D])).
fof(f295_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f3226,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl9 | $spl56)),
  inference(backward_demodulation,[],[f923,f271])).
fof(f3224,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f3223,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f3222,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f3221,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))),
  introduced(theory_axiom,[])).
fof(f3220,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925,f2926,f2927,f2934,f2935,f2936])).
fof(f2936,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(superposition,[],[f888,f162])).
fof(f2935,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7)) ) | $spl50),
  inference(resolution,[],[f888,f239])).
fof(f2934,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl50),
  inference(resolution,[],[f888,f233])).
fof(f2927,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(resolution,[],[f888,f168])).
fof(f2926,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(resolution,[],[f888,f166])).
fof(f2925,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) | $spl50),
  inference(resolution,[],[f888,f216])).
fof(f216,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CLAUSES|split|4')).
fof(f2924,plain,(
  ( ! [X4] : (p(s(bool,X4))) ) | $spl50),
  inference(resolution,[],[f888,f192])).
fof(f192,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'F_0'))) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.OR_CLAUSES|split|3')).
fof(f2923,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl50),
  inference(resolution,[],[f888,f188])).
fof(f188,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) | p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.OR_CLAUSES|split|2')).
fof(f2922,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | $spl50),
  inference(resolution,[],[f888,f186])).
fof(f186,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.AND_CLAUSES|split|3')).
fof(f2921,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl50),
  inference(resolution,[],[f888,f181])).
fof(f181,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) & p(s(bool,X11))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.AND_CLAUSES|split|2')).
fof(f2920,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(resolution,[],[f888,f163])).
fof(f163,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.FALSITY')).
fof(f2606,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f163])).
fof(f2608,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f181])).
fof(f2610,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f186])).
fof(f2612,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f188])).
fof(f2614,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f192])).
fof(f2616,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f888,f216])).
fof(f2618,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f888,f166])).
fof(f2620,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f888,f168])).
fof(f2632,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f156,f888,f233])).
fof(f2636,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f156,f888,f233])).
fof(f2650,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f156,f888,f239])).
fof(f2654,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f156,f888,f239])).
fof(f3219,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925,f2926,f2927,f2934,f2935])).
fof(f3218,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925,f2926,f2927,f2934])).
fof(f3217,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925,f2926,f2927])).
fof(f3216,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925,f2926])).
fof(f3215,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924,f2925])).
fof(f3214,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923,f2924])).
fof(f3213,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922,f2923])).
fof(f3212,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921,f2922])).
fof(f3211,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920,f2921])).
fof(f3210,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606,f2920])).
fof(f3209,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608,f2606])).
fof(f3208,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610,f2608])).
fof(f3207,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612,f2610])).
fof(f3206,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614,f2612])).
fof(f3205,plain,(
  $false | $spl50),
  inference(global_subsumption,[],[f155,f888,f2654,f2650,f2636,f2632,f2620,f2618,f2616,f2614])).
fof(f3204,plain,(
  $false | ($spl33 | $spl50)),
  inference(global_subsumption,[],[f155,f314,f888,f2656,f2654,f2652,f2650,f2638,f2636,f2634])).
fof(f2634,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f233])).
fof(f2638,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f233])).
fof(f2652,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f239])).
fof(f2656,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f239])).
fof(f3203,plain,(
  $false | ($spl33 | $spl50)),
  inference(global_subsumption,[],[f155,f314,f888,f2656,f2654,f2652,f2650,f2638])).
fof(f3194,plain,(
  $false | ($spl33 | $spl50)),
  inference(global_subsumption,[],[f155,f314,f888,f2656,f2654,f2652])).
fof(f2977,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2914])).
fof(f2914,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f161,f888,f257])).
fof(f2976,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2881])).
fof(f2881,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f294,f888,f254])).
fof(f294,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl24),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2975,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2880])).
fof(f2880,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f287,f888,f254])).
fof(f2974,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2879])).
fof(f2879,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f280,f888,f254])).
fof(f2973,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2878])).
fof(f2878,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f275,f888,f254])).
fof(f2972,plain,(
  $false | ($spl22 | $spl50)),
  inference(trivial_inequality_removal,[],[f2877])).
fof(f2877,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl22 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f293,f888,f254])).
fof(f2971,plain,(
  $false | ($spl18 | $spl50)),
  inference(trivial_inequality_removal,[],[f2876])).
fof(f2876,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl18 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f286,f888,f254])).
fof(f2970,plain,(
  $false | ($spl14 | $spl50)),
  inference(trivial_inequality_removal,[],[f2875])).
fof(f2875,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f888,f254])).
fof(f2969,plain,(
  $false | ($spl10 | $spl50)),
  inference(trivial_inequality_removal,[],[f2874])).
fof(f2874,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f888,f254])).
fof(f2968,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2861])).
fof(f2861,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f294,f888,f254])).
fof(f2967,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2847])).
fof(f2847,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f287,f888,f254])).
fof(f2966,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2833])).
fof(f2833,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f280,f888,f254])).
fof(f2965,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2819])).
fof(f2819,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f275,f888,f254])).
fof(f2964,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2805])).
fof(f2805,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f161,f888,f888,f253])).
fof(f2963,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2795])).
fof(f2795,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f161,f888,f888,f253])).
fof(f2962,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2794])).
fof(f2794,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f161,f156,f888,f253])).
fof(f2961,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2777])).
fof(f2777,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f294,f888,f252])).
fof(f2960,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2776])).
fof(f2776,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f287,f888,f252])).
fof(f2959,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2775])).
fof(f2775,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f888,f252])).
fof(f2958,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2774])).
fof(f2774,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f888,f252])).
fof(f2957,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2741])).
fof(f2741,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f294,f888,f250])).
fof(f2956,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2740])).
fof(f2740,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f287,f888,f250])).
fof(f2955,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2739])).
fof(f2739,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f280,f888,f250])).
fof(f2954,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2738])).
fof(f2738,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f888,f275,f888,f250])).
fof(f2953,plain,(
  $false | ($spl22 | $spl50)),
  inference(trivial_inequality_removal,[],[f2737])).
fof(f2737,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl22 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f293,f888,f250])).
fof(f2952,plain,(
  $false | ($spl18 | $spl50)),
  inference(trivial_inequality_removal,[],[f2736])).
fof(f2736,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl18 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f286,f888,f250])).
fof(f2951,plain,(
  $false | ($spl14 | $spl50)),
  inference(trivial_inequality_removal,[],[f2735])).
fof(f2735,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f279,f888,f250])).
fof(f2950,plain,(
  $false | ($spl10 | $spl50)),
  inference(trivial_inequality_removal,[],[f2734])).
fof(f2734,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f156,f274,f888,f250])).
fof(f2949,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2721])).
fof(f2721,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f294,f888,f250])).
fof(f2948,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2707])).
fof(f2707,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f287,f888,f250])).
fof(f2947,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2693])).
fof(f2693,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f280,f888,f250])).
fof(f2946,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2679])).
fof(f2679,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f888,f161,f275,f888,f250])).
fof(f2945,plain,(
  $false | ($spl24 | $spl50)),
  inference(trivial_inequality_removal,[],[f2661])).
fof(f2661,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f294,f888,f248])).
fof(f2944,plain,(
  $false | ($spl20 | $spl50)),
  inference(trivial_inequality_removal,[],[f2660])).
fof(f2660,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl20 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f287,f888,f248])).
fof(f2943,plain,(
  $false | ($spl16 | $spl50)),
  inference(trivial_inequality_removal,[],[f2659])).
fof(f2659,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f280,f888,f248])).
fof(f2942,plain,(
  $false | ($spl12 | $spl50)),
  inference(trivial_inequality_removal,[],[f2658])).
fof(f2658,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl50)),
  inference(unit_resulting_resolution,[],[f161,f161,f275,f888,f248])).
fof(f2657,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f156,f314,f888,f239])).
fof(f2653,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f156,f314,f888,f239])).
fof(f2941,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2645])).
fof(f2645,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f235])).
fof(f2940,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2640])).
fof(f2640,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f234])).
fof(f2639,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f156,f314,f888,f233])).
fof(f2635,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f156,f314,f888,f233])).
fof(f2631,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f888,f888,f216])).
fof(f2939,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2623])).
fof(f2623,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f888,f171])).
fof(f2938,plain,(
  $false | $spl50),
  inference(trivial_inequality_removal,[],[f2622])).
fof(f2622,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl50),
  inference(unit_resulting_resolution,[],[f888,f169])).
fof(f2621,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f168])).
fof(f2619,plain,(
  $false | ($spl33 | $spl50)),
  inference(unit_resulting_resolution,[],[f314,f888,f166])).
fof(f2617,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f156,f888,f216])).
fof(f2615,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f192])).
fof(f2613,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f188])).
fof(f2611,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f186])).
fof(f2609,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f181])).
fof(f2607,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888,f163])).
fof(f2605,plain,(
  $false | $spl50),
  inference(unit_resulting_resolution,[],[f161,f888])).
fof(f2919,plain,(
  $false | $spl50),
  inference(resolution,[],[f888,f161])).
fof(f293,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(cnf_transformation,[],[f293_D])).
fof(f293_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2054,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl57),
  inference(trivial_inequality_removal,[],[f2053])).
fof(f2053,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl57),
  inference(superposition,[],[f924,f162])).
fof(f2056,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl70),
  inference(cnf_transformation,[],[f2056_D])).
fof(f2056_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2046,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f924,f172])).
fof(f2047,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f924,f170])).
fof(f2048,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl57),
  inference(unit_resulting_resolution,[],[f924,f162])).
fof(f2049,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f161,f924,f236])).
fof(f2050,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f161,f924,f230])).
fof(f2051,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f161,f924,f236])).
fof(f2052,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl57),
  inference(unit_resulting_resolution,[],[f161,f924,f230])).
fof(f2043,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl40)),
  inference(backward_demodulation,[],[f287,f772])).
fof(f286,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(cnf_transformation,[],[f286_D])).
fof(f286_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2039,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl45),
  inference(trivial_inequality_removal,[],[f2038])).
fof(f2038,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl45),
  inference(superposition,[],[f777,f162])).
fof(f2041,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl68),
  inference(cnf_transformation,[],[f2041_D])).
fof(f2041_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2031,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f777,f172])).
fof(f2032,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f777,f170])).
fof(f2033,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl45),
  inference(unit_resulting_resolution,[],[f777,f162])).
fof(f2034,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f161,f777,f236])).
fof(f2035,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f161,f777,f230])).
fof(f2036,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f161,f777,f236])).
fof(f2037,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f161,f777,f230])).
fof(f292,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2013,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ($spl9 | $spl54)),
  inference(backward_demodulation,[],[f921,f271])).
fof(f285,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl19),
  inference(cnf_transformation,[],[f285_D])).
fof(f285_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f2010,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ($spl5 | $spl42)),
  inference(backward_demodulation,[],[f774,f267])).
fof(f2009,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl66),
  inference(cnf_transformation,[],[f2009_D])).
fof(f2009_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1253,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl22),
  inference(superposition,[],[f224,f293])).
fof(f1040,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl61),
  inference(cnf_transformation,[],[f1040_D])).
fof(f1040_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1254,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl22),
  inference(superposition,[],[f225,f293])).
fof(f2005,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704,f1705,f1709,f1710,f1711,f1712])).
fof(f1712,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6)) ) | $spl65),
  inference(resolution,[],[f1192,f236])).
fof(f1711,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5)) ) | $spl65),
  inference(resolution,[],[f1192,f236])).
fof(f1710,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) ) | $spl65),
  inference(resolution,[],[f1192,f230])).
fof(f1709,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') = s(bool,X3)) ) | $spl65),
  inference(resolution,[],[f1192,f230])).
fof(f1705,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(resolution,[],[f1192,f172])).
fof(f1704,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(resolution,[],[f1192,f170])).
fof(f1702,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | $spl65),
  inference(resolution,[],[f1192,f203])).
fof(f203,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CLAUSES|split|0')).
fof(f1701,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl65),
  inference(resolution,[],[f1192,f206])).
fof(f206,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CLAUSES|split|1')).
fof(f1700,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl65),
  inference(resolution,[],[f1192,f209])).
fof(f209,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CLAUSES|split|3')).
fof(f1363,plain,(
  p(s(bool,'F_0')) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f157])).
fof(f157,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f1365,plain,(
  p(s(bool,'F_0')) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f159])).
fof(f159,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f1367,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f175])).
fof(f175,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'T_0'))) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.AND_CLAUSES|split|1')).
fof(f1369,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f177])).
fof(f177,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) & p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.AND_CLAUSES|split|0')).
fof(f1371,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f197])).
fof(f197,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.OR_CLAUSES|split|1')).
fof(f1373,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f201])).
fof(f201,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) | p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.OR_CLAUSES|split|0')).
fof(f1375,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f203])).
fof(f1377,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f206])).
fof(f1379,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f209])).
fof(f1381,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f210])).
fof(f210,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f134])).
fof(f1383,plain,(
  p(s(bool,'F_0')) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f212])).
fof(f212,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.bool.IMP_CLAUSES|split|2')).
fof(f1385,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f213])).
fof(f213,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f136])).
fof(f1390,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f170])).
fof(f1392,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f172])).
fof(f1404,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f230])).
fof(f1408,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f230])).
fof(f1422,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f236])).
fof(f1426,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f236])).
fof(f2004,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704,f1705,f1709,f1710,f1711])).
fof(f2003,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704,f1705,f1709,f1710])).
fof(f2002,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704,f1705,f1709])).
fof(f2001,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704,f1705])).
fof(f2000,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702,f1704])).
fof(f1999,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701,f1702])).
fof(f1998,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700,f1701])).
fof(f1997,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363,f1700])).
fof(f1996,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365,f1363])).
fof(f1995,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367,f1365])).
fof(f1994,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369,f1367])).
fof(f1993,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371,f1369])).
fof(f1992,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373,f1371])).
fof(f1991,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375,f1373])).
fof(f1990,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377,f1375])).
fof(f1989,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379,f1377])).
fof(f1988,plain,(
  $false | $spl65),
  inference(global_subsumption,[],[f155,f1192,f1426,f1422,f1408,f1404,f1392,f1390,f1385,f1383,f1381,f1379])).
fof(f1987,plain,(
  $false | ($spl33 | $spl65)),
  inference(global_subsumption,[],[f155,f314,f1192,f1428,f1426,f1424,f1422,f1410,f1408,f1406])).
fof(f1406,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f230])).
fof(f1410,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f230])).
fof(f1424,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f236])).
fof(f1428,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f236])).
fof(f1986,plain,(
  $false | ($spl33 | $spl65)),
  inference(global_subsumption,[],[f155,f314,f1192,f1428,f1426,f1424,f1422,f1410])).
fof(f1977,plain,(
  $false | ($spl33 | $spl65)),
  inference(global_subsumption,[],[f155,f314,f1192,f1428,f1426,f1424])).
fof(f1752,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1687])).
fof(f1687,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f1192,f257])).
fof(f1751,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1681])).
fof(f1681,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f156,f1192,f257])).
fof(f1750,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1664])).
fof(f1664,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f293,f1192,f254])).
fof(f1749,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1663])).
fof(f1663,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f286,f1192,f254])).
fof(f1748,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1662])).
fof(f1662,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1192,f254])).
fof(f1747,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1661])).
fof(f1661,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1192,f254])).
fof(f1746,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1648])).
fof(f1648,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f156,f1192,f253])).
fof(f1745,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1642])).
fof(f1642,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f161,f156,f156,f1192,f253])).
fof(f1744,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1629])).
fof(f1629,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f293,f1192,f252])).
fof(f1743,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1615])).
fof(f1615,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f286,f1192,f252])).
fof(f1742,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1601])).
fof(f1601,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f279,f1192,f252])).
fof(f1741,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1587])).
fof(f1587,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f274,f1192,f252])).
fof(f1740,plain,(
  $false | ($spl24 | $spl65)),
  inference(trivial_inequality_removal,[],[f1553])).
fof(f1553,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f294,f1192,f252])).
fof(f1739,plain,(
  $false | ($spl20 | $spl65)),
  inference(trivial_inequality_removal,[],[f1552])).
fof(f1552,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f287,f1192,f252])).
fof(f1738,plain,(
  $false | ($spl16 | $spl65)),
  inference(trivial_inequality_removal,[],[f1551])).
fof(f1551,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1192,f252])).
fof(f1737,plain,(
  $false | ($spl12 | $spl65)),
  inference(trivial_inequality_removal,[],[f1550])).
fof(f1550,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1192,f252])).
fof(f1736,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1549])).
fof(f1549,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f293,f1192,f252])).
fof(f1735,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1548])).
fof(f1548,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f286,f1192,f252])).
fof(f1734,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1547])).
fof(f1547,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f279,f1192,f252])).
fof(f1733,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1546])).
fof(f1546,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f274,f1192,f252])).
fof(f1732,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1529])).
fof(f1529,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f293,f1192,f250])).
fof(f1731,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1528])).
fof(f1528,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f286,f1192,f250])).
fof(f1730,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1527])).
fof(f1527,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f279,f1192,f250])).
fof(f1729,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1526])).
fof(f1526,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f156,f274,f1192,f250])).
fof(f1728,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1513])).
fof(f1513,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f293,f1192,f248])).
fof(f1727,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1499])).
fof(f1499,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f286,f1192,f248])).
fof(f1726,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1485])).
fof(f1485,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f279,f1192,f248])).
fof(f1725,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1471])).
fof(f1471,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f1192,f156,f274,f1192,f248])).
fof(f1724,plain,(
  $false | ($spl24 | $spl65)),
  inference(trivial_inequality_removal,[],[f1437])).
fof(f1437,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f294,f1192,f248])).
fof(f1723,plain,(
  $false | ($spl20 | $spl65)),
  inference(trivial_inequality_removal,[],[f1436])).
fof(f1436,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f287,f1192,f248])).
fof(f1722,plain,(
  $false | ($spl16 | $spl65)),
  inference(trivial_inequality_removal,[],[f1435])).
fof(f1435,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f280,f1192,f248])).
fof(f1721,plain,(
  $false | ($spl12 | $spl65)),
  inference(trivial_inequality_removal,[],[f1434])).
fof(f1434,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f161,f275,f1192,f248])).
fof(f1720,plain,(
  $false | ($spl22 | $spl65)),
  inference(trivial_inequality_removal,[],[f1433])).
fof(f1433,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f293,f1192,f248])).
fof(f1719,plain,(
  $false | ($spl18 | $spl65)),
  inference(trivial_inequality_removal,[],[f1432])).
fof(f1432,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl18 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f286,f1192,f248])).
fof(f1718,plain,(
  $false | ($spl14 | $spl65)),
  inference(trivial_inequality_removal,[],[f1431])).
fof(f1431,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f279,f1192,f248])).
fof(f1717,plain,(
  $false | ($spl10 | $spl65)),
  inference(trivial_inequality_removal,[],[f1430])).
fof(f1430,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl65)),
  inference(unit_resulting_resolution,[],[f156,f1192,f274,f1192,f248])).
fof(f1429,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f161,f314,f1192,f236])).
fof(f1425,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f161,f314,f1192,f236])).
fof(f1716,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1417])).
fof(f1417,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f235])).
fof(f1715,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1412])).
fof(f1412,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f234])).
fof(f1411,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f161,f314,f1192,f230])).
fof(f1407,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f161,f314,f1192,f230])).
fof(f1403,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f1192,f209])).
fof(f1401,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f1192,f206])).
fof(f1399,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f1192,f203])).
fof(f1393,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f172])).
fof(f1391,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f314,f1192,f170])).
fof(f1714,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1389])).
fof(f1389,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f167])).
fof(f1713,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f1388])).
fof(f1388,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f1192,f165])).
fof(f1386,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f213])).
fof(f1384,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f212])).
fof(f1382,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f210])).
fof(f1380,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f209])).
fof(f1378,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f206])).
fof(f1376,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f203])).
fof(f1374,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f201])).
fof(f1372,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f197])).
fof(f1370,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f177])).
fof(f1368,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192,f175])).
fof(f1366,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f159])).
fof(f1364,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f161,f1192,f157])).
fof(f1362,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f156,f1192])).
fof(f1703,plain,(
  $false | $spl65),
  inference(resolution,[],[f1192,f156])).
fof(f920,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl55),
  inference(cnf_transformation,[],[f920_D])).
fof(f920_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f1345,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(global_subsumption,[],[f155,f293,f1252])).
fof(f1252,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl22),
  inference(superposition,[],[f223,f293])).
fof(f223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X15,X11] : (s(fun(X6,bool),X11) = s(fun(X6,bool),X15) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X15))))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.pred_set.EXTENSION')).
fof(f1101,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl18),
  inference(superposition,[],[f225,f286])).
fof(f1190,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl62),
  inference(cnf_transformation,[],[f1190_D])).
fof(f1190_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1100,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f224,f286])).
fof(f773,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl43),
  inference(cnf_transformation,[],[f773_D])).
fof(f773_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1188,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl18),
  inference(global_subsumption,[],[f155,f286,f1099])).
fof(f1099,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl18),
  inference(superposition,[],[f223,f286])).
fof(f1039,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl59),
  inference(cnf_transformation,[],[f1039_D])).
fof(f1039_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f960,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(superposition,[],[f226,f294])).
fof(f924,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl57),
  inference(cnf_transformation,[],[f924_D])).
fof(f924_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f922,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f294,f271])).
fof(f921,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl54),
  inference(cnf_transformation,[],[f921_D])).
fof(f921_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f917,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(trivial_inequality_removal,[],[f916])).
fof(f916,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(superposition,[],[f295,f162])).
fof(f919,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl52),
  inference(cnf_transformation,[],[f919_D])).
fof(f919_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f909,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f295,f172])).
fof(f910,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f295,f170])).
fof(f911,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl25),
  inference(unit_resulting_resolution,[],[f295,f162])).
fof(f912,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f161,f295,f236])).
fof(f913,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f161,f295,f230])).
fof(f914,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f161,f295,f236])).
fof(f915,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f161,f295,f230])).
fof(f809,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl20),
  inference(superposition,[],[f226,f287])).
fof(f777,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl45),
  inference(cnf_transformation,[],[f777_D])).
fof(f777_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f775,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f287,f267])).
fof(f774,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl42),
  inference(cnf_transformation,[],[f774_D])).
fof(f774_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f770,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl21),
  inference(trivial_inequality_removal,[],[f769])).
fof(f769,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl21),
  inference(superposition,[],[f288,f162])).
fof(f772,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl40),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f762,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f288,f172])).
fof(f763,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f288,f170])).
fof(f764,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl21),
  inference(unit_resulting_resolution,[],[f288,f162])).
fof(f765,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f161,f288,f236])).
fof(f766,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f161,f288,f230])).
fof(f767,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f161,f288,f236])).
fof(f768,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f161,f288,f230])).
fof(f533,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f533_D])).
fof(f533_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f640,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f155,f279,f600])).
fof(f600,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl14),
  inference(superposition,[],[f240,f279])).
fof(f240,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1',aHL_EXT)).
fof(f537,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f537_D])).
fof(f537_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f535,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f280,f269])).
fof(f534,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f534_D])).
fof(f534_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f530,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f529])).
fof(f529,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(superposition,[],[f281,f162])).
fof(f532,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl34),
  inference(cnf_transformation,[],[f532_D])).
fof(f532_D,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f522,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f281,f172])).
fof(f523,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f281,f170])).
fof(f524,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f281,f162])).
fof(f525,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f161,f281,f236])).
fof(f526,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f161,f281,f230])).
fof(f527,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f161,f281,f236])).
fof(f528,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f161,f281,f230])).
fof(f307,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f416,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f155,f274,f384])).
fof(f384,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl10),
  inference(superposition,[],[f240,f274])).
fof(f314,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl33),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f312,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl31)),
  inference(forward_demodulation,[],[f311,f274])).
fof(f311,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f309,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f275,f265])).
fof(f308,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f308_D])).
fof(f308_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f304,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f303])).
fof(f303,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl13),
  inference(superposition,[],[f276,f162])).
fof(f306,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl26),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f296,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f276,f172])).
fof(f297,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f276,f170])).
fof(f298,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f276,f162])).
fof(f299,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f161,f276,f236])).
fof(f300,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f161,f276,f230])).
fof(f301,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f161,f276,f236])).
fof(f302,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f161,f276,f230])).
fof(f291,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl9),
  inference(superposition,[],[f271,f162])).
fof(f284,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl5),
  inference(superposition,[],[f267,f162])).
fof(f279,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f281,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f281_D])).
fof(f281_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f277,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl7),
  inference(superposition,[],[f269,f162])).
fof(f274,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f274_D])).
fof(f274_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f276,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f276_D])).
fof(f276_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f272,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(superposition,[],[f265,f162])).
fof(f271,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl9),
  inference(cnf_transformation,[],[f271_D])).
fof(f271_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f260,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f259,f223])).
fof(f269,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f269_D])).
fof(f269_D,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f261,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f259,f240])).
fof(f267,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl5),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f262,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f259,f223])).
fof(f265,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f263,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f259,f240])).
fof(f259,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) | $spl1),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f155,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X8)))),
  file('thm.pred_set.DIFF_SAME_UNION|split|1','thm.pred_set.DIFF_SAME_UNION|split|1')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 97
% Passive clauses: 116
% Generated clauses: 4427
% Final active clauses: 60
% Final passive clauses: 25
% Input formulas: 37
% Initial clauses: 100
% 
% Duplicate literals: 30
% Trivial inequalities: 120
% Fw demodulations: 1383
% Bw demodulations: 38
% Global subsumptions: 46
% 
% Simple tautologies: 97
% Equational tautologies: 42
% Forward subsumptions: 2151
% Backward subsumptions: 7
% 
% Binary resolution: 109
% Unit resulting resolution: 1698
% Forward superposition: 39
% Backward superposition: 821
% 
% Split clauses: 15
% Split components: 34
% SAT solver clauses: 974
% SAT solver unit clauses: 78
% SAT solver binary clauses: 627
% 
% Sat splits: 15
% Sat splitting refutations: 158
% 
% Memory used [KB]: 2302
% Time elapsed: 0.100 s
% ------------------------------
% ------------------------------
% Success in time 18.118 s
