Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f6570,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6569,f135,f208,f206,f6544,f134,f207,f205])).
fof(f205,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) | $spl0),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f207,plain,(
  s('type.quote.index',sK1) = s('type.quote.index',sK2) | $spl2),
  inference(cnf_transformation,[],[f207_D])).
fof(f207_D,plain,(
  s('type.quote.index',sK1) = s('type.quote.index',sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f134,plain,(
  s('type.quote.index',sK1) = s('type.quote.index',sK2) | s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2)))),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) | s('type.quote.index',sK1) = s('type.quote.index',sK2)) & (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) | s('type.quote.index',sK1) != s('type.quote.index',sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f113])).
fof(f113,plain,(
  ? [X0,X1,X2] : ((s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) | s('type.quote.index',X1) = s('type.quote.index',X2)) & (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) | s('type.quote.index',X1) != s('type.quote.index',X2)))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ? [X0,X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) <~> s('type.quote.index',X1) = s('type.quote.index',X2))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) <=> s('type.quote.index',X1) = s('type.quote.index',X2))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X1))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))) <=> s('type.quote.index',X1) = s('type.quote.index',X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X7))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))) <=> s('type.quote.index',X7) = s('type.quote.index',X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X7))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X8))) <=> s('type.quote.index',X7) = s('type.quote.index',X8))),
  file('thm.ringNorm.polynom_|49||49||split|0','thm.ringNorm.polynom_|49||49||split|0')).
fof(f6544,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f6543,f208])).
fof(f6543,plain,(
  s('type.quote.index',sK1) = s('type.quote.index',sK2) | $spl0),
  inference(forward_demodulation,[],[f6518,f166])).
fof(f166,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X10] : s(X6,X7) = s(X6,'const.ind_type.FCONS_3'(s(X6,X7),s(fun('type.num.num',X6),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.ringNorm.polynom_|49||49||split|0','thm.ind_type.FCONS|split|0')).
fof(f6518,plain,(
  ( ! [X2] : (s('type.quote.index',sK2) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK1),s(fun('type.num.num','type.quote.index'),X2),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f3922,f166])).
fof(f3922,plain,(
  ( ! [X154,X153] : (s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK1),s(fun('type.num.num','type.quote.index'),X153),s('type.num.num',X154))) = s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK2),s(fun('type.num.num','type.quote.index'),X153),s('type.num.num',X154)))) ) | $spl0),
  inference(forward_demodulation,[],[f3858,f176])).
fof(f176,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X9,X16,X18] : s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X16),s('type.num.num',X18))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X9))),s(fun('type.num.num',X6),X16))),s('type.num.num',X18)))),
  file('thm.ringNorm.polynom_|49||49||split|0','ARITY_0_3_const.ind_type.FCONS')).
fof(f3858,plain,(
  ( ! [X154,X153] : (s('type.quote.index','const.ind_type.FCONS_3'(s('type.quote.index',sK2),s(fun('type.num.num','type.quote.index'),X153),s('type.num.num',X154))) = s('type.quote.index',happ(s(fun('type.num.num','type.quote.index'),happ(s(fun(fun('type.num.num','type.quote.index'),fun('type.num.num','type.quote.index')),happ(s(fun('type.quote.index',fun(fun('type.num.num','type.quote.index'),fun('type.num.num','type.quote.index'))),'const.ind_type.FCONS_0'),s('type.quote.index',sK1))),s(fun('type.num.num','type.quote.index'),X153))),s('type.num.num',X154)))) ) | $spl0),
  inference(superposition,[],[f176,f3747])).
fof(f3747,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK1))) = s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK2)))) ) | $spl0),
  inference(forward_demodulation,[],[f3611,f187])).
fof(f187,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun('type.quote.index',X1),X3),s('type.quote.index',X2))) = s(X1,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X0),'const.ringNorm.Pvar_1'(s('type.quote.index',X2))),s(fun('type.quote.index',X1),X3),s(fun(X0,X1),X4),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X5),s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),X1)),X6),s(fun('type.ringNorm.polynom'(X0),X1),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X11] : ! [X7,X10,X12,X13,X14,X15] : s(X11,happ(s(fun('type.quote.index',X11),X10),s('type.quote.index',X7))) = s(X11,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pvar_1'(s('type.quote.index',X7))),s(fun('type.quote.index',X11),X10),s(fun(X6,X11),X12),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X11)),X13),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X11)),X14),s(fun('type.ringNorm.polynom'(X6),X11),X15)))),
  file('thm.ringNorm.polynom_|49||49||split|0','thm.ringNorm.polynom_case_def|split|0')).
fof(f3611,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))),s(fun('type.quote.index',X0),X1),s(fun(sK0,X0),X2),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) ) | $spl0),
  inference(superposition,[],[f187,f205])).
fof(f206,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) | $spl1),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f208,plain,(
  s('type.quote.index',sK1) != s('type.quote.index',sK2) | $spl3),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  s('type.quote.index',sK1) != s('type.quote.index',sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f135,plain,(
  s('type.quote.index',sK1) != s('type.quote.index',sK2) | s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK2)))),
  inference(cnf_transformation,[],[f114])).
fof(f6569,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f6568])).
fof(f6568,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) != s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pvar_1'(s('type.quote.index',sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f206,f207])).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 87
% Passive clauses: 1547
% Generated clauses: 6308
% Final active clauses: 73
% Final passive clauses: 1281
% Input formulas: 37
% Initial clauses: 71
% 
% Duplicate literals: 251
% Trivial inequalities: 103
% Fw subsumption resolutions: 671
% Fw demodulations: 284
% Bw demodulations: 5
% 
% Simple tautologies: 159
% Equational tautologies: 40
% Forward subsumptions: 3039
% 
% Binary resolution: 256
% Factoring: 2
% Forward superposition: 2199
% Backward superposition: 2143
% Self superposition: 148
% Equality factoring: 94
% Equality resolution: 13
% 
% Split clauses: 175
% Split components: 356
% SAT solver clauses: 320
% SAT solver unit clauses: 28
% SAT solver binary clauses: 121
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 175
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 317
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 7547
% Time elapsed: 0.185 s
% ------------------------------
% ------------------------------
% Success in time 0.181 s
