Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quote
% (113844)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 170
% Passive clauses: 3573
% Generated clauses: 21122
% Final active clauses: 143
% Final passive clauses: 3211
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 666
% Trivial inequalities: 370
% Fw subsumption resolutions: 1902
% Fw demodulations: 3070
% Bw demodulations: 25
% 
% Simple tautologies: 330
% Equational tautologies: 52
% Forward subsumptions: 10621
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 457
% Factoring: 1
% Forward superposition: 6840
% Backward superposition: 7283
% Self superposition: 199
% Equality factoring: 92
% Equality resolution: 31
% 
% Split clauses: 286
% Split components: 575
% SAT solver clauses: 730
% SAT solver unit clauses: 7
% SAT solver binary clauses: 146
% 
% Sat splits: 286
% 
% TWLsolver clauses: 730
% TWLsolver calls for satisfiability: 184
% 
% Memory used [KB]: 20596
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quote
% (113858)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 222
% Passive clauses: 1857
% Generated clauses: 19496
% Final active clauses: 82
% Final passive clauses: 519
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 360
% Trivial inequalities: 157
% Fw subsumption resolutions: 1351
% Fw demodulations: 2444
% Bw demodulations: 1095
% Condensations: 902
% Global subsumptions: 12
% 
% Simple tautologies: 173
% Equational tautologies: 48
% Forward subsumptions: 9214
% Bw demodulations to eq. taut.: 35
% 
% Binary resolution: 292
% Factoring: 1
% Forward superposition: 4378
% Backward superposition: 4611
% Self superposition: 233
% Equality factoring: 58
% Equality resolution: 29
% 
% Split clauses: 265
% Split components: 530
% SAT solver clauses: 2458
% SAT solver unit clauses: 162
% SAT solver binary clauses: 109
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 5
% 
% Sat splits: 265
% Sat splitting refutations: 93
% 
% TWLsolver clauses: 2405
% TWLsolver calls for satisfiability: 514
% 
% Memory used [KB]: 15351
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quote
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quote
% SZS output start Proof for thm.quote
fof(f7888,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f238,f7870])).
fof(f7870,plain,(
  ( ! [X216] : (~p(s(bool,X216))) )),
  inference(forward_demodulation,[],[f7869,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : ! [X10,X11,X12] : s(X9,X12) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|1','thm.quote.index_case_def|split|2')).
fof(f7869,plain,(
  ( ! [X216,X215,X214] : (~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X214),s(fun('type.quote.index',bool),X215),s(bool,X216))))) )),
  inference(subsumption_resolution,[],[f7755,f7720])).
fof(f7720,plain,(
  ( ! [X4,X7,X5] : (s(X4,X7) = s(X4,happ(s(fun('type.quote.index',X4),X5),s('type.quote.index',sK0)))) )),
  inference(superposition,[],[f2718,f306])).
fof(f2718,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))) )),
  inference(superposition,[],[f311,f237])).
fof(f237,plain,(
  s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0)))),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f170])).
fof(f170,plain,(
  ? [X0] : s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0] : s('type.quote.index','const.quote.End_idx_0') != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.quote.index','const.quote.End_idx_0') = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6)))),
  file('thm.quote.index_distinct|split|1','thm.quote.index_distinct|split|1')).
fof(f311,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : ! [X6,X10,X11,X12] : s(X9,happ(s(fun('type.quote.index',X9),X10),s('type.quote.index',X6))) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6))),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_distinct|split|1','thm.quote.index_case_def|split|0')).
fof(f7755,plain,(
  ( ! [X216,X215,X214] : (s(bool,'F_0') != s(bool,happ(s(fun('type.quote.index',bool),X214),s('type.quote.index',sK0))) | ~p(s(bool,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.End_idx_0'),s(fun('type.quote.index',bool),X214),s(fun('type.quote.index',bool),X215),s(bool,X216))))) )),
  inference(superposition,[],[f258,f2718])).
fof(f258,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X28] : (s(bool,'F_0') = s(bool,X28) <=> ~p(s(bool,X28)))),
  file('thm.quote.index_distinct|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f238,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quote.index_distinct|split|1','thm.bool.TRUTH')).
% SZS output end Proof for thm.quote
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 144
% Passive clauses: 1071
% Generated clauses: 7667
% Final active clauses: 120
% Final passive clauses: 836
% Input formulas: 69
% Initial clauses: 127
% 
% Duplicate literals: 4
% Trivial inequalities: 17
% Fw subsumption resolutions: 164
% Fw demodulations: 1332
% Bw demodulations: 112
% Condensations: 26
% 
% Simple tautologies: 46
% Equational tautologies: 65
% Forward subsumptions: 4793
% Fw demodulations to eq. taut.: 52
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 1244
% Factoring: 1
% Forward superposition: 2120
% Backward superposition: 2560
% Self superposition: 106
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 1144
% SAT solver unit clauses: 500
% SAT solver binary clauses: 478
% 
% Sat splits: 10
% Sat splitting refutations: 51
% 
% Memory used [KB]: 22131
% Time elapsed: 0.450 s
% ------------------------------
% ------------------------------
% Success in time 1.79 s
