Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (68577)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 2320
% Generated clauses: 20695
% Final active clauses: 123
% Final passive clauses: 2021
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 1180
% Trivial inequalities: 236
% Fw subsumption resolutions: 3227
% Fw demodulations: 2262
% Bw demodulations: 7
% 
% Simple tautologies: 437
% Equational tautologies: 65
% Forward subsumptions: 8972
% 
% Binary resolution: 405
% Forward superposition: 5806
% Backward superposition: 6856
% Self superposition: 224
% Equality factoring: 175
% Equality resolution: 26
% 
% Split clauses: 1414
% Split components: 2911
% SAT solver clauses: 2071
% SAT solver unit clauses: 29
% SAT solver binary clauses: 286
% 
% Sat splits: 1414
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2059
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (68591)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 472
% Generated clauses: 10254
% Final active clauses: 81
% Final passive clauses: 354
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 1170
% Trivial inequalities: 99
% Fw subsumption resolutions: 719
% Fw demodulations: 583
% Bw demodulations: 31
% Condensations: 1627
% Global subsumptions: 4
% 
% Simple tautologies: 202
% Equational tautologies: 44
% Forward subsumptions: 3544
% 
% Binary resolution: 159
% Factoring: 1
% Forward superposition: 2416
% Backward superposition: 2680
% Self superposition: 157
% Equality factoring: 124
% Equality resolution: 14
% 
% Split clauses: 127
% Split components: 262
% SAT solver clauses: 1030
% SAT solver unit clauses: 43
% SAT solver binary clauses: 61
% SAT solver learnt clauses: 8
% 
% Sat splits: 127
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1014
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (68605)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 571
% Passive clauses: 4818
% Generated clauses: 120332
% Final active clauses: 47
% Final passive clauses: 141
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 6112
% Trivial inequalities: 712
% Fw subsumption resolutions: 8899
% Fw demodulations: 17762
% Bw demodulations: 3653
% Condensations: 11662
% Global subsumptions: 3637
% 
% Simple tautologies: 955
% Equational tautologies: 3348
% Forward subsumptions: 58107
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 139
% 
% Unit resulting resolution: 13578
% Forward superposition: 24049
% Backward superposition: 14973
% Self superposition: 564
% Equality factoring: 352
% Equality resolution: 102
% 
% Split clauses: 364
% Split components: 739
% SAT solver clauses: 13268
% SAT solver unit clauses: 80
% SAT solver binary clauses: 2286
% 
% Sat splits: 364
% Sat splitting refutations: 4082
% 
% Memory used [KB]: 69593
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2281,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f238,f258,f256,f254,f252,f262,f263,f273,f271,f312,f443,f450,f444,f454,f445,f446,f452,f458,f447,f448,f456,f457,f523,f524,f525,f753,f760,f754,f764,f755,f756,f762,f768,f757,f758,f766,f767,f824,f825,f826,f1594,f1601,f1595,f1596,f1597,f1603,f1609,f1598,f1599,f1673,f1674,f1675,f1717,f1724,f1718,f1728,f1719,f1720,f1726,f1732,f1721,f1722,f1730,f1731,f1897,f1898,f1899,f2071,f2132,f1605,f2135,f2136,f1607,f1608,f2277])).
fof(f2277,plain,(
  $false | ($spl38 | $spl40)),
  inference(subsumption_resolution,[],[f2276,f108])).
fof(f108,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_UNIV|split|1',aHL_FALSITY)).
fof(f2276,plain,(
  p(s(bool,'F_0')) | ($spl38 | $spl40)),
  inference(forward_demodulation,[],[f2246,f1606])).
fof(f1606,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl38),
  inference(cnf_transformation,[],[f1606_D])).
fof(f1606_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2246,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl40),
  inference(resolution,[],[f1608,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.pred_set.IN_INTER')).
fof(f1608,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) | $spl40),
  inference(cnf_transformation,[],[f1608_D])).
fof(f1608_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1607,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl39),
  inference(cnf_transformation,[],[f1607_D])).
fof(f1607_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f2136,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f79])).
fof(f79,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f556,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(equality_resolution,[],[f122])).
fof(f122,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X0) | s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),X0,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) )),
  inference(superposition,[],[f73,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.pred_set.EXTENSION')).
fof(f73,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f47])).
fof(f47,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0')))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.pred_set.INTER_UNIV|split|1')).
fof(f2135,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f1605,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) | $spl36),
  inference(cnf_transformation,[],[f1605_D])).
fof(f1605_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2132,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(superposition,[],[f556,f86])).
fof(f86,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_UNIV|split|1',aHL_BOOL_CASES)).
fof(f2071,plain,(
  $false | ($spl34 | $spl41)),
  inference(subsumption_resolution,[],[f2070,f110])).
fof(f110,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_UNIV|split|1',aHL_TRUTH)).
fof(f2070,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl41)),
  inference(forward_demodulation,[],[f2069,f1602])).
fof(f1602,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl34),
  inference(cnf_transformation,[],[f1602_D])).
fof(f1602_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2069,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl41),
  inference(subsumption_resolution,[],[f2040,f90])).
fof(f90,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X9] : p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.pred_set.IN_UNIV')).
fof(f2040,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.UNIV_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl41),
  inference(resolution,[],[f1609,f89])).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f1899,plain,(
  $false | ($spl45 | $spl49)),
  inference(subsumption_resolution,[],[f1894,f1827])).
fof(f1827,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl45),
  inference(trivial_inequality_removal,[],[f1824])).
fof(f1824,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl45),
  inference(superposition,[],[f1726,f83])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f1894,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl49),
  inference(trivial_inequality_removal,[],[f1893])).
fof(f1893,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl49),
  inference(superposition,[],[f1730,f79])).
fof(f1898,plain,(
  $false | ($spl45 | $spl49)),
  inference(subsumption_resolution,[],[f1895,f1827])).
fof(f1895,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl49),
  inference(trivial_inequality_removal,[],[f1892])).
fof(f1892,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl49),
  inference(superposition,[],[f1730,f81])).
fof(f1897,plain,(
  $false | ($spl45 | $spl49)),
  inference(subsumption_resolution,[],[f1896,f1726])).
fof(f1896,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl49),
  inference(trivial_inequality_removal,[],[f1889])).
fof(f1889,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl49),
  inference(superposition,[],[f1730,f86])).
fof(f1731,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) | $spl50),
  inference(cnf_transformation,[],[f1731_D])).
fof(f1731_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1730,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl49),
  inference(cnf_transformation,[],[f1730_D])).
fof(f1730_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1722,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f724,f79])).
fof(f724,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(equality_resolution,[],[f123])).
fof(f123,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X1) | s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) )),
  inference(superposition,[],[f73,f75])).
fof(f1721,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f724,f81])).
fof(f1732,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) | $spl51),
  inference(cnf_transformation,[],[f1732_D])).
fof(f1732_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f1726,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl45),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1720,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f724,f83])).
fof(f1719,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f724,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f1728,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) | $spl46),
  inference(cnf_transformation,[],[f1728_D])).
fof(f1728_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1718,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(superposition,[],[f724,f86])).
fof(f1724,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) | $spl42),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1717,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(superposition,[],[f724,f86])).
fof(f1675,plain,(
  $false | ($spl35 | $spl39)),
  inference(subsumption_resolution,[],[f1670,f1626])).
fof(f1626,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl35),
  inference(trivial_inequality_removal,[],[f1623])).
fof(f1623,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl35),
  inference(superposition,[],[f1603,f83])).
fof(f1670,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(trivial_inequality_removal,[],[f1669])).
fof(f1669,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(superposition,[],[f1607,f79])).
fof(f1674,plain,(
  $false | ($spl35 | $spl39)),
  inference(subsumption_resolution,[],[f1671,f1626])).
fof(f1671,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(trivial_inequality_removal,[],[f1668])).
fof(f1668,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(superposition,[],[f1607,f81])).
fof(f1673,plain,(
  $false | ($spl35 | $spl39)),
  inference(subsumption_resolution,[],[f1672,f1603])).
fof(f1672,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl39),
  inference(trivial_inequality_removal,[],[f1665])).
fof(f1665,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl39),
  inference(superposition,[],[f1607,f86])).
fof(f1599,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f79])).
fof(f1598,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f81])).
fof(f1609,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) | $spl41),
  inference(cnf_transformation,[],[f1609_D])).
fof(f1609_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1603,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl35),
  inference(cnf_transformation,[],[f1603_D])).
fof(f1603_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1597,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f83])).
fof(f1596,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))))),
  inference(superposition,[],[f556,f85])).
fof(f1595,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(superposition,[],[f556,f86])).
fof(f1601,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) | $spl32),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0'))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1594,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')))))),
  inference(superposition,[],[f556,f86])).
fof(f826,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f821,f783])).
fof(f783,plain,(
  ~p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl25),
  inference(trivial_inequality_removal,[],[f780])).
fof(f780,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl25),
  inference(superposition,[],[f762,f83])).
fof(f821,plain,(
  p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl29),
  inference(trivial_inequality_removal,[],[f820])).
fof(f820,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl29),
  inference(superposition,[],[f766,f79])).
fof(f825,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f822,f783])).
fof(f822,plain,(
  p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl29),
  inference(trivial_inequality_removal,[],[f819])).
fof(f819,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl29),
  inference(superposition,[],[f766,f81])).
fof(f824,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f823,f762])).
fof(f823,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl29),
  inference(trivial_inequality_removal,[],[f816])).
fof(f816,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl29),
  inference(superposition,[],[f766,f86])).
fof(f767,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl30),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f766,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f758,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))))),
  inference(superposition,[],[f467,f79])).
fof(f467,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))),
  inference(equality_resolution,[],[f125])).
fof(f125,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X3) | s(bool,f(s(sK0,sK3(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) != s(bool,g(s(sK0,sK3(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) )),
  inference(superposition,[],[f73,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_UNIV|split|1',aHL_EXT)).
fof(f757,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))))),
  inference(superposition,[],[f467,f81])).
fof(f768,plain,(
  ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) | $spl31),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f762,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl25),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f756,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))))),
  inference(superposition,[],[f467,f83])).
fof(f755,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))))),
  inference(superposition,[],[f467,f85])).
fof(f764,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f754,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))),
  inference(superposition,[],[f467,f86])).
fof(f760,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | $spl22),
  inference(cnf_transformation,[],[f760_D])).
fof(f760_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f753,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),bool,sK0))))),
  inference(superposition,[],[f467,f86])).
fof(f525,plain,(
  $false | ($spl15 | $spl19)),
  inference(subsumption_resolution,[],[f520,f482])).
fof(f482,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl15),
  inference(trivial_inequality_removal,[],[f479])).
fof(f479,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl15),
  inference(superposition,[],[f452,f83])).
fof(f520,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl19),
  inference(trivial_inequality_removal,[],[f519])).
fof(f519,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl19),
  inference(superposition,[],[f456,f79])).
fof(f524,plain,(
  $false | ($spl15 | $spl19)),
  inference(subsumption_resolution,[],[f521,f482])).
fof(f521,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl19),
  inference(trivial_inequality_removal,[],[f518])).
fof(f518,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl19),
  inference(superposition,[],[f456,f81])).
fof(f523,plain,(
  $false | ($spl15 | $spl19)),
  inference(subsumption_resolution,[],[f522,f452])).
fof(f522,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl19),
  inference(trivial_inequality_removal,[],[f515])).
fof(f515,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl19),
  inference(superposition,[],[f456,f86])).
fof(f457,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl20),
  inference(cnf_transformation,[],[f457_D])).
fof(f457_D,plain,(
  p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f456,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl19),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f448,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))))),
  inference(superposition,[],[f426,f79])).
fof(f426,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))),
  inference(equality_resolution,[],[f124])).
fof(f124,plain,(
  ( ! [X2] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X2) | s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),X2,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),X2,bool,sK0))))) )),
  inference(superposition,[],[f73,f111])).
fof(f447,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))))),
  inference(superposition,[],[f426,f81])).
fof(f458,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) | $spl21),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f452,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f446,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))))),
  inference(superposition,[],[f426,f83])).
fof(f445,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | ~p(s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))))),
  inference(superposition,[],[f426,f85])).
fof(f454,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f454_D])).
fof(f454_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f444,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))),
  inference(superposition,[],[f426,f86])).
fof(f450,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f443,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.UNIV_0')),sK1,bool,sK0))))),
  inference(superposition,[],[f426,f86])).
fof(f312,plain,(
  $false | ($spl1 | $spl9)),
  inference(subsumption_resolution,[],[f271,f266])).
fof(f266,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) | $spl1),
  inference(trivial_inequality_removal,[],[f265])).
fof(f265,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) | $spl1),
  inference(superposition,[],[f252,f79])).
fof(f271,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) | $spl9),
  inference(cnf_transformation,[],[f271_D])).
fof(f271_D,plain,(
  ~p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f273,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f263,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) | $spl1),
  inference(superposition,[],[f252,f83])).
fof(f262,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))) | $spl1),
  inference(superposition,[],[f252,f85])).
fof(f252,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0)))) | $spl1),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f254,plain,(
  s(fun(sK0,bool),'const.pred_set.UNIV_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) | $spl3),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  s(fun(sK0,bool),'const.pred_set.UNIV_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f256,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0)))) | $spl5),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f258,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | $spl7),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f238,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.UNIV_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',sK1,bool,sK0))))),
  inference(resolution,[],[f140,f136])).
fof(f136,plain,(
  ( ! [X4] : (p(s(bool,g(s(sK0,sK3(X4,sK1,bool,sK0))))) | s(fun(sK0,bool),X4) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') != s(bool,f(s(sK0,sK3(X4,sK1,bool,sK0))))) )),
  inference(superposition,[],[f117,f81])).
fof(f117,plain,(
  ( ! [X3] : (s(bool,f(s(sK0,sK3(X3,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X3,sK1,bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.UNIV_0')))) )),
  inference(superposition,[],[f73,f111])).
fof(f140,plain,(
  ( ! [X2] : (~p(s(bool,g(s(sK0,sK3('const.pred_set.UNIV_0',X2,bool,sK0))))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))) | s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',X2,bool,sK0))))) )),
  inference(superposition,[],[f120,f85])).
fof(f120,plain,(
  ( ! [X2] : (s(bool,f(s(sK0,sK3('const.pred_set.UNIV_0',X2,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.UNIV_0',X2,bool,sK0)))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2)))) )),
  inference(superposition,[],[f73,f111])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 155
% Passive clauses: 1131
% Generated clauses: 2161
% Final active clauses: 124
% Final passive clauses: 786
% Input formulas: 20
% Initial clauses: 39
% 
% Selected by SInE selection: 20
% SInE iterations: 1
% 
% Duplicate literals: 9
% Trivial inequalities: 72
% Fw subsumption resolutions: 111
% Fw demodulations: 172
% Bw demodulations: 17
% 
% Simple tautologies: 33
% Equational tautologies: 1
% Forward subsumptions: 630
% 
% Binary resolution: 185
% Forward superposition: 1294
% Backward superposition: 251
% Equality resolution: 4
% 
% Split clauses: 30
% Split components: 62
% SAT solver clauses: 46
% SAT solver binary clauses: 42
% 
% Sat splits: 30
% Sat splitting refutations: 15
% 
% Memory used [KB]: 1791
% Time elapsed: 0.070 s
% ------------------------------
% ------------------------------
% Success in time 11.812 s
