Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (109106)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 606
% Passive clauses: 10097
% Generated clauses: 17462
% Final active clauses: 604
% Final passive clauses: 9114
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 271
% Trivial inequalities: 15
% Fw subsumption resolutions: 1188
% Fw demodulations: 86
% Bw demodulations: 2
% 
% Simple tautologies: 254
% Equational tautologies: 57
% Forward subsumptions: 5239
% 
% Binary resolution: 11972
% Factoring: 1014
% Forward superposition: 1457
% Backward superposition: 1002
% Self superposition: 30
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 313
% Split components: 628
% SAT solver clauses: 641
% SAT solver unit clauses: 20
% SAT solver binary clauses: 196
% 
% Sat splits: 313
% 
% TWLsolver clauses: 641
% TWLsolver calls for satisfiability: 379
% 
% Memory used [KB]: 8955
% Time elapsed: 0.609 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (109128)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 413
% Passive clauses: 2929
% Generated clauses: 11214
% Final active clauses: 408
% Final passive clauses: 2513
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 194
% Trivial inequalities: 70
% Fw subsumption resolutions: 642
% Fw demodulations: 223
% Bw demodulations: 7
% Condensations: 315
% Global subsumptions: 22
% 
% Simple tautologies: 292
% Equational tautologies: 17
% Forward subsumptions: 3193
% 
% Binary resolution: 4139
% Factoring: 202
% Forward superposition: 1167
% Backward superposition: 957
% Self superposition: 30
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 3553
% SAT solver unit clauses: 240
% SAT solver binary clauses: 98
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 33
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3465
% TWLsolver calls for satisfiability: 1965
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 285 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (109152)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3500
% Passive clauses: 21782
% Generated clauses: 139698
% Final active clauses: 737
% Final passive clauses: 4747
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 2578
% Trivial inequalities: 327
% Fw subsumption resolutions: 2570
% Fw demodulations: 21380
% Bw demodulations: 12184
% Condensations: 4234
% Global subsumptions: 376
% 
% Simple tautologies: 171
% Equational tautologies: 4067
% Forward subsumptions: 80779
% Bw demodulations to eq. taut.: 21
% 
% Unit resulting resolution: 32274
% Factoring: 5
% Forward superposition: 27030
% Backward superposition: 23935
% Self superposition: 270
% Equality factoring: 126
% Equality resolution: 73
% 
% Split clauses: 183
% Split components: 371
% SAT solver clauses: 23824
% SAT solver unit clauses: 7732
% SAT solver binary clauses: 5755
% 
% Sat splits: 183
% Sat splitting refutations: 1298
% 
% Memory used [KB]: 104518
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 181 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% (109448)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3081
% Passive clauses: 47228
% Generated clauses: 191688
% Final active clauses: 1875
% Final passive clauses: 33722
% Input formulas: 64
% Initial clauses: 227
% 
% Selected by SInE selection: 56
% SInE iterations: 1
% 
% Duplicate literals: 4670
% Trivial inequalities: 797
% Fw subsumption resolutions: 18879
% Fw demodulations: 6720
% Bw demodulations: 802
% 
% Simple tautologies: 5687
% Equational tautologies: 587
% Forward subsumptions: 95739
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 131689
% Factoring: 316
% Forward superposition: 17274
% Backward superposition: 5792
% Self superposition: 39
% Equality factoring: 2
% Equality resolution: 143
% 
% Split clauses: 14487
% Split components: 33239
% SAT solver clauses: 15669
% SAT solver unit clauses: 31
% SAT solver binary clauses: 1045
% 
% Sat splits: 14487
% Sat splitting refutations: 93
% 
% Memory used [KB]: 53346
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 141 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.integer
% (109577)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 395
% Passive clauses: 2923
% Generated clauses: 14637
% Final active clauses: 387
% Final passive clauses: 2435
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 561
% Trivial inequalities: 213
% Fw subsumption resolutions: 1994
% Fw demodulations: 157
% Bw demodulations: 42
% Global subsumptions: 16
% 
% Simple tautologies: 638
% Equational tautologies: 45
% Forward subsumptions: 7881
% 
% Unit resulting resolution: 2035
% Factoring: 18
% Forward superposition: 4597
% Backward superposition: 4576
% Self superposition: 93
% Equality factoring: 6
% Equality resolution: 23
% 
% Split clauses: 22
% Split components: 45
% SAT solver clauses: 3022
% SAT solver unit clauses: 324
% SAT solver binary clauses: 157
% 
% Sat splits: 22
% Sat splitting refutations: 59
% 
% Memory used [KB]: 9594
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 134 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.integer
% (109602)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1043
% Passive clauses: 5607
% Generated clauses: 17386
% Final active clauses: 731
% Final passive clauses: 4137
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 844
% Trivial inequalities: 64
% Fw subsumption resolutions: 1054
% Bw subsumption resolutions: 184
% Fw demodulations: 598
% Bw demodulations: 36
% Condensations: 1564
% Global subsumptions: 2
% 
% Simple tautologies: 1306
% Equational tautologies: 50
% Forward subsumptions: 6129
% Backward subsumptions: 657
% 
% Binary resolution: 8275
% Factoring: 42
% Forward superposition: 1954
% Backward superposition: 2006
% Self superposition: 46
% Equality factoring: 4
% Equality resolution: 15
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 5648
% SAT solver unit clauses: 117
% SAT solver binary clauses: 449
% 
% Sat splits: 1
% Sat splitting refutations: 35
% 
% Memory used [KB]: 10362
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 118 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f4101,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f343,f600,f342,f602,f341,f604,f810,f809,f808,f807,f806,f851,f805,f804,f803,f802,f800,f799,f798,f797,f855,f796,f853,f795,f794,f793,f857,f783,f859,f781,f780,f779,f778,f863,f777,f861,f776,f775,f774,f865,f770,f867,f751,f869,f748,f747,f746,f745,f873,f744,f871,f743,f742,f741,f875,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f883,f722,f721,f720,f719,f879,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f887,f706,f877,f705,f704,f703,f885,f702,f701,f700,f699,f698,f697,f696,f695,f694,f881,f693,f692,f691,f889,f689,f688,f687,f686,f893,f685,f891,f684,f683,f682,f895,f675,f674,f673,f672,f899,f671,f897,f670,f669,f668,f901,f666,f665,f664,f663,f905,f662,f903,f661,f660,f659,f907,f658,f657,f656,f655,f911,f654,f909,f653,f652,f651,f913,f650,f649,f648,f647,f646,f645,f644,f643,f642,f641,f640,f923,f639,f921,f638,f637,f636,f635,f919,f634,f915,f633,f632,f631,f917,f625,f624,f623,f621,f620,f619,f611,f610,f849,f609,f608,f606,f811,f927,f812,f813,f928,f929,f925,f951,f950,f949,f959,f948,f961,f941,f940,f939,f963,f938,f965,f1033,f1032,f1031,f1036,f1104,f1103,f1102,f1107,f1116,f1108,f1114,f1109,f1112,f1324,f1330,f1332,f1341,f1340,f1339,f1343,f1353,f1344,f1351,f1345,f1349,f1412,f1418,f1423,f1434,f1440,f1445,f1468,f1472,f1478,f1482,f1488,f1492,f1498,f1502,f1575,f1574,f1573,f1686,f1550,f1549,f1548,f1700,f1798,f1814,f1907,f1913,f1920,f1940,f1944,f1954,f1958,f2089,f2088,f2093,f2101,f2094,f2099,f2106,f2110,f2114,f2160,f2162,f2201,f2200,f2199,f2212,f2297,f2294,f2303,f2306,f2377,f2381,f2379,f2353,f2383,f2405,f2433,f2773,f2772,f2801,f2770,f2769,f2805,f2768,f2803,f2767,f2807,f2766,f2809,f2765,f2811,f2764,f2763,f2762,f2760,f2759,f2758,f2754,f2753,f2752,f2751,f2750,f2749,f2748,f2747,f2746,f2745,f2744,f2743,f2742,f2741,f2740,f2739,f2738,f2737,f2735,f2734,f2733,f2729,f2728,f2727,f2723,f2722,f2721,f2717,f2716,f2715,f2714,f2713,f2712,f2711,f2710,f2709,f2698,f2697,f2696,f2695,f2694,f2693,f2692,f2691,f2690,f2819,f2685,f2684,f2683,f2682,f2681,f2680,f2823,f2679,f2678,f2677,f2676,f2675,f2821,f2674,f2825,f2671,f2670,f2669,f2668,f2667,f2666,f2665,f2664,f2663,f2827,f2645,f2644,f2643,f2642,f2641,f2640,f2834,f2638,f2637,f2636,f2635,f2634,f2832,f2633,f2837,f2631,f2630,f2629,f2628,f2846,f2627,f2626,f2844,f2625,f2624,f2850,f2622,f2840,f2621,f2842,f2620,f2848,f2619,f2853,f2616,f2615,f2614,f2613,f2612,f2611,f2858,f2609,f2608,f2607,f2856,f2606,f2605,f2604,f2861,f2587,f2586,f2585,f2584,f2583,f2582,f2581,f2580,f2579,f2865,f2560,f2559,f2558,f2557,f2556,f2555,f2554,f2553,f2552,f2869,f2545,f2871,f2544,f2873,f2543,f2875,f2542,f2877,f2539,f2538,f2537,f2536,f2887,f2533,f2532,f2883,f2531,f2530,f2893,f2527,f2881,f2526,f2891,f2525,f2885,f2524,f2897,f2521,f2520,f2519,f2518,f2517,f2516,f2515,f2514,f2513,f2899,f2510,f2509,f2508,f2507,f2506,f2505,f2504,f2503,f2502,f2901,f2501,f2500,f2499,f2498,f2497,f2496,f2495,f2494,f2493,f2903,f2492,f2491,f2490,f2489,f2488,f2487,f2486,f2485,f2484,f2483,f2482,f2481,f2480,f2913,f2479,f2478,f2477,f2911,f2476,f2475,f2909,f2474,f2473,f2907,f2471,f2470,f2468,f2467,f2465,f2464,f2918,f2462,f2916,f2461,f2921,f2459,f2457,f2448,f2447,f2446,f2445,f2813,f2774,f2775,f2925,f2926,f3221,f3220,f2800,f3216,f3211,f3210,f3209,f3208,f3206,f3205,f3204,f3203,f3202,f2810,f3201,f2808,f3200,f3196,f3195,f3194,f3192,f3188,f3187,f3186,f3185,f3184,f3183,f3182,f3179,f3178,f3177,f3176,f2822,f3174,f3272,f3173,f3172,f3171,f3275,f3160,f3278,f3157,f3156,f3155,f3154,f3283,f3152,f3281,f3151,f3150,f3149,f3286,f3128,f2849,f3105,f3296,f3094,f3093,f3092,f3091,f3304,f3090,f3302,f3089,f3088,f3087,f3306,f3084,f3083,f3082,f3081,f3080,f3079,f3078,f3077,f3076,f3075,f3074,f3073,f3072,f3071,f3070,f3069,f2864,f3068,f3067,f3066,f3065,f3310,f3064,f3063,f3062,f3061,f3060,f3059,f3058,f3057,f3056,f3055,f3054,f3053,f3316,f3052,f3308,f3051,f3050,f3049,f3314,f3048,f3047,f3046,f3045,f3044,f3043,f3042,f3041,f3040,f3312,f3039,f3038,f3037,f3318,f3034,f3033,f3032,f3031,f3323,f3029,f3321,f3028,f3027,f3026,f3326,f3022,f3021,f3020,f3019,f3268,f3017,f3264,f3016,f3015,f3014,f3266,f3011,f2874,f3010,f3009,f2872,f3008,f2870,f3002,f3001,f3000,f2999,f3333,f2997,f3331,f2996,f2995,f2994,f3336,f2991,f2990,f2989,f2988,f3341,f2986,f3339,f2985,f2984,f2983,f3344,f2982,f3262,f2981,f2980,f2979,f2978,f3348,f2977,f3346,f2976,f2975,f2974,f2876,f2973,f2972,f2971,f2970,f2969,f2968,f2967,f2966,f2965,f2964,f2963,f2908,f2962,f2906,f2961,f2960,f2959,f2958,f2910,f2957,f2905,f2956,f2955,f2954,f2912,f2948,f2947,f2946,f2944,f2943,f2942,f2934,f2933,f2924,f2932,f2931,f2929,f2928,f2927,f3222,f3354,f3223,f3352,f3224,f3350,f3225,f3226,f3355,f3356,f2812,f3403,f3407,f3411,f3450,f3489,f3451,f3491,f3461,f3497,f3553,f3568,f3605,f3627,f3632,f3634,f3685,f3687,f3642,f3697,f3643,f3699,f3653,f1815,f3705,f3724,f3726,f3725,f3741,f3743,f3728,f3742,f3744,f3756,f3757,f3770,f2301,f3773,f3775,f2299,f3854,f3858,f3860,f2305,f3863,f3856,f3771,f4096,f4097])).
fof(f4097,plain,(
  $false | ($spl3 | $spl150)),
  inference(global_subsumption,[],[f343,f342,f341,f602,f605,f1817,f1816,f2244,f2243,f2303,f2345,f2351,f2384,f2386,f3451,f3492,f3493,f3461,f3510,f3514,f3515,f3519,f3522,f3535,f3540,f3539,f4055,f4066])).
fof(f4066,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))) | ($spl3 | $spl150)),
  inference(superposition,[],[f2351,f2303])).
fof(f4055,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0)))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f605,f2351,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15(X1,X0)),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15(X1,X0)),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f213])).
fof(f213,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | ? [X3] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X3),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))) & (! [X2] : s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) <=> ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) <=> ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) <=> ? [X2] : s('type.integer.int',X1) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X23] : (~s('type.integer.int',X23) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X8] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X8),s('type.integer.int',X23))) <=> ? [X24] : s('type.integer.int',X8) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X24),s('type.integer.int',X23)))))),
  file('thm.integer.INT_EXP_MOD','thm.integer.INT_MOD_EQ|48|')).
fof(f3539,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num',X2))),s('type.num.num',X3)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f2243,f453])).
fof(f453,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) <=> (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) <=> (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X8,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X7))) <=> (s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.integer.INT_EXP_MOD','thm.integer.INT_EXP_EQ|48|')).
fof(f3540,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num',X3)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f496,f2243,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X2),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f214])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X7,X6,X8] : s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X7))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X6)))))),
  file('thm.integer.INT_EXP_MOD','thm.integer.INT_EXP_ADD_EXPONENTS')).
fof(f3535,plain,(
  ( ! [X24,X23] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X24) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23))),s('type.num.num',X24))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f3525])).
fof(f3525,plain,(
  ( ! [X24,X23] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23))),s('type.num.num',X24))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X23)) ) | $spl3),
  inference(superposition,[],[f455,f2386])).
fof(f455,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f279])).
fof(f3522,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X17),s('type.integer.int',X18))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X16))),s('type.integer.int',X18))) | s('type.integer.int',X18) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) ) | $spl3),
  inference(superposition,[],[f364,f2386])).
fof(f3519,plain,(
  ( ! [X8,X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) | s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)),X9)),s('type.integer.int',X9))) | s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)) ) | $spl3),
  inference(superposition,[],[f363,f2386])).
fof(f3515,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) ) | $spl3),
  inference(superposition,[],[f496,f2386])).
fof(f3514,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl3),
  inference(superposition,[],[f496,f2386])).
fof(f3510,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl3),
  inference(superposition,[],[f2386,f446])).
fof(f446,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.integer.INT_EXP_MOD','thm.arithmetic.ADD_SYM')).
fof(f3493,plain,(
  ( ! [X4] : (s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X4,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X4))))) ) | $spl150),
  inference(forward_demodulation,[],[f3456,f2303])).
fof(f3456,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X4)))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))),s('type.num.num',sK28(X4,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))))) ) | $spl150),
  inference(superposition,[],[f451,f2303])).
fof(f451,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK28(X1,X0))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK28(X1,X0))))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28])],[f276])).
fof(f276,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X3] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X3)))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ? [X2] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))) & (! [X2] : s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> ? [X2] : s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> ? [X8] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8))))),
  file('thm.integer.INT_EXP_MOD','thm.arithmetic.LESS_EQ_EXISTS')).
fof(f3492,plain,(
  ( ! [X3] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12))))) ) | $spl150),
  inference(forward_demodulation,[],[f3455,f2303])).
fof(f3455,plain,(
  ( ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),X3))))) ) | $spl150),
  inference(superposition,[],[f451,f2303])).
fof(f2386,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) ) | $spl3),
  inference(forward_demodulation,[],[f2375,f1817])).
fof(f2375,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X27))))),s('type.num.num',X28)))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2365])).
fof(f2365,plain,(
  ( ! [X28,X26,X27] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X27))))),s('type.num.num',X28)))) ) | $spl3),
  inference(superposition,[],[f455,f1817])).
fof(f2384,plain,(
  ( ! [X6,X7,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X6)))),X7)),s('type.integer.int',X7))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl3),
  inference(forward_demodulation,[],[f2359,f1817])).
fof(f2359,plain,(
  ( ! [X6,X7,X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))) | s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X6))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X6)))),X7)),s('type.integer.int',X7))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl3),
  inference(superposition,[],[f363,f1817])).
fof(f2351,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1)))))) ) | $spl3),
  inference(superposition,[],[f1817,f446])).
fof(f2345,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1)))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f605,f1817,f363])).
fof(f2243,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num',X2)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1816,f453])).
fof(f2244,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X2)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f496,f1816,f364])).
fof(f1816,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0))),s('type.num.num',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f605,f453])).
fof(f1817,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f496,f605,f364])).
fof(f605,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f602,f453])).
fof(f4096,plain,(
  $false | ($spl3 | $spl84)),
  inference(global_subsumption,[],[f343,f342,f341,f602,f605,f927,f1789,f1809,f1810,f1798,f1817,f1816,f2189,f2244,f2243,f2267,f2272,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2287,f2288,f2289,f2290,f2291,f2297,f2307,f2312,f2314,f2315,f2316,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2324,f2325,f2326,f2327,f2328,f2329,f2330,f2331,f2345,f2351,f2353,f2384,f2386,f2926,f3510,f3514,f3515,f3519,f3522,f3535,f3540,f3539,f3575,f3579,f3716,f3717,f3720,f3724,f3734,f3735,f3736,f3737,f3741,f3752,f3756,f3766,f3770,f3873,f3890,f3905,f3916,f3922,f3938,f3939,f3949,f3950,f3951,f3957,f3958,f3959,f3960,f3976,f3982,f3995,f4008,f4021,f4030,f4034,f4046,f4047,f4055,f4065])).
fof(f4065,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))) | ($spl3 | $spl84)),
  inference(superposition,[],[f2351,f927])).
fof(f4047,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2330,f446])).
fof(f4046,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num',X12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X12)),X13,X14) | p(s(bool,X13)) | ~p(s(bool,X14))) ) | $spl84),
  inference(resolution,[],[f2330,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f332])).
fof(f332,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f4034,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2329,f446])).
fof(f4030,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X3,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP11(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X3))) | ~p(s(bool,X4))) ) | $spl84),
  inference(resolution,[],[f2329,f597])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f4021,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2325,f446])).
fof(f4008,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2320,f446])).
fof(f3995,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X2,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2319,f446])).
fof(f3982,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2317,f446])).
fof(f3976,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X2,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2316,f446])).
fof(f3960,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2315,f446])).
fof(f3959,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X9,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)))) | sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)),X10,X11) | ~sP3(X10,X11,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9))) | ~sP2(X11,X10,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)))) ) | $spl84),
  inference(resolution,[],[f2315,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f305])).
fof(f305,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f3958,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X6,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | p(s(bool,X7)) | sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6)),X7,X8) | ~sP3(X7,X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6))) | ~sP2(X8,X7,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6)))) ) | $spl84),
  inference(resolution,[],[f2315,f514])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f306])).
fof(f3957,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X3,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ~p(s(bool,X4)) | sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X3)),X5,X4) | ~sP3(X5,X4,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X3))) | ~sP2(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X3)))) ) | $spl84),
  inference(resolution,[],[f2315,f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f306])).
fof(f3951,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X2,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X3) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2314,f446])).
fof(f3950,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X9,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | p(s(bool,X10)) | sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9))) | ~sP3(X11,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)),X10) | ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)),X11,X10)) ) | $spl84),
  inference(resolution,[],[f2314,f515])).
fof(f3949,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X6,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | p(s(bool,X7)) | sP5(X8,X7,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6))) | ~sP3(X7,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6)),X8) | ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X6)),X7,X8)) ) | $spl84),
  inference(resolution,[],[f2314,f514])).
fof(f3939,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X2,X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2290,f446])).
fof(f3938,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X12)))) | sP11('const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK12)),X13,X14) | p(s(bool,X13)) | ~p(s(bool,X14))) ) | $spl84),
  inference(resolution,[],[f2290,f572])).
fof(f3922,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2289,f446])).
fof(f3916,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X3)))) | sP11(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12))) | ~p(s(bool,X4))) ) | $spl84),
  inference(resolution,[],[f2289,f597])).
fof(f3905,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X2,X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2285,f446])).
fof(f3890,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X2,X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2280,f446])).
fof(f3873,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2279,f446])).
fof(f3766,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X2,X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2277,f446])).
fof(f3752,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2276,f446])).
fof(f3737,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X2,X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2275,f446])).
fof(f3736,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X9)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)))) | sP5('const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)),X10,X11) | ~sP3(X10,X11,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12))) | ~sP2(X11,X10,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)))) ) | $spl84),
  inference(resolution,[],[f2275,f515])).
fof(f3735,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X6)))) | p(s(bool,X7)) | sP5('const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12)),X7,X8) | ~sP3(X7,X8,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12))) | ~sP2(X8,X7,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12)))) ) | $spl84),
  inference(resolution,[],[f2275,f514])).
fof(f3734,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X3)))) | ~p(s(bool,X4)) | sP5('const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12)),X5,X4) | ~sP3(X5,X4,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12))) | ~sP2(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12)))) ) | $spl84),
  inference(resolution,[],[f2275,f513])).
fof(f3720,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X2,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)),X3) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f2274,f446])).
fof(f3717,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X9)))) | p(s(bool,X10)) | sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12))) | ~sP3(X11,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)),X10) | ~sP2('const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)),X11,X10)) ) | $spl84),
  inference(resolution,[],[f2274,f515])).
fof(f3716,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X6)))) | p(s(bool,X7)) | sP5(X8,X7,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12))) | ~sP3(X7,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12)),X8) | ~sP2('const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK12)),X7,X8)) ) | $spl84),
  inference(resolution,[],[f2274,f514])).
fof(f3579,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))))) ) | $spl84),
  inference(superposition,[],[f2189,f446])).
fof(f3575,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))))) ) | $spl84),
  inference(superposition,[],[f2189,f446])).
fof(f2331,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f1810,f446])).
fof(f2330,plain,(
  ( ! [X59,X57,X58] : (sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X57)),X58,X59) | s('type.num.num',X57) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X57,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f576])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f335])).
fof(f335,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f2329,plain,(
  ( ! [X54,X56,X55] : (sP10(X55,X56,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X54))) | s('type.num.num',X54) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X54,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f336])).
fof(f2328,plain,(
  ( ! [X52,X53,X51] : (s('type.num.num',X51) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X51,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP9(X52,X53,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X51))) | p(s(bool,X53)) | p(s(bool,X52))) ) | $spl84),
  inference(resolution,[],[f1810,f560])).
fof(f560,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f155,f204,f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X25,X23,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X23)) | p(s(bool,X25)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X23))) & (p(s(bool,X8)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X8)))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_25_prove_ltb')).
fof(f2327,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num',X48) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X48,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X48)),X49,X50) | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl84),
  inference(resolution,[],[f1810,f560])).
fof(f2326,plain,(
  ( ! [X47,X45,X46] : (s('type.num.num',X45) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X45,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP9(X46,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X45)),X47) | p(s(bool,X47)) | p(s(bool,X46))) ) | $spl84),
  inference(resolution,[],[f1810,f560])).
fof(f2325,plain,(
  ( ! [X43,X44,X42] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X42)),X43,X44) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X42,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f327])).
fof(f327,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f203])).
fof(f2324,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',X39) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X39,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X39)),X40,X41) | p(s(bool,X40))) ) | $spl84),
  inference(resolution,[],[f1810,f544])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f153,f201,f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f201,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X25,X23,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X23)) & p(s(bool,X25)))) <=> ((p(s(bool,X8)) | ~p(s(bool,X23)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X8))) & (p(s(bool,X25)) | ~p(s(bool,X8)))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_23_prove_ltb')).
fof(f2323,plain,(
  ( ! [X37,X38,X36] : (s('type.num.num',X36) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X36,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP7(X37,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X36)),X38) | p(s(bool,X37))) ) | $spl84),
  inference(resolution,[],[f1810,f544])).
fof(f2322,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num',X33) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X33,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP7(X34,X35,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X33))) | p(s(bool,X34))) ) | $spl84),
  inference(resolution,[],[f1810,f543])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f322])).
fof(f2321,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num',X30) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X30,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X30)),X31,X32) | p(s(bool,X32))) ) | $spl84),
  inference(resolution,[],[f1810,f543])).
fof(f2320,plain,(
  ( ! [X28,X29,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X27)),X28,X29) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X27,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f542])).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f200])).
fof(f2319,plain,(
  ( ! [X26,X24,X25] : (sP6(X25,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X24)),X26) | s('type.num.num',X24) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X24,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f320])).
fof(f2318,plain,(
  ( ! [X23,X21,X22] : (s('type.num.num',X21) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X21,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X21)),X22,X23) | s(bool,X22) = s(bool,X23)) ) | $spl84),
  inference(resolution,[],[f1810,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f151,f198,f197,f196,f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f196,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X25,X23,X8] : ((p(s(bool,X8)) <=> s(bool,X23) = s(bool,X25)) <=> ((p(s(bool,X8)) | p(s(bool,X23)) | p(s(bool,X25))) & (p(s(bool,X8)) | ~p(s(bool,X25)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X25)) | ~p(s(bool,X8))) & (p(s(bool,X25)) | ~p(s(bool,X23)) | ~p(s(bool,X8)))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_2|49|_prove_ltb')).
fof(f2317,plain,(
  ( ! [X19,X20,X18] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X18)),X19,X20) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X18,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f311])).
fof(f311,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f196])).
fof(f2316,plain,(
  ( ! [X17,X15,X16] : (sP3(X16,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X15)),X17) | s('type.num.num',X15) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X15,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f312])).
fof(f2315,plain,(
  ( ! [X14,X12,X13] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X12)),X13,X14) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f308])).
fof(f308,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f2314,plain,(
  ( ! [X10,X11,X9] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X9)),X11) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X9,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(resolution,[],[f1810,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f309])).
fof(f2312,plain,(
  ( ! [X6,X5] : (s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X5,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | p(s(bool,X6)) | s(bool,X6) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X5)))) ) | $spl84),
  inference(resolution,[],[f1810,f447])).
fof(f447,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.integer.INT_EXP_MOD','thm.bool.IMP_ANTISYM_AX')).
fof(f2307,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X0,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X0)))) ) | $spl84),
  inference(resolution,[],[f1810,f370])).
fof(f370,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.integer.INT_EXP_MOD','thm.bool.EQ_CLAUSES|split|3')).
fof(f2291,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f1809,f446])).
fof(f2290,plain,(
  ( ! [X59,X57,X58] : (sP10('const.arithmetic.<=_2'(s('type.num.num',X57),s('type.num.num',sK12)),X58,X59) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X57))))) ) | $spl84),
  inference(resolution,[],[f1809,f576])).
fof(f2289,plain,(
  ( ! [X54,X56,X55] : (sP10(X55,X56,'const.arithmetic.<=_2'(s('type.num.num',X54),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X54),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X54))))) ) | $spl84),
  inference(resolution,[],[f1809,f574])).
fof(f2288,plain,(
  ( ! [X52,X53,X51] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X51)))) | sP9(X52,X53,'const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num',sK12))) | p(s(bool,X53)) | p(s(bool,X52))) ) | $spl84),
  inference(resolution,[],[f1809,f560])).
fof(f2287,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X48)))) | sP9('const.arithmetic.<=_2'(s('type.num.num',X48),s('type.num.num',sK12)),X49,X50) | p(s(bool,X49)) | p(s(bool,X50))) ) | $spl84),
  inference(resolution,[],[f1809,f560])).
fof(f2286,plain,(
  ( ! [X47,X45,X46] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X45)))) | sP9(X46,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',sK12)),X47) | p(s(bool,X47)) | p(s(bool,X46))) ) | $spl84),
  inference(resolution,[],[f1809,f560])).
fof(f2285,plain,(
  ( ! [X43,X44,X42] : (sP8('const.arithmetic.<=_2'(s('type.num.num',X42),s('type.num.num',sK12)),X43,X44) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X42))))) ) | $spl84),
  inference(resolution,[],[f1809,f559])).
fof(f2284,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X39)))) | sP7('const.arithmetic.<=_2'(s('type.num.num',X39),s('type.num.num',sK12)),X40,X41) | p(s(bool,X40))) ) | $spl84),
  inference(resolution,[],[f1809,f544])).
fof(f2283,plain,(
  ( ! [X37,X38,X36] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X36)))) | sP7(X37,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num',sK12)),X38) | p(s(bool,X37))) ) | $spl84),
  inference(resolution,[],[f1809,f544])).
fof(f2282,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X33),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X33)))) | sP7(X34,X35,'const.arithmetic.<=_2'(s('type.num.num',X33),s('type.num.num',sK12))) | p(s(bool,X34))) ) | $spl84),
  inference(resolution,[],[f1809,f543])).
fof(f2281,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X30)))) | sP7('const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',sK12)),X31,X32) | p(s(bool,X32))) ) | $spl84),
  inference(resolution,[],[f1809,f543])).
fof(f2280,plain,(
  ( ! [X28,X29,X27] : (sP6('const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',sK12)),X28,X29) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X27))))) ) | $spl84),
  inference(resolution,[],[f1809,f542])).
fof(f2279,plain,(
  ( ! [X26,X24,X25] : (sP6(X25,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',sK12)),X26) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X24))))) ) | $spl84),
  inference(resolution,[],[f1809,f541])).
fof(f2278,plain,(
  ( ! [X23,X21,X22] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X21)))) | sP5('const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',sK12)),X22,X23) | s(bool,X22) = s(bool,X23)) ) | $spl84),
  inference(resolution,[],[f1809,f528])).
fof(f2277,plain,(
  ( ! [X19,X20,X18] : (sP3('const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num',sK12)),X19,X20) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X18))))) ) | $spl84),
  inference(resolution,[],[f1809,f523])).
fof(f2276,plain,(
  ( ! [X17,X15,X16] : (sP3(X16,'const.arithmetic.<=_2'(s('type.num.num',X15),s('type.num.num',sK12)),X17) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X15))))) ) | $spl84),
  inference(resolution,[],[f1809,f522])).
fof(f2275,plain,(
  ( ! [X14,X12,X13] : (sP4('const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',sK12)),X13,X14) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X12))))) ) | $spl84),
  inference(resolution,[],[f1809,f519])).
fof(f2274,plain,(
  ( ! [X10,X11,X9] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',sK12)),X11) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X9))))) ) | $spl84),
  inference(resolution,[],[f1809,f518])).
fof(f2272,plain,(
  ( ! [X6,X5] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X5)))) | p(s(bool,X6)) | s(bool,X6) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',sK12)))) ) | $spl84),
  inference(resolution,[],[f1809,f447])).
fof(f2267,plain,(
  ( ! [X0] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X0)))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK12)))) ) | $spl84),
  inference(resolution,[],[f1809,f370])).
fof(f2189,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) ) | $spl84),
  inference(superposition,[],[f1789,f446])).
fof(f1810,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X4)))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(X4,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(forward_demodulation,[],[f1793,f927])).
fof(f1793,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X4)))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28(X4,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f451,f927])).
fof(f1809,plain,(
  ( ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X3))))) ) | $spl84),
  inference(forward_demodulation,[],[f1792,f927])).
fof(f1792,plain,(
  ( ! [X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),X3))))) ) | $spl84),
  inference(superposition,[],[f451,f927])).
fof(f1789,plain,(
  ( ! [X0] : (s('type.num.num',sK12) != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',X0))))) ) | $spl84),
  inference(superposition,[],[f452,f927])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f3771,plain,(
  ( ! [X8,X9] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) | $spl358),
  inference(cnf_transformation,[],[f3771_D])).
fof(f3771_D,plain,(
  ( ! [X8,X9] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f3856,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) | $spl362),
  inference(cnf_transformation,[],[f3856_D])).
fof(f3856_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f3863,plain,(
  ( ! [X39,X40] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) | sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X39,X40)) ) | ($spl84 | $spl148)),
  inference(forward_demodulation,[],[f3827,f2301])).
fof(f3827,plain,(
  ( ! [X39,X40] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X39,X40) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))))) ) | ($spl84 | $spl148)),
  inference(superposition,[],[f2277,f2301])).
fof(f2305,plain,(
  ~p(s(bool,'T_0')) | $spl153),
  inference(cnf_transformation,[],[f2305_D])).
fof(f2305_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f3860,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl364),
  inference(cnf_transformation,[],[f3860_D])).
fof(f3860_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f3858,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ~p(s(bool,'T_0')) | ($spl84 | $spl148 | $spl172)),
  inference(forward_demodulation,[],[f3857,f2301])).
fof(f3857,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ($spl84 | $spl148 | $spl172)),
  inference(forward_demodulation,[],[f3809,f2812])).
fof(f3809,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ($spl84 | $spl148)),
  inference(superposition,[],[f1810,f2301])).
fof(f3854,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) | ~p(s(bool,'T_0')) | ($spl84 | $spl148 | $spl172)),
  inference(forward_demodulation,[],[f3853,f2301])).
fof(f3853,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) | ($spl84 | $spl148 | $spl172)),
  inference(forward_demodulation,[],[f3808,f2812])).
fof(f3808,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))))) | ($spl84 | $spl148)),
  inference(superposition,[],[f1809,f2301])).
fof(f2299,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl147),
  inference(cnf_transformation,[],[f2299_D])).
fof(f2299_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f3775,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl360),
  inference(cnf_transformation,[],[f3775_D])).
fof(f3775_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f3773,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | ($spl84 | $spl150)),
  inference(forward_demodulation,[],[f3468,f2303])).
fof(f3468,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ($spl84 | $spl150)),
  inference(superposition,[],[f1810,f2303])).
fof(f2301,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl148),
  inference(cnf_transformation,[],[f2301_D])).
fof(f2301_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3770,plain,(
  ( ! [X8,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) | $spl84),
  inference(forward_demodulation,[],[f3768,f927])).
fof(f3768,plain,(
  ( ! [X8,X9] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2277,f927])).
fof(f3757,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) | $spl356),
  inference(cnf_transformation,[],[f3757_D])).
fof(f3757_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f3756,plain,(
  ( ! [X8,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) | $spl84),
  inference(forward_demodulation,[],[f3754,f927])).
fof(f3754,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2276,f927])).
fof(f3744,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ~p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))))),
  introduced(theory_axiom,[])).
fof(f3742,plain,(
  ( ! [X8,X9] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) | $spl354),
  inference(cnf_transformation,[],[f3742_D])).
fof(f3742_D,plain,(
  ( ! [X8,X9] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f3728,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))))) | $spl352),
  inference(cnf_transformation,[],[f3728_D])).
fof(f3728_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))))) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f3743,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))))) | sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X10,X11)) ) | ($spl84 | $spl150)),
  inference(forward_demodulation,[],[f3740,f2303])).
fof(f3740,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X10,X11) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))))) ) | ($spl84 | $spl150)),
  inference(superposition,[],[f2275,f2303])).
fof(f3741,plain,(
  ( ! [X8,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9)) ) | $spl84),
  inference(forward_demodulation,[],[f3739,f927])).
fof(f3739,plain,(
  ( ! [X8,X9] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X8,X9) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2275,f927])).
fof(f3725,plain,(
  ( ! [X8,X9] : (sP4(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) | $spl350),
  inference(cnf_transformation,[],[f3725_D])).
fof(f3725_D,plain,(
  ( ! [X8,X9] : (sP4(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f3726,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))))) | sP4(X10,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X11)) ) | ($spl84 | $spl150)),
  inference(forward_demodulation,[],[f3723,f2303])).
fof(f3723,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X11) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))))) ) | ($spl84 | $spl150)),
  inference(superposition,[],[f2274,f2303])).
fof(f3724,plain,(
  ( ! [X8,X9] : (s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | sP4(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9)) ) | $spl84),
  inference(forward_demodulation,[],[f3722,f927])).
fof(f3722,plain,(
  ( ! [X8,X9] : (sP4(X8,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X9) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl84),
  inference(superposition,[],[f2274,f927])).
fof(f3705,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl349),
  inference(cnf_transformation,[],[f3705_D])).
fof(f3705_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f1815,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X12),s('type.num.num',sK12)))) ) | $spl128),
  inference(cnf_transformation,[],[f1815_D])).
fof(f1815_D,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X12),s('type.num.num',sK12)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f3653,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X12),s('type.num.num',sK12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))))))) ) | $spl148),
  inference(superposition,[],[f454,f2301])).
fof(f454,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f279])).
fof(f3699,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) | $spl346),
  inference(cnf_transformation,[],[f3699_D])).
fof(f3699_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f3643,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) | ($spl3 | $spl148)),
  inference(superposition,[],[f1817,f2301])).
fof(f3697,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) | $spl344),
  inference(cnf_transformation,[],[f3697_D])).
fof(f3697_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f3642,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))))) | ($spl1 | $spl148)),
  inference(superposition,[],[f2023,f2301])).
fof(f2023,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',X1)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f496,f933,f364])).
fof(f933,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f453])).
fof(f3687,plain,(
  p(s(bool,'F_0')) | $spl342),
  inference(cnf_transformation,[],[f3687_D])).
fof(f3687_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f3685,plain,(
  p(s(bool,'F_0')) | ($spl148 | $spl258)),
  inference(forward_demodulation,[],[f3640,f2925])).
fof(f3640,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl148),
  inference(unit_resulting_resolution,[],[f2301,f452])).
fof(f3634,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))) | $spl340),
  inference(cnf_transformation,[],[f3634_D])).
fof(f3634_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f3632,plain,(
  s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12)),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | $spl338),
  inference(cnf_transformation,[],[f3632_D])).
fof(f3632_D,plain,(
  s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12)),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f3627,plain,(
  s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12)),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))) | $spl156),
  inference(trivial_inequality_removal,[],[f3608])).
fof(f3608,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12)),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))) | $spl156),
  inference(superposition,[],[f363,f2383])).
fof(f3605,plain,(
  s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15('const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12)),'const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | ($spl3 | $spl156)),
  inference(unit_resulting_resolution,[],[f605,f2383,f363])).
fof(f3568,plain,(
  sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl336),
  inference(cnf_transformation,[],[f3568_D])).
fof(f3568_D,plain,(
  sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f3553,plain,(
  sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f678,f2339])).
fof(f2339,plain,(
  ( ! [X0,X1] : (~sP2(X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0) | p(s(bool,X0)) | sP5(X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X1) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f1015,f522])).
fof(f1015,plain,(
  ( ! [X6,X7] : (~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X7,X6) | sP5(X6,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X7) | p(s(bool,X6)) | ~sP2(X7,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X6)) ) | $spl4),
  inference(resolution,[],[f667,f515])).
fof(f667,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f517])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f309])).
fof(f678,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f195])).
fof(f350,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_EXP_MOD',aHL_FALSITY)).
fof(f3497,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) | $spl335),
  inference(cnf_transformation,[],[f3497_D])).
fof(f3497_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f3461,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X12),s('type.num.num',sK12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13))))) ) | $spl150),
  inference(superposition,[],[f454,f2303])).
fof(f3491,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) | $spl332),
  inference(cnf_transformation,[],[f3491_D])).
fof(f3491_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f3451,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) | ($spl3 | $spl150)),
  inference(superposition,[],[f1817,f2303])).
fof(f3489,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) | $spl330),
  inference(cnf_transformation,[],[f3489_D])).
fof(f3489_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f3450,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))))) | ($spl1 | $spl150)),
  inference(superposition,[],[f2023,f2303])).
fof(f3411,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl328),
  inference(cnf_transformation,[],[f3411_D])).
fof(f3411_D,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f3407,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl144),
  inference(resolution,[],[f2212,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f306])).
fof(f3403,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl144),
  inference(unit_resulting_resolution,[],[f2212,f510])).
fof(f2812,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl172),
  inference(cnf_transformation,[],[f2812_D])).
fof(f2812_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f3356,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(global_subsumption,[],[f343,f342,f341,f2298,f3200,f3192,f3188,f3187,f3186,f3185,f3180,f3179,f3178,f3177,f3176,f3174,f3173,f3172,f3171,f3170,f3169,f3165,f3162,f3160,f3159,f3157,f3156,f3155,f3154,f3152,f3151,f3150,f3149,f3133,f3129,f3128,f3127,f3126,f3121,f3118,f3117,f3112,f3109,f3108,f3105,f3094,f3093,f3092,f3091,f3090,f3089,f3088,f3087,f3086,f3085,f3084,f3083,f3082,f3081,f3080,f3079,f3078,f3077,f3076,f3075,f3074,f3073,f3072,f3071,f3070,f3069,f3068,f3067,f3066,f3065,f3064,f3063,f3062,f3061,f3060,f3059,f3058,f3057,f3056,f3055,f3054,f3053,f3052,f3051,f3050,f3049,f3048,f3047,f3046,f3045,f3044,f3043,f3042,f3041,f3040,f3039,f3038,f3037,f3036,f3034,f3033,f3032,f3031,f3029,f3028,f3027,f3026,f3007,f3006,f3005,f3004,f3002,f3001,f3000,f2999,f2997,f2996,f2995,f2994,f2993,f2991,f2990,f2989,f2988,f2986,f2985,f2984,f2983,f2981,f2980,f2979,f2978,f2977,f2976,f2975,f2974,f2973,f2972,f2971,f2970,f2969,f2968,f2967,f2966,f2965,f2964,f2963,f2962,f2961,f2960,f2959,f2958,f2957,f2956,f2955,f2954,f2953,f2952,f2948,f2947,f2946,f2944,f2943,f2942,f2934,f2933,f2932,f2931,f2929,f3224,f3225,f3226,f3233,f3238,f3239,f3240,f3243,f3250,f3253,f3254,f3256,f3257,f3258,f3259,f3355,f3260])).
fof(f3260,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(superposition,[],[f2298,f352])).
fof(f352,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_EXP_MOD',aHL_BOOL_CASES)).
fof(f3259,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(superposition,[],[f2298,f353])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X9))),
  file('thm.integer.INT_EXP_MOD','thm.bool.BOOL_CASES_AX')).
fof(f3258,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X39,X40)) ) | $spl146),
  inference(resolution,[],[f2298,f579])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f188,f207,f206])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X25,X23,X8] : ((p(s(bool,X8)) <=> (p(s(bool,X23)) => p(s(bool,X25)))) <=> ((p(s(bool,X8)) | p(s(bool,X23))) & (p(s(bool,X8)) | ~p(s(bool,X25))) & (~p(s(bool,X23)) | p(s(bool,X25)) | ~p(s(bool,X8)))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_27_prove_ltb')).
fof(f3257,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X38,X37)) ) | $spl146),
  inference(resolution,[],[f2298,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f3256,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | sP11(X36,X35,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | p(s(bool,X36))) ) | $spl146),
  inference(resolution,[],[f2298,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f338])).
fof(f3254,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X31,X32)) ) | $spl146),
  inference(resolution,[],[f2298,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f330])).
fof(f3253,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X30,X29)) ) | $spl146),
  inference(resolution,[],[f2298,f561])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f330])).
fof(f3250,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X23,X24)) ) | $spl146),
  inference(resolution,[],[f2298,f545])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f322])).
fof(f3243,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X10)) ) | $spl146),
  inference(resolution,[],[f2298,f483])).
fof(f483,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f138,f193])).
fof(f193,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X23,X8] : ((p(s(bool,X8)) <=> ~p(s(bool,X23))) <=> ((p(s(bool,X8)) | p(s(bool,X23))) & (~p(s(bool,X23)) | ~p(s(bool,X8)))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_29_prove_ltb')).
fof(f3240,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(resolution,[],[f2298,f480])).
fof(f480,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f193])).
fof(f3239,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X6)) ) | $spl146),
  inference(resolution,[],[f2298,f477])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f3238,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~sP1(X5,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(resolution,[],[f2298,f476])).
fof(f476,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f3233,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(resolution,[],[f2298,f450])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f2952,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f466])).
fof(f466,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f180,f191])).
fof(f191,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f180,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X17,X12] : ((~(p(s(bool,X12)) | p(s(bool,X17))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X12)) => p(s(bool,'F_0'))) => (~p(s(bool,X17)) => p(s(bool,'F_0')))))),
  file('thm.integer.INT_EXP_MOD','thm.sat.sat_9_prove_ltb')).
fof(f2953,plain,(
  ( ! [X0] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f467])).
fof(f467,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f2993,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f517])).
fof(f3004,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f312])).
fof(f3005,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f3006,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X1)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f526])).
fof(f3007,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0,X1)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f3036,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f540])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f320])).
fof(f3085,plain,(
  ( ! [X0] : (~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0,'F_0')) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f546])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f322])).
fof(f3086,plain,(
  ( ! [X0] : (~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0',X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f547])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f3108,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f557])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f328])).
fof(f3109,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X1)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f328])).
fof(f3112,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f561])).
fof(f344,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_EXP_MOD','thm.bool.TRUTH')).
fof(f3117,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0,'T_0')) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f561])).
fof(f3118,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f561])).
fof(f3121,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f562])).
fof(f3126,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0',X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f562])).
fof(f3127,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f562])).
fof(f3129,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f3133,plain,(
  ( ! [X0] : (~sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f3159,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X1)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f336])).
fof(f3162,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0,'F_0')) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f578])).
fof(f3165,plain,(
  ( ! [X0] : (sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f579])).
fof(f3169,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0',X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f579])).
fof(f3170,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f579])).
fof(f3180,plain,(
  ( ! [X0] : (~sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),X0)) ) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f338])).
fof(f2298,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl146),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f3355,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(global_subsumption,[],[f343,f342,f341,f2298,f3200,f3192,f3188,f3187,f3186,f3185,f3180,f3179,f3178,f3177,f3176,f3174,f3173,f3172,f3171,f3170,f3169,f3165,f3162,f3160,f3159,f3157,f3156,f3155,f3154,f3152,f3151,f3150,f3149,f3133,f3129,f3128,f3127,f3126,f3121,f3118,f3117,f3112,f3109,f3108,f3105,f3094,f3093,f3092,f3091,f3090,f3089,f3088,f3087,f3086,f3085,f3084,f3083,f3082,f3081,f3080,f3079,f3078,f3077,f3076,f3075,f3074,f3073,f3072,f3071,f3070,f3069,f3068,f3067,f3066,f3065,f3064,f3063,f3062,f3061,f3060,f3059,f3058,f3057,f3056,f3055,f3054,f3053,f3052,f3051,f3050,f3049,f3048,f3047,f3046,f3045,f3044,f3043,f3042,f3041,f3040,f3039,f3038,f3037,f3036,f3034,f3033,f3032,f3031,f3029,f3028,f3027,f3026,f3007,f3006,f3005,f3004,f3002,f3001,f3000,f2999,f2997,f2996,f2995,f2994,f2993,f2991,f2990,f2989,f2988,f2986,f2985,f2984,f2983,f2981,f2980,f2979,f2978,f2977,f2976,f2975,f2974,f2973,f2972,f2971,f2970,f2969,f2968,f2967,f2966,f2965,f2964,f2963,f2962,f2961,f2960,f2959,f2958,f2957,f2956,f2955,f2954,f2953,f2952,f2948,f2947,f2946,f2944,f2943,f2942,f2934,f2933,f2932,f2931,f2929,f3224,f3225,f3226,f3233,f3238,f3239,f3240,f3243,f3250,f3253,f3254,f3256,f3257,f3258,f3259])).
fof(f3226,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(resolution,[],[f2298,f368])).
fof(f368,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.integer.INT_EXP_MOD','thm.bool.EQ_CLAUSES|split|1')).
fof(f3225,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(resolution,[],[f2298,f366])).
fof(f366,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.integer.INT_EXP_MOD','thm.bool.EQ_CLAUSES|split|0')).
fof(f3350,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,sK12)))) | $spl322),
  inference(cnf_transformation,[],[f3350_D])).
fof(f3350_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,sK12)))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3224,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,sK12)))) | $spl146),
  inference(resolution,[],[f2298,f451])).
fof(f3352,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK12)))) | $spl324),
  inference(cnf_transformation,[],[f3352_D])).
fof(f3352_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK12)))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3223,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK12)))) | ($spl84 | $spl146)),
  inference(resolution,[],[f2298,f1809])).
fof(f3354,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl326),
  inference(cnf_transformation,[],[f3354_D])).
fof(f3354_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f3222,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ($spl84 | $spl146)),
  inference(resolution,[],[f2298,f1810])).
fof(f2927,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ($spl84 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f1810])).
fof(f2928,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK12)))) | ($spl84 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f1809])).
fof(f2929,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28(sK12,sK12)))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f451])).
fof(f2931,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f366])).
fof(f2932,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f368])).
fof(f2924,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl259),
  inference(cnf_transformation,[],[f2924_D])).
fof(f2924_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f2933,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f369])).
fof(f369,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f2934,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f371])).
fof(f371,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.integer.INT_EXP_MOD','thm.bool.EQ_CLAUSES|split|2')).
fof(f2942,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f450])).
fof(f594,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f422])).
fof(f422,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : ! [X13] : (s(X12,X13) = s(X12,X13) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_EXP_MOD','thm.bool.REFL_CLAUSE')).
fof(f2943,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f450])).
fof(f345,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_EXP_MOD',aHL_TRUTH)).
fof(f2944,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f450])).
fof(f2946,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f450])).
fof(f2947,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f450])).
fof(f2948,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f450])).
fof(f2912,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl251),
  inference(cnf_transformation,[],[f2912_D])).
fof(f2912_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f2954,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f476])).
fof(f2955,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f476])).
fof(f2956,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f476])).
fof(f2905,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl243),
  inference(cnf_transformation,[],[f2905_D])).
fof(f2905_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f2957,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f476])).
fof(f2910,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl249),
  inference(cnf_transformation,[],[f2910_D])).
fof(f2910_D,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f2958,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f476])).
fof(f2959,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f476])).
fof(f2960,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f476])).
fof(f2961,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f476])).
fof(f2906,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl244),
  inference(cnf_transformation,[],[f2906_D])).
fof(f2906_D,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2962,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f477])).
fof(f2908,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl246),
  inference(cnf_transformation,[],[f2908_D])).
fof(f2908_D,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f2963,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f480])).
fof(f2964,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f481])).
fof(f481,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f2965,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f482])).
fof(f482,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f291])).
fof(f2966,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f483])).
fof(f2967,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f483])).
fof(f2968,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f483])).
fof(f2969,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f483])).
fof(f2970,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f483])).
fof(f2971,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f483])).
fof(f2972,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f483])).
fof(f2973,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f483])).
fof(f2876,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl221),
  inference(cnf_transformation,[],[f2876_D])).
fof(f2876_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f2974,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f2975,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f512])).
fof(f2976,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f512])).
fof(f3346,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl319),
  inference(cnf_transformation,[],[f3346_D])).
fof(f3346_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f2977,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f512])).
fof(f3348,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl321),
  inference(cnf_transformation,[],[f3348_D])).
fof(f3348_D,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl321),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl321])])).
fof(f2978,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f512])).
fof(f2979,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f512])).
fof(f2980,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f512])).
fof(f2981,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f512])).
fof(f3262,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl260),
  inference(cnf_transformation,[],[f3262_D])).
fof(f3262_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f2982,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f678,f676,f667,f2298,f513])).
fof(f676,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f521])).
fof(f3344,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl317),
  inference(cnf_transformation,[],[f3344_D])).
fof(f3344_D,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f2983,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f309])).
fof(f2984,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f516])).
fof(f2985,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f516])).
fof(f3339,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl313),
  inference(cnf_transformation,[],[f3339_D])).
fof(f3339_D,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f2986,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f516])).
fof(f3341,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl315),
  inference(cnf_transformation,[],[f3341_D])).
fof(f3341_D,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl315),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl315])])).
fof(f2988,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f516])).
fof(f2989,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f516])).
fof(f2990,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f516])).
fof(f2991,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f516])).
fof(f3336,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl311),
  inference(cnf_transformation,[],[f3336_D])).
fof(f3336_D,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f2994,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f520])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f312])).
fof(f2995,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f520])).
fof(f2996,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f520])).
fof(f3331,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl307),
  inference(cnf_transformation,[],[f3331_D])).
fof(f3331_D,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f2997,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f520])).
fof(f3333,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl309),
  inference(cnf_transformation,[],[f3333_D])).
fof(f3333_D,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f2999,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f520])).
fof(f3000,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f520])).
fof(f3001,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f520])).
fof(f3002,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f520])).
fof(f2870,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl214),
  inference(cnf_transformation,[],[f2870_D])).
fof(f2870_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f3008,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f925,f2298,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f315])).
fof(f2872,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl216),
  inference(cnf_transformation,[],[f2872_D])).
fof(f2872_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f3009,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f925,f2298,f529])).
fof(f3010,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | ($spl9 | $spl146)),
  inference(unit_resulting_resolution,[],[f851,f2298,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f2874,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl219),
  inference(cnf_transformation,[],[f2874_D])).
fof(f2874_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f3011,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | ($spl9 | $spl146)),
  inference(unit_resulting_resolution,[],[f851,f2298,f530])).
fof(f3266,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl264),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f3014,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f594,f690,f2298,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f317])).
fof(f317,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f201])).
fof(f690,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f540])).
fof(f3015,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f345,f690,f2298,f535])).
fof(f3016,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f344,f690,f2298,f535])).
fof(f3264,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl262),
  inference(cnf_transformation,[],[f3264_D])).
fof(f3264_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f3017,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f690,f2298,f535])).
fof(f3268,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl266),
  inference(cnf_transformation,[],[f3268_D])).
fof(f3268_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f3019,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f594,f690,f2298,f535])).
fof(f3020,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f345,f690,f2298,f535])).
fof(f3021,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f344,f690,f2298,f535])).
fof(f3022,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f690,f2298,f535])).
fof(f3326,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl305),
  inference(cnf_transformation,[],[f3326_D])).
fof(f3326_D,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f3026,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f3027,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f539])).
fof(f3028,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f539])).
fof(f3321,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl301),
  inference(cnf_transformation,[],[f3321_D])).
fof(f3321_D,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f3029,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f539])).
fof(f3323,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl303),
  inference(cnf_transformation,[],[f3323_D])).
fof(f3323_D,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f3031,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f539])).
fof(f3032,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f539])).
fof(f3033,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f539])).
fof(f3034,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f539])).
fof(f3318,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl298),
  inference(cnf_transformation,[],[f3318_D])).
fof(f3318_D,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f3037,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f594,f2298,f545])).
fof(f3038,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f594,f2298,f545])).
fof(f3039,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f594,f2298,f545])).
fof(f3312,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl292),
  inference(cnf_transformation,[],[f3312_D])).
fof(f3312_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f3040,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f594,f2298,f545])).
fof(f3041,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f345,f2298,f545])).
fof(f3042,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f345,f2298,f545])).
fof(f3043,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f345,f2298,f545])).
fof(f3044,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f345,f2298,f545])).
fof(f3045,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f344,f2298,f545])).
fof(f3046,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f344,f2298,f545])).
fof(f3047,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f344,f2298,f545])).
fof(f3048,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f344,f2298,f545])).
fof(f3314,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl294),
  inference(cnf_transformation,[],[f3314_D])).
fof(f3314_D,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f3049,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f2298,f545])).
fof(f3050,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f2298,f545])).
fof(f3051,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f2298,f545])).
fof(f3308,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl288),
  inference(cnf_transformation,[],[f3308_D])).
fof(f3308_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f3052,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f2298,f545])).
fof(f3316,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl296),
  inference(cnf_transformation,[],[f3316_D])).
fof(f3316_D,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f3053,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f594,f2298,f545])).
fof(f3054,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f594,f2298,f545])).
fof(f3055,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f594,f2298,f545])).
fof(f3056,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f594,f2298,f545])).
fof(f3057,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f345,f2298,f545])).
fof(f3058,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f345,f2298,f545])).
fof(f3059,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f345,f2298,f545])).
fof(f3060,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f345,f2298,f545])).
fof(f3061,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f344,f2298,f545])).
fof(f3062,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f344,f2298,f545])).
fof(f3063,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f344,f2298,f545])).
fof(f3064,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f344,f2298,f545])).
fof(f3310,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl290),
  inference(cnf_transformation,[],[f3310_D])).
fof(f3310_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f3065,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f2298,f545])).
fof(f3066,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f2298,f545])).
fof(f3067,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f2298,f545])).
fof(f3068,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f2298,f545])).
fof(f2864,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl210),
  inference(cnf_transformation,[],[f2864_D])).
fof(f2864_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f3069,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f594,f2298,f545])).
fof(f3070,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f594,f2298,f545])).
fof(f3071,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f594,f2298,f545])).
fof(f3072,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f594,f2298,f545])).
fof(f3073,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f345,f2298,f545])).
fof(f3074,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f345,f2298,f545])).
fof(f3075,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f345,f2298,f545])).
fof(f3076,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f345,f2298,f545])).
fof(f3077,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f594,f344,f2298,f545])).
fof(f3078,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f345,f344,f2298,f545])).
fof(f3079,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f344,f344,f2298,f545])).
fof(f3080,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f344,f2298,f545])).
fof(f3081,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f2298,f2298,f545])).
fof(f3082,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f2298,f2298,f545])).
fof(f3083,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f2298,f2298,f545])).
fof(f3084,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f2298,f2298,f545])).
fof(f3306,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl287),
  inference(cnf_transformation,[],[f3306_D])).
fof(f3306_D,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f3087,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f594,f2298,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f322])).
fof(f3088,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f345,f2298,f548])).
fof(f3089,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f344,f2298,f548])).
fof(f3302,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl283),
  inference(cnf_transformation,[],[f3302_D])).
fof(f3302_D,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f3090,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f2298,f548])).
fof(f3304,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl285),
  inference(cnf_transformation,[],[f3304_D])).
fof(f3304_D,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f3091,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f594,f2298,f548])).
fof(f3092,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f345,f2298,f548])).
fof(f3093,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f344,f2298,f548])).
fof(f3094,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f2298,f548])).
fof(f3296,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl281),
  inference(cnf_transformation,[],[f3296_D])).
fof(f3296_D,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f3105,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f328])).
fof(f2849,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl201),
  inference(cnf_transformation,[],[f2849_D])).
fof(f2849_D,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f3128,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f330])).
fof(f3286,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl279),
  inference(cnf_transformation,[],[f3286_D])).
fof(f3286_D,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f3149,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f594,f2298,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f3150,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f345,f2298,f573])).
fof(f3151,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f344,f2298,f573])).
fof(f3281,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl275),
  inference(cnf_transformation,[],[f3281_D])).
fof(f3281_D,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f3152,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f2298,f573])).
fof(f3283,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl277),
  inference(cnf_transformation,[],[f3283_D])).
fof(f3283_D,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f3154,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f594,f350,f2298,f573])).
fof(f3155,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f345,f350,f2298,f573])).
fof(f3156,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f344,f350,f2298,f573])).
fof(f3157,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f2298,f350,f2298,f573])).
fof(f3278,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl272),
  inference(cnf_transformation,[],[f3278_D])).
fof(f3278_D,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f3160,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f577])).
fof(f3275,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl271),
  inference(cnf_transformation,[],[f3275_D])).
fof(f3275_D,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f3171,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f594,f2298,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f3172,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f345,f2298,f580])).
fof(f3173,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f344,f2298,f580])).
fof(f3272,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl269),
  inference(cnf_transformation,[],[f3272_D])).
fof(f3272_D,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f3174,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f2298,f580])).
fof(f2822,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl179),
  inference(cnf_transformation,[],[f2822_D])).
fof(f2822_D,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f3176,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f594,f2298,f580])).
fof(f3177,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f345,f2298,f580])).
fof(f3178,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl146),
  inference(unit_resulting_resolution,[],[f350,f344,f2298,f580])).
fof(f3179,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f2298,f2298,f580])).
fof(f3182,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f594,f925,f2298,f586])).
fof(f586,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X13,X20,X14,X21] : ((s(bool,X13) = s(bool,X20) & (p(s(bool,X20)) => s(bool,X14) = s(bool,X21))) => ((p(s(bool,X13)) => p(s(bool,X14))) <=> (p(s(bool,X20)) => p(s(bool,X21)))))),
  file('thm.integer.INT_EXP_MOD','thm.bool.IMP_CONG')).
fof(f3183,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f345,f925,f2298,f586])).
fof(f3184,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f344,f925,f2298,f586])).
fof(f3185,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f594,f2298,f589])).
fof(f589,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f340])).
fof(f3186,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f345,f2298,f589])).
fof(f3187,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f344,f2298,f589])).
fof(f3188,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f2298,f589])).
fof(f3192,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f2298,f589])).
fof(f3194,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f594,f925,f2298,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f340])).
fof(f3195,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f345,f925,f2298,f590])).
fof(f3196,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f350,f344,f925,f2298,f590])).
fof(f3200,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl146),
  inference(unit_resulting_resolution,[],[f350,f350,f2298,f598])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f587])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f340])).
fof(f2808,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl169),
  inference(cnf_transformation,[],[f2808_D])).
fof(f2808_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f3201,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f825])).
fof(f825,plain,(
  ( ! [X5] : (~sP1(X5,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f604,f476])).
fof(f2810,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl171),
  inference(cnf_transformation,[],[f2810_D])).
fof(f2810_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f3202,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f830])).
fof(f830,plain,(
  ( ! [X10] : (~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X10) | ~p(s(bool,X10))) ) | $spl4),
  inference(resolution,[],[f604,f483])).
fof(f3203,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f594,f2298,f837])).
fof(f837,plain,(
  ( ! [X24,X23] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl4),
  inference(resolution,[],[f604,f545])).
fof(f3204,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f345,f2298,f837])).
fof(f3205,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f344,f2298,f837])).
fof(f3206,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f2298,f837])).
fof(f3208,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f594,f2298,f837])).
fof(f3209,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f345,f2298,f837])).
fof(f3210,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f344,f2298,f837])).
fof(f3211,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f2298,f2298,f837])).
fof(f3216,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl146)),
  inference(unit_resulting_resolution,[],[f676,f2298,f1020])).
fof(f1020,plain,(
  ( ! [X2,X3] : (~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X2,X3) | sP5(X3,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X2) | ~p(s(bool,X2))) ) | $spl4),
  inference(global_subsumption,[],[f343,f342,f341,f604,f810,f809,f808,f807,f806,f805,f804,f803,f802,f801,f800,f799,f798,f797,f796,f795,f794,f793,f792,f791,f787,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f772,f771,f770,f769,f768,f764,f761,f760,f756,f753,f752,f751,f748,f747,f746,f745,f744,f743,f742,f741,f740,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f682,f679,f678,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f665,f664,f663,f662,f661,f660,f659,f658,f657,f656,f655,f654,f653,f652,f651,f650,f649,f648,f647,f646,f645,f644,f643,f642,f641,f640,f639,f638,f637,f636,f635,f634,f633,f632,f631,f630,f629,f625,f624,f623,f621,f620,f619,f611,f610,f609,f608,f606,f811,f812,f813,f820,f825,f826,f827,f830,f837,f840,f841,f843,f844,f845,f846,f928,f847,f929,f1013])).
fof(f1013,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | sP5(X3,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X2) | ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X2,X3) | ~sP2(X2,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X3)) ) | $spl4),
  inference(resolution,[],[f667,f513])).
fof(f847,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(superposition,[],[f604,f352])).
fof(f846,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(superposition,[],[f604,f353])).
fof(f845,plain,(
  ( ! [X39,X40] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X39,X40) | ~p(s(bool,X39))) ) | $spl4),
  inference(resolution,[],[f604,f579])).
fof(f844,plain,(
  ( ! [X37,X38] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X38,X37) | p(s(bool,X37))) ) | $spl4),
  inference(resolution,[],[f604,f578])).
fof(f843,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl4),
  inference(resolution,[],[f604,f577])).
fof(f841,plain,(
  ( ! [X31,X32] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X31,X32) | ~p(s(bool,X31))) ) | $spl4),
  inference(resolution,[],[f604,f562])).
fof(f840,plain,(
  ( ! [X30,X29] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X30,X29) | ~p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f604,f561])).
fof(f827,plain,(
  ( ! [X7] : (sP1(X7,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | p(s(bool,X7))) ) | $spl4),
  inference(resolution,[],[f604,f480])).
fof(f826,plain,(
  ( ! [X6] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X6) | p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f604,f477])).
fof(f820,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f604,f450])).
fof(f629,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f466])).
fof(f630,plain,(
  ( ! [X0] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f467])).
fof(f677,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f525])).
fof(f679,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f527])).
fof(f739,plain,(
  ( ! [X0] : (~sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f546])).
fof(f740,plain,(
  ( ! [X0] : (~sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f547])).
fof(f752,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f557])).
fof(f753,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f558])).
fof(f756,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f561])).
fof(f760,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f561])).
fof(f761,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f561])).
fof(f764,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f562])).
fof(f768,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f562])).
fof(f769,plain,(
  ( ! [X0] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f562])).
fof(f771,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f564])).
fof(f772,plain,(
  ( ! [X0] : (~sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f565])).
fof(f782,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f575])).
fof(f784,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f578])).
fof(f787,plain,(
  ( ! [X0] : (sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f579])).
fof(f791,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f579])).
fof(f792,plain,(
  ( ! [X0] : (sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f579])).
fof(f801,plain,(
  ( ! [X0] : (~sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f582])).
fof(f2800,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl160),
  inference(cnf_transformation,[],[f2800_D])).
fof(f2800_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f3220,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f925,f2298,f1633])).
fof(f1633,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') != s(bool,X33) | sP5(X34,X33,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,X34))) ) | $spl82),
  inference(superposition,[],[f529,f925])).
fof(f3221,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl146)),
  inference(unit_resulting_resolution,[],[f925,f2298,f1634])).
fof(f1634,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP5(X36,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X35) | ~p(s(bool,X36))) ) | $spl82),
  inference(superposition,[],[f529,f925])).
fof(f2926,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl84),
  inference(forward_demodulation,[],[f2333,f927])).
fof(f2333,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl84),
  inference(superposition,[],[f1810,f927])).
fof(f2925,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl258),
  inference(cnf_transformation,[],[f2925_D])).
fof(f2925_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f2775,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(resolution,[],[f2299,f372])).
fof(f372,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f218])).
fof(f2774,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(resolution,[],[f2299,f370])).
fof(f2813,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl173),
  inference(cnf_transformation,[],[f2813_D])).
fof(f2813_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f2445,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f365])).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f2446,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f367])).
fof(f367,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f2447,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f370])).
fof(f2448,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f372])).
fof(f2457,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f447])).
fof(f2459,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f447])).
fof(f2921,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl257),
  inference(cnf_transformation,[],[f2921_D])).
fof(f2921_D,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f2461,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f460])).
fof(f460,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f2916,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl253),
  inference(cnf_transformation,[],[f2916_D])).
fof(f2916_D,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f2462,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f460])).
fof(f2918,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl255),
  inference(cnf_transformation,[],[f2918_D])).
fof(f2918_D,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f2464,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f460])).
fof(f2465,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f460])).
fof(f2467,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f285])).
fof(f2468,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f464])).
fof(f2470,plain,(
  ~sP0('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f464])).
fof(f2471,plain,(
  ~sP0('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f464])).
fof(f2907,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl245),
  inference(cnf_transformation,[],[f2907_D])).
fof(f2907_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f2473,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f475])).
fof(f475,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f290])).
fof(f2474,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f475])).
fof(f2909,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl247),
  inference(cnf_transformation,[],[f2909_D])).
fof(f2909_D,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl247),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl247])])).
fof(f2475,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f475])).
fof(f2476,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f475])).
fof(f2911,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl248),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f2477,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f477])).
fof(f2478,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f477])).
fof(f2479,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f477])).
fof(f2913,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl250),
  inference(cnf_transformation,[],[f2913_D])).
fof(f2913_D,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f2480,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f480])).
fof(f2481,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f480])).
fof(f2482,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f480])).
fof(f2483,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f481])).
fof(f2484,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f481])).
fof(f2485,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f481])).
fof(f2486,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f482])).
fof(f2487,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f482])).
fof(f2488,plain,(
  sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f482])).
fof(f2489,plain,(
  ~sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f2490,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f484])).
fof(f2491,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f484])).
fof(f2492,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f484])).
fof(f2903,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl241),
  inference(cnf_transformation,[],[f2903_D])).
fof(f2903_D,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f2493,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f512])).
fof(f2494,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f512])).
fof(f2495,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f512])).
fof(f2496,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f512])).
fof(f2497,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f512])).
fof(f2498,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f512])).
fof(f2499,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f512])).
fof(f2500,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f512])).
fof(f2501,plain,(
  ~sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f512])).
fof(f2901,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl239),
  inference(cnf_transformation,[],[f2901_D])).
fof(f2901_D,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f2502,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f516])).
fof(f2503,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f516])).
fof(f2504,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f516])).
fof(f2505,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f516])).
fof(f2506,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f516])).
fof(f2507,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f516])).
fof(f2508,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f516])).
fof(f2509,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f516])).
fof(f2510,plain,(
  ~sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f516])).
fof(f2899,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl237),
  inference(cnf_transformation,[],[f2899_D])).
fof(f2899_D,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f2513,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f520])).
fof(f2514,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f520])).
fof(f2515,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f520])).
fof(f2516,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f520])).
fof(f2517,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f520])).
fof(f2518,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f520])).
fof(f2519,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f520])).
fof(f2520,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f520])).
fof(f2521,plain,(
  ~sP3('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f520])).
fof(f2897,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl235),
  inference(cnf_transformation,[],[f2897_D])).
fof(f2897_D,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f2524,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f2885,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl227),
  inference(cnf_transformation,[],[f2885_D])).
fof(f2885_D,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f2525,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f524])).
fof(f2891,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl231),
  inference(cnf_transformation,[],[f2891_D])).
fof(f2891_D,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f2526,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f524])).
fof(f2881,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl223),
  inference(cnf_transformation,[],[f2881_D])).
fof(f2881_D,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f2527,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f524])).
fof(f2893,plain,(
  ~sP2('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl233),
  inference(cnf_transformation,[],[f2893_D])).
fof(f2893_D,plain,(
  ~sP2('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f2530,plain,(
  ~sP2('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f524])).
fof(f2531,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f524])).
fof(f2883,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl225),
  inference(cnf_transformation,[],[f2883_D])).
fof(f2883_D,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f2532,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f524])).
fof(f2533,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f524])).
fof(f2887,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl229),
  inference(cnf_transformation,[],[f2887_D])).
fof(f2887_D,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl229),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl229])])).
fof(f2536,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f524])).
fof(f2537,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f524])).
fof(f2538,plain,(
  ~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f524])).
fof(f2539,plain,(
  ~sP2('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f524])).
fof(f2877,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | $spl220),
  inference(cnf_transformation,[],[f2877_D])).
fof(f2877_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f2542,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','F_0') | ($spl9 | $spl147)),
  inference(unit_resulting_resolution,[],[f851,f2299,f528])).
fof(f2875,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl218),
  inference(cnf_transformation,[],[f2875_D])).
fof(f2875_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f2543,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | ($spl9 | $spl147)),
  inference(unit_resulting_resolution,[],[f851,f2299,f528])).
fof(f2873,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl217),
  inference(cnf_transformation,[],[f2873_D])).
fof(f2873_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2544,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f925,f2299,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f315])).
fof(f2871,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl215),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f2545,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f925,f2299,f531])).
fof(f2869,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl213),
  inference(cnf_transformation,[],[f2869_D])).
fof(f2869_D,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f2552,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f539])).
fof(f2553,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f539])).
fof(f2554,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f539])).
fof(f2555,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f539])).
fof(f2556,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f539])).
fof(f2557,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f539])).
fof(f2558,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f539])).
fof(f2559,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f539])).
fof(f2560,plain,(
  ~sP6('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f539])).
fof(f2865,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl211),
  inference(cnf_transformation,[],[f2865_D])).
fof(f2865_D,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f2579,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f548])).
fof(f2580,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f548])).
fof(f2581,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f548])).
fof(f2582,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f548])).
fof(f2583,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f548])).
fof(f2584,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f548])).
fof(f2585,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f548])).
fof(f2586,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f548])).
fof(f2587,plain,(
  ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f548])).
fof(f2861,plain,(
  ~sP8('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl209),
  inference(cnf_transformation,[],[f2861_D])).
fof(f2861_D,plain,(
  ~sP8('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f2604,plain,(
  ~sP8('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f350,f2299,f556])).
fof(f2605,plain,(
  ~sP8('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f350,f2299,f556])).
fof(f2606,plain,(
  ~sP8('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f350,f2299,f556])).
fof(f2856,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl205),
  inference(cnf_transformation,[],[f2856_D])).
fof(f2856_D,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f2607,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f2299,f556])).
fof(f2608,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f2299,f556])).
fof(f2609,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f2299,f556])).
fof(f2858,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl207),
  inference(cnf_transformation,[],[f2858_D])).
fof(f2858_D,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f2611,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f350,f2299,f556])).
fof(f2612,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f350,f2299,f556])).
fof(f2613,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f350,f2299,f556])).
fof(f2614,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f2299,f556])).
fof(f2615,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f2299,f556])).
fof(f2616,plain,(
  ~sP8('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f2299,f556])).
fof(f2853,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl202),
  inference(cnf_transformation,[],[f2853_D])).
fof(f2853_D,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f2619,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f560])).
fof(f2848,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl198),
  inference(cnf_transformation,[],[f2848_D])).
fof(f2848_D,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f2620,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f560])).
fof(f2842,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl192),
  inference(cnf_transformation,[],[f2842_D])).
fof(f2842_D,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f2621,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f560])).
fof(f2840,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl190),
  inference(cnf_transformation,[],[f2840_D])).
fof(f2840_D,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f2622,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f560])).
fof(f2850,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl200),
  inference(cnf_transformation,[],[f2850_D])).
fof(f2850_D,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f2624,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f560])).
fof(f2625,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f560])).
fof(f2844,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl194),
  inference(cnf_transformation,[],[f2844_D])).
fof(f2844_D,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f2626,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f560])).
fof(f2627,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f560])).
fof(f2846,plain,(
  sP9('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl196),
  inference(cnf_transformation,[],[f2846_D])).
fof(f2846_D,plain,(
  sP9('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f2628,plain,(
  sP9('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f350,f2299,f560])).
fof(f2629,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f350,f2299,f560])).
fof(f2630,plain,(
  sP9('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f2299,f2299,f560])).
fof(f2631,plain,(
  sP9('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f2299,f2299,f560])).
fof(f2837,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl189),
  inference(cnf_transformation,[],[f2837_D])).
fof(f2837_D,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f2633,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f2299,f563])).
fof(f2832,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl185),
  inference(cnf_transformation,[],[f2832_D])).
fof(f2832_D,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f2634,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f594,f2299,f563])).
fof(f2635,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f2299,f563])).
fof(f2636,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f345,f2299,f563])).
fof(f2637,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f2299,f563])).
fof(f2638,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f344,f2299,f563])).
fof(f2834,plain,(
  ~sP9('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl187),
  inference(cnf_transformation,[],[f2834_D])).
fof(f2834_D,plain,(
  ~sP9('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f2640,plain,(
  ~sP9('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f2299,f563])).
fof(f2641,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f594,f2299,f563])).
fof(f2642,plain,(
  ~sP9('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f2299,f563])).
fof(f2643,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f345,f2299,f563])).
fof(f2644,plain,(
  ~sP9('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f2299,f563])).
fof(f2645,plain,(
  ~sP9('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f344,f2299,f563])).
fof(f2827,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl183),
  inference(cnf_transformation,[],[f2827_D])).
fof(f2827_D,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f2663,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f573])).
fof(f2664,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f573])).
fof(f2665,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f573])).
fof(f2666,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f573])).
fof(f2667,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f573])).
fof(f2668,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f573])).
fof(f2669,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f573])).
fof(f2670,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f573])).
fof(f2671,plain,(
  ~sP10('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f573])).
fof(f2825,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl180),
  inference(cnf_transformation,[],[f2825_D])).
fof(f2825_D,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f2674,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f2299,f577])).
fof(f2821,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl176),
  inference(cnf_transformation,[],[f2821_D])).
fof(f2821_D,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f2675,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f594,f2299,f577])).
fof(f2676,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f2299,f577])).
fof(f2677,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f345,f2299,f577])).
fof(f2678,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f2299,f577])).
fof(f2679,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f2299,f344,f2299,f577])).
fof(f2823,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl178),
  inference(cnf_transformation,[],[f2823_D])).
fof(f2823_D,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f2680,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f350,f2299,f577])).
fof(f2681,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f350,f2299,f577])).
fof(f2682,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f350,f2299,f577])).
fof(f2683,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f2299,f577])).
fof(f2684,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f2299,f577])).
fof(f2685,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f2299,f577])).
fof(f2819,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl175),
  inference(cnf_transformation,[],[f2819_D])).
fof(f2819_D,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f2690,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f594,f2299,f580])).
fof(f2691,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f594,f2299,f580])).
fof(f2692,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f594,f2299,f580])).
fof(f2693,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f345,f2299,f580])).
fof(f2694,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f345,f2299,f580])).
fof(f2695,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f345,f2299,f580])).
fof(f2696,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f594,f344,f2299,f580])).
fof(f2697,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f345,f344,f2299,f580])).
fof(f2698,plain,(
  ~sP11('T_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'T_0') | $spl147),
  inference(unit_resulting_resolution,[],[f344,f344,f2299,f580])).
fof(f2709,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f594,f925,f2299,f586])).
fof(f2710,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f594,f925,f2299,f586])).
fof(f2711,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f594,f925,f2299,f586])).
fof(f2712,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f345,f925,f2299,f586])).
fof(f2713,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f345,f925,f2299,f586])).
fof(f2714,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f345,f925,f2299,f586])).
fof(f2715,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f344,f925,f2299,f586])).
fof(f2716,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f344,f925,f2299,f586])).
fof(f2717,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f344,f925,f2299,f586])).
fof(f2721,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f594,f2299,f589])).
fof(f2722,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f594,f2299,f589])).
fof(f2723,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f594,f2299,f589])).
fof(f2727,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f345,f2299,f589])).
fof(f2728,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f345,f2299,f589])).
fof(f2729,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f345,f2299,f589])).
fof(f2733,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f344,f2299,f589])).
fof(f2734,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f344,f2299,f589])).
fof(f2735,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f344,f2299,f589])).
fof(f2737,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f594,f2299,f589])).
fof(f2738,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f594,f2299,f589])).
fof(f2739,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f594,f2299,f589])).
fof(f2740,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f345,f2299,f589])).
fof(f2741,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f345,f2299,f589])).
fof(f2742,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f345,f2299,f589])).
fof(f2743,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f344,f2299,f589])).
fof(f2744,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f344,f2299,f589])).
fof(f2745,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f344,f2299,f589])).
fof(f2746,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f594,f925,f2299,f590])).
fof(f2747,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f594,f925,f2299,f590])).
fof(f2748,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f594,f925,f2299,f590])).
fof(f2749,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f345,f925,f2299,f590])).
fof(f2750,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f345,f925,f2299,f590])).
fof(f2751,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f345,f925,f2299,f590])).
fof(f2752,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f594,f344,f925,f2299,f590])).
fof(f2753,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f345,f344,f925,f2299,f590])).
fof(f2754,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f344,f344,f925,f2299,f590])).
fof(f2758,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f594,f2299,f2299,f598])).
fof(f2759,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f345,f2299,f2299,f598])).
fof(f2760,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f344,f2299,f2299,f598])).
fof(f2762,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f594,f2299,f598])).
fof(f2763,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f345,f2299,f598])).
fof(f2764,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl147),
  inference(unit_resulting_resolution,[],[f350,f344,f2299,f598])).
fof(f2811,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | $spl170),
  inference(cnf_transformation,[],[f2811_D])).
fof(f2811_D,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f2765,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f2299,f826])).
fof(f2809,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl168),
  inference(cnf_transformation,[],[f2809_D])).
fof(f2809_D,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f2766,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f2299,f827])).
fof(f2807,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl166),
  inference(cnf_transformation,[],[f2807_D])).
fof(f2807_D,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f2767,plain,(
  sP11('F_0','const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f350,f2299,f843])).
fof(f2803,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl162),
  inference(cnf_transformation,[],[f2803_D])).
fof(f2803_D,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f2768,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f2299,f2299,f843])).
fof(f2805,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl164),
  inference(cnf_transformation,[],[f2805_D])).
fof(f2805_D,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f2769,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f350,f2299,f843])).
fof(f2770,plain,(
  sP11('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl4 | $spl147)),
  inference(unit_resulting_resolution,[],[f2299,f2299,f843])).
fof(f2801,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl161),
  inference(cnf_transformation,[],[f2801_D])).
fof(f2801_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f2772,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f925,f2299,f1635])).
fof(f1635,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X38)) | ~sP5(X38,X37,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) ) | $spl82),
  inference(superposition,[],[f531,f925])).
fof(f2773,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ($spl82 | $spl147)),
  inference(unit_resulting_resolution,[],[f925,f2299,f1636])).
fof(f1636,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X40)) | ~sP5(X40,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),X39)) ) | $spl82),
  inference(superposition,[],[f531,f925])).
fof(f2433,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28(sK12,sK13)))))) | $spl158),
  inference(cnf_transformation,[],[f2433_D])).
fof(f2433_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28(sK12,sK13)))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f2405,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))),s('type.num.num',sK28(sK12,sK13)))))) | ($spl1 | $spl84)),
  inference(superposition,[],[f2023,f927])).
fof(f2383,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | $spl156),
  inference(cnf_transformation,[],[f2383_D])).
fof(f2383_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f2353,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK28(sK12,sK13)))))) | ($spl3 | $spl84)),
  inference(superposition,[],[f1817,f927])).
fof(f2379,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK12))) | $spl154),
  inference(cnf_transformation,[],[f2379_D])).
fof(f2379_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK12))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f2381,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK12))) | ($spl3 | $spl84 | $spl127)),
  inference(forward_demodulation,[],[f2380,f1817])).
fof(f2380,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1))))),s('type.num.num',sK12)))) ) | ($spl3 | $spl84 | $spl127)),
  inference(forward_demodulation,[],[f2346,f927])).
fof(f2346,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) ) | ($spl3 | $spl127)),
  inference(unit_resulting_resolution,[],[f1814,f1817,f455])).
fof(f2377,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK12))) | ($spl3 | $spl143)),
  inference(forward_demodulation,[],[f2347,f1817])).
fof(f2347,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',X1))))),s('type.num.num',sK12)))) ) | ($spl3 | $spl143)),
  inference(unit_resulting_resolution,[],[f2162,f1817,f455])).
fof(f2306,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f2303,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) | $spl150),
  inference(cnf_transformation,[],[f2303_D])).
fof(f2303_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f2294,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),sK13)))) | ($spl82 | $spl84)),
  inference(superposition,[],[f1809,f925])).
fof(f2297,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | $spl84),
  inference(forward_demodulation,[],[f2293,f927])).
fof(f2293,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))),s('type.num.num',sK28('const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))),'const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))))) | $spl84),
  inference(superposition,[],[f1809,f927])).
fof(f2212,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl144),
  inference(cnf_transformation,[],[f2212_D])).
fof(f2212_D,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f2199,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f594,f925,f1633])).
fof(f2200,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f345,f925,f1633])).
fof(f2201,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f344,f925,f1633])).
fof(f2162,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | $spl143),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f2160,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK12) | ($spl84 | $spl127)),
  inference(superposition,[],[f1814,f927])).
fof(f2114,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl140),
  inference(cnf_transformation,[],[f2114_D])).
fof(f2114_D,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2110,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl124),
  inference(resolution,[],[f1700,f510])).
fof(f2106,plain,(
  sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl124),
  inference(unit_resulting_resolution,[],[f1700,f510])).
fof(f2099,plain,(
  sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl136),
  inference(cnf_transformation,[],[f2099_D])).
fof(f2099_D,plain,(
  sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f2094,plain,(
  sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl122),
  inference(resolution,[],[f1686,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f306])).
fof(f2101,plain,(
  sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl138),
  inference(cnf_transformation,[],[f2101_D])).
fof(f2101_D,plain,(
  sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2093,plain,(
  sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl122),
  inference(resolution,[],[f1686,f510])).
fof(f2088,plain,(
  sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl122),
  inference(unit_resulting_resolution,[],[f1686,f510])).
fof(f2089,plain,(
  sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl122),
  inference(unit_resulting_resolution,[],[f1686,f511])).
fof(f1958,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl135),
  inference(cnf_transformation,[],[f1958_D])).
fof(f1958_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1954,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl61),
  inference(unit_resulting_resolution,[],[f903,f511])).
fof(f1944,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl133),
  inference(cnf_transformation,[],[f1944_D])).
fof(f1944_D,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f1940,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl55),
  inference(unit_resulting_resolution,[],[f897,f510])).
fof(f1920,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl130),
  inference(cnf_transformation,[],[f1920_D])).
fof(f1920_D,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1913,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl42),
  inference(resolution,[],[f885,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f318])).
fof(f1907,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f885,f532])).
fof(f1814,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) | $spl127),
  inference(cnf_transformation,[],[f1814_D])).
fof(f1814_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f1798,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X12),s('type.num.num',sK12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13))))) ) | $spl84),
  inference(superposition,[],[f454,f927])).
fof(f1700,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl124),
  inference(cnf_transformation,[],[f1700_D])).
fof(f1700_D,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1548,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f594,f925,f529])).
fof(f1549,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f345,f925,f529])).
fof(f1550,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f344,f925,f529])).
fof(f1686,plain,(
  sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl122),
  inference(cnf_transformation,[],[f1686_D])).
fof(f1686_D,plain,(
  sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1573,plain,(
  sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f594,f925,f529])).
fof(f1574,plain,(
  sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f345,f925,f529])).
fof(f1575,plain,(
  sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(unit_resulting_resolution,[],[f344,f925,f529])).
fof(f1502,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl121),
  inference(cnf_transformation,[],[f1502_D])).
fof(f1502_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f1498,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f907,f511])).
fof(f1492,plain,(
  ~sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl119),
  inference(cnf_transformation,[],[f1492_D])).
fof(f1492_D,plain,(
  ~sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1488,plain,(
  ~sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl63),
  inference(unit_resulting_resolution,[],[f905,f511])).
fof(f1482,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl117),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1478,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl59),
  inference(unit_resulting_resolution,[],[f901,f510])).
fof(f1472,plain,(
  ~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl115),
  inference(cnf_transformation,[],[f1472_D])).
fof(f1472_D,plain,(
  ~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1468,plain,(
  ~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl57),
  inference(unit_resulting_resolution,[],[f899,f510])).
fof(f1445,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl112),
  inference(cnf_transformation,[],[f1445_D])).
fof(f1445_D,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1440,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl46),
  inference(resolution,[],[f889,f532])).
fof(f1434,plain,(
  sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f889,f532])).
fof(f1423,plain,(
  sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl110),
  inference(cnf_transformation,[],[f1423_D])).
fof(f1423_D,plain,(
  sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1418,plain,(
  sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl44),
  inference(resolution,[],[f887,f532])).
fof(f1412,plain,(
  sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl44),
  inference(unit_resulting_resolution,[],[f887,f532])).
fof(f1349,plain,(
  sP4('F_0','T_0','F_0') | $spl104),
  inference(cnf_transformation,[],[f1349_D])).
fof(f1349_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1345,plain,(
  sP4('F_0','T_0','F_0') | $spl92),
  inference(resolution,[],[f965,f511])).
fof(f1351,plain,(
  sP3('F_0','T_0','F_0') | $spl106),
  inference(cnf_transformation,[],[f1351_D])).
fof(f1351_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1344,plain,(
  sP3('F_0','T_0','F_0') | $spl92),
  inference(resolution,[],[f965,f510])).
fof(f1353,plain,(
  sP2('T_0','F_0','F_0') | $spl108),
  inference(cnf_transformation,[],[f1353_D])).
fof(f1353_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1343,plain,(
  sP2('T_0','F_0','F_0') | $spl92),
  inference(resolution,[],[f965,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f306])).
fof(f1339,plain,(
  sP2('T_0','F_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f965,f509])).
fof(f1340,plain,(
  sP3('F_0','T_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f965,f510])).
fof(f1341,plain,(
  sP4('F_0','T_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f965,f511])).
fof(f1332,plain,(
  sP10('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl102),
  inference(cnf_transformation,[],[f1332_D])).
fof(f1332_D,plain,(
  sP10('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1330,plain,(
  sP10('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl16),
  inference(resolution,[],[f859,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f333])).
fof(f1324,plain,(
  sP10('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl16),
  inference(unit_resulting_resolution,[],[f859,f568])).
fof(f1112,plain,(
  sP4('F_0','F_0','T_0') | $spl96),
  inference(cnf_transformation,[],[f1112_D])).
fof(f1112_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1109,plain,(
  sP4('F_0','F_0','T_0') | $spl88),
  inference(resolution,[],[f961,f511])).
fof(f1114,plain,(
  sP3('T_0','F_0','F_0') | $spl98),
  inference(cnf_transformation,[],[f1114_D])).
fof(f1114_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1108,plain,(
  sP3('T_0','F_0','F_0') | $spl88),
  inference(resolution,[],[f961,f510])).
fof(f1116,plain,(
  sP2('F_0','T_0','F_0') | $spl100),
  inference(cnf_transformation,[],[f1116_D])).
fof(f1116_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1107,plain,(
  sP2('F_0','T_0','F_0') | $spl88),
  inference(resolution,[],[f961,f509])).
fof(f1102,plain,(
  sP2('F_0','T_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f961,f509])).
fof(f1103,plain,(
  sP3('T_0','F_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f961,f510])).
fof(f1104,plain,(
  sP4('F_0','F_0','T_0') | $spl88),
  inference(unit_resulting_resolution,[],[f961,f511])).
fof(f1036,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl94),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1031,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f667,f676,f677,f513])).
fof(f1032,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f667,f676,f677,f513])).
fof(f1033,plain,(
  sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f667,f676,f677,f513])).
fof(f965,plain,(
  sP5('F_0','F_0','T_0') | $spl92),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f938,plain,(
  sP5('F_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f350,f851,f528])).
fof(f963,plain,(
  ~sP5('T_0','F_0','T_0') | $spl91),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f939,plain,(
  ~sP5('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f594,f851,f530])).
fof(f940,plain,(
  ~sP5('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f345,f851,f530])).
fof(f941,plain,(
  ~sP5('T_0','F_0','T_0') | $spl9),
  inference(unit_resulting_resolution,[],[f344,f851,f530])).
fof(f961,plain,(
  sP5('F_0','T_0','F_0') | $spl88),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f948,plain,(
  sP5('F_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f350,f851,f528])).
fof(f959,plain,(
  ~sP5('T_0','T_0','F_0') | $spl87),
  inference(cnf_transformation,[],[f959_D])).
fof(f959_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f949,plain,(
  ~sP5('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f594,f851,f530])).
fof(f950,plain,(
  ~sP5('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f345,f851,f530])).
fof(f951,plain,(
  ~sP5('T_0','T_0','F_0') | $spl9),
  inference(unit_resulting_resolution,[],[f344,f851,f530])).
fof(f925,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl82),
  inference(cnf_transformation,[],[f925_D])).
fof(f925_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f929,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(global_subsumption,[],[f343,f342,f341,f604,f810,f809,f808,f807,f806,f805,f804,f803,f802,f801,f800,f799,f798,f797,f796,f795,f794,f793,f792,f791,f787,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f772,f771,f770,f769,f768,f764,f761,f760,f756,f753,f752,f751,f748,f747,f746,f745,f744,f743,f742,f741,f740,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f682,f679,f678,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f665,f664,f663,f662,f661,f660,f659,f658,f657,f656,f655,f654,f653,f652,f651,f650,f649,f648,f647,f646,f645,f644,f643,f642,f641,f640,f639,f638,f637,f636,f635,f634,f633,f632,f631,f630,f629,f625,f624,f623,f621,f620,f619,f611,f610,f609,f608,f606,f811,f812,f813,f820,f825,f826,f827,f830,f837,f840,f841,f843,f844,f845,f846,f928,f847])).
fof(f928,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(global_subsumption,[],[f343,f342,f341,f604,f810,f809,f808,f807,f806,f805,f804,f803,f802,f801,f800,f799,f798,f797,f796,f795,f794,f793,f792,f791,f787,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f772,f771,f770,f769,f768,f764,f761,f760,f756,f753,f752,f751,f748,f747,f746,f745,f744,f743,f742,f741,f740,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f682,f679,f678,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f665,f664,f663,f662,f661,f660,f659,f658,f657,f656,f655,f654,f653,f652,f651,f650,f649,f648,f647,f646,f645,f644,f643,f642,f641,f640,f639,f638,f637,f636,f635,f634,f633,f632,f631,f630,f629,f625,f624,f623,f621,f620,f619,f611,f610,f609,f608,f606,f811,f812,f813,f820,f825,f826,f827,f830,f837,f840,f841,f843,f844,f845,f846])).
fof(f813,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(resolution,[],[f604,f368])).
fof(f812,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(resolution,[],[f604,f366])).
fof(f927,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) | $spl84),
  inference(cnf_transformation,[],[f927_D])).
fof(f927_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f811,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) | $spl4),
  inference(resolution,[],[f604,f451])).
fof(f606,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',sK28(sK12,sK13)))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f451])).
fof(f608,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f366])).
fof(f609,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f368])).
fof(f849,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl7),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f610,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f369])).
fof(f611,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f371])).
fof(f619,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f450])).
fof(f620,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f450])).
fof(f621,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f450])).
fof(f623,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f450])).
fof(f624,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f450])).
fof(f625,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f450])).
fof(f917,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl75),
  inference(cnf_transformation,[],[f917_D])).
fof(f917_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f631,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f476])).
fof(f632,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f476])).
fof(f633,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f476])).
fof(f915,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl73),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f634,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f476])).
fof(f919,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl77),
  inference(cnf_transformation,[],[f919_D])).
fof(f919_D,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f635,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f476])).
fof(f636,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f476])).
fof(f637,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f476])).
fof(f638,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f476])).
fof(f921,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl78),
  inference(cnf_transformation,[],[f921_D])).
fof(f921_D,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f639,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f477])).
fof(f923,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl80),
  inference(cnf_transformation,[],[f923_D])).
fof(f923_D,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f640,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f480])).
fof(f641,plain,(
  sP1('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f481])).
fof(f642,plain,(
  sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f482])).
fof(f643,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f483])).
fof(f644,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f483])).
fof(f645,plain,(
  ~sP1('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f483])).
fof(f646,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f483])).
fof(f647,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f483])).
fof(f648,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f483])).
fof(f649,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f483])).
fof(f650,plain,(
  ~sP1('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f483])).
fof(f913,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl71),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f651,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f512])).
fof(f652,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f512])).
fof(f653,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f512])).
fof(f909,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl67),
  inference(cnf_transformation,[],[f909_D])).
fof(f909_D,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f654,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f512])).
fof(f911,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl69),
  inference(cnf_transformation,[],[f911_D])).
fof(f911_D,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f655,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f512])).
fof(f656,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f512])).
fof(f657,plain,(
  ~sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f512])).
fof(f658,plain,(
  ~sP5('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f512])).
fof(f907,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl65),
  inference(cnf_transformation,[],[f907_D])).
fof(f907_D,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f659,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f516])).
fof(f660,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f516])).
fof(f661,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f516])).
fof(f903,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl61),
  inference(cnf_transformation,[],[f903_D])).
fof(f903_D,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f662,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f516])).
fof(f905,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl63),
  inference(cnf_transformation,[],[f905_D])).
fof(f905_D,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f663,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f516])).
fof(f664,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f516])).
fof(f665,plain,(
  ~sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f516])).
fof(f666,plain,(
  ~sP4('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f516])).
fof(f901,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl59),
  inference(cnf_transformation,[],[f901_D])).
fof(f901_D,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f668,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f520])).
fof(f669,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f520])).
fof(f670,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f520])).
fof(f897,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl55),
  inference(cnf_transformation,[],[f897_D])).
fof(f897_D,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f671,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f520])).
fof(f899,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl57),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f672,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f520])).
fof(f673,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f520])).
fof(f674,plain,(
  ~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f520])).
fof(f675,plain,(
  ~sP3('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f520])).
fof(f895,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl53),
  inference(cnf_transformation,[],[f895_D])).
fof(f895_D,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f682,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f539])).
fof(f683,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f539])).
fof(f684,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f539])).
fof(f891,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl49),
  inference(cnf_transformation,[],[f891_D])).
fof(f891_D,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f685,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f539])).
fof(f893,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl51),
  inference(cnf_transformation,[],[f893_D])).
fof(f893_D,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f686,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f539])).
fof(f687,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f539])).
fof(f688,plain,(
  ~sP6('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f539])).
fof(f689,plain,(
  ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f539])).
fof(f889,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl46),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f691,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f594,f604,f545])).
fof(f692,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f594,f604,f545])).
fof(f693,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f594,f604,f545])).
fof(f881,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl38),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f694,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f594,f604,f545])).
fof(f695,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f345,f604,f545])).
fof(f696,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f545])).
fof(f697,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f545])).
fof(f698,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f545])).
fof(f699,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f344,f604,f545])).
fof(f700,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f545])).
fof(f701,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f545])).
fof(f702,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f545])).
fof(f885,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl42),
  inference(cnf_transformation,[],[f885_D])).
fof(f885_D,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f703,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f604,f545])).
fof(f704,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f545])).
fof(f705,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f545])).
fof(f877,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl34),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f706,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f545])).
fof(f887,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl44),
  inference(cnf_transformation,[],[f887_D])).
fof(f887_D,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f707,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f594,f604,f545])).
fof(f708,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f594,f604,f545])).
fof(f709,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f594,f604,f545])).
fof(f710,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f594,f604,f545])).
fof(f711,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f345,f604,f545])).
fof(f712,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f545])).
fof(f713,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f545])).
fof(f714,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f545])).
fof(f715,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f344,f604,f545])).
fof(f716,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f545])).
fof(f717,plain,(
  sP7('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f545])).
fof(f718,plain,(
  sP7('T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f545])).
fof(f879,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl36),
  inference(cnf_transformation,[],[f879_D])).
fof(f879_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f719,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f604,f545])).
fof(f720,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f545])).
fof(f721,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f545])).
fof(f722,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f545])).
fof(f883,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl40),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f723,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f594,f604,f545])).
fof(f724,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f594,f604,f545])).
fof(f725,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f594,f604,f545])).
fof(f726,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f594,f604,f545])).
fof(f727,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f345,f604,f545])).
fof(f728,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f545])).
fof(f729,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f545])).
fof(f730,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f545])).
fof(f731,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f594,f344,f604,f545])).
fof(f732,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f545])).
fof(f733,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f545])).
fof(f734,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f545])).
fof(f735,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f604,f604,f545])).
fof(f736,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f545])).
fof(f737,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f545])).
fof(f738,plain,(
  sP7('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f545])).
fof(f875,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl33),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f741,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f594,f604,f548])).
fof(f742,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f345,f604,f548])).
fof(f743,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f344,f604,f548])).
fof(f871,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl29),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f744,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f604,f548])).
fof(f873,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl31),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f745,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f594,f604,f548])).
fof(f746,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f345,f604,f548])).
fof(f747,plain,(
  ~sP7('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f344,f604,f548])).
fof(f748,plain,(
  ~sP7('F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f604,f548])).
fof(f869,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') | $spl27),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f751,plain,(
  ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f556])).
fof(f867,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') | $spl25),
  inference(cnf_transformation,[],[f867_D])).
fof(f867_D,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f770,plain,(
  ~sP9('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f563])).
fof(f865,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl23),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f774,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f594,f604,f573])).
fof(f775,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f345,f604,f573])).
fof(f776,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f344,f604,f573])).
fof(f861,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl19),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f777,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f604,f573])).
fof(f863,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl21),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f778,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f594,f350,f604,f573])).
fof(f779,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f350,f604,f573])).
fof(f780,plain,(
  ~sP10('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f350,f604,f573])).
fof(f781,plain,(
  ~sP10('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f350,f604,f573])).
fof(f859,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl16),
  inference(cnf_transformation,[],[f859_D])).
fof(f859_D,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f783,plain,(
  sP11('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f577])).
fof(f857,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl15),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f793,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f594,f604,f580])).
fof(f794,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f345,f604,f580])).
fof(f795,plain,(
  ~sP11('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f344,f604,f580])).
fof(f853,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl11),
  inference(cnf_transformation,[],[f853_D])).
fof(f853_D,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f796,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f604,f580])).
fof(f855,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl13),
  inference(cnf_transformation,[],[f855_D])).
fof(f855_D,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f797,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f594,f604,f580])).
fof(f798,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f345,f604,f580])).
fof(f799,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f344,f604,f580])).
fof(f800,plain,(
  ~sP11('const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)),'F_0','const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f604,f604,f580])).
fof(f802,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f594,f604,f589])).
fof(f803,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f345,f604,f589])).
fof(f804,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f344,f604,f589])).
fof(f805,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f604,f589])).
fof(f851,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f806,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f594,f604,f589])).
fof(f807,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f345,f604,f589])).
fof(f808,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f344,f604,f589])).
fof(f809,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f604,f589])).
fof(f810,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))) | $spl4),
  inference(unit_resulting_resolution,[],[f350,f350,f604,f598])).
fof(f604,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) | $spl4),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f341,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12))))),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK12)))) & s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f160])).
fof(f160,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & ~s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X0))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X2),s('type.num.num',X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) & ~s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X6))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))) & ~s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X6))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',X8),s('type.num.num',X7))))))),
  file('thm.integer.INT_EXP_MOD','thm.integer.INT_EXP_MOD')).
fof(f602,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f342,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f209])).
fof(f600,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))) | $spl1),
  inference(cnf_transformation,[],[f600_D])).
fof(f600_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f343,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK12))),s('type.integer.int','const.integer.int_exp_2'(s('type.integer.int',sK14),s('type.num.num',sK13)))))),
  inference(cnf_transformation,[],[f209])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 330
% Passive clauses: 452
% Generated clauses: 3345
% Final active clauses: 322
% Final passive clauses: 183
% Input formulas: 64
% Initial clauses: 250
% 
% Duplicate literals: 54
% Trivial inequalities: 24
% Fw demodulations: 488
% Bw demodulations: 3
% Global subsumptions: 7
% 
% Simple tautologies: 206
% Equational tautologies: 17
% Forward subsumptions: 1555
% Backward subsumptions: 1
% 
% Binary resolution: 653
% Unit resulting resolution: 1372
% Forward superposition: 220
% Backward superposition: 321
% Equality resolution: 5
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 1739
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1374
% 
% Sat splits: 17
% Sat splitting refutations: 2
% 
% Memory used [KB]: 4349
% Time elapsed: 0.093 s
% ------------------------------
% ------------------------------
% Success in time 18.228 s
