Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numeral
% (81644)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 170
% Passive clauses: 1484
% Generated clauses: 23072
% Final active clauses: 163
% Final passive clauses: 1136
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 1004
% Trivial inequalities: 189
% Fw subsumption resolutions: 2261
% Fw demodulations: 2102
% Bw demodulations: 6
% 
% Simple tautologies: 368
% Equational tautologies: 60
% Forward subsumptions: 13215
% Fw demodulations to eq. taut.: 70
% 
% Binary resolution: 179
% Forward superposition: 8110
% Backward superposition: 8509
% Self superposition: 312
% Equality factoring: 173
% Equality resolution: 20
% 
% Split clauses: 1047
% Split components: 2160
% SAT solver clauses: 1558
% SAT solver unit clauses: 22
% SAT solver binary clauses: 131
% 
% Sat splits: 1047
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1546
% TWLsolver calls for satisfiability: 143
% 
% Memory used [KB]: 19445
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numeral
% (81657)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 458
% Generated clauses: 16692
% Final active clauses: 137
% Final passive clauses: 275
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 1086
% Trivial inequalities: 75
% Fw subsumption resolutions: 449
% Fw demodulations: 1323
% Bw demodulations: 36
% Condensations: 1370
% Global subsumptions: 4
% 
% Simple tautologies: 193
% Equational tautologies: 39
% Forward subsumptions: 9798
% Fw demodulations to eq. taut.: 87
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 35
% Forward superposition: 5911
% Backward superposition: 5571
% Self superposition: 226
% Equality factoring: 110
% Equality resolution: 12
% 
% Split clauses: 24
% Split components: 53
% SAT solver clauses: 724
% SAT solver unit clauses: 62
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 8
% 
% Sat splits: 24
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 708
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 20340
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numeral
% (81689)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 399
% Passive clauses: 3680
% Generated clauses: 108642
% Final active clauses: 160
% Final passive clauses: 809
% Input formulas: 30
% Initial clauses: 46
% Discarded non-redundant clauses: 1712
% 
% Duplicate literals: 4048
% Trivial inequalities: 566
% Fw subsumption resolutions: 4811
% Fw demodulations: 15392
% Bw demodulations: 940
% Condensations: 7000
% Global subsumptions: 2986
% 
% Simple tautologies: 820
% Equational tautologies: 1720
% Forward subsumptions: 63881
% Fw demodulations to eq. taut.: 515
% Bw demodulations to eq. taut.: 32
% 
% Unit resulting resolution: 4677
% Forward superposition: 31630
% Backward superposition: 21007
% Self superposition: 692
% Equality factoring: 252
% Equality resolution: 44
% 
% Split clauses: 208
% Split components: 421
% SAT solver clauses: 10090
% SAT solver unit clauses: 2227
% SAT solver binary clauses: 181
% 
% Sat splits: 208
% Sat splitting refutations: 3006
% 
% Memory used [KB]: 212107
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.numeral
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 28
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 28
% Input formulas: 30
% Initial clauses: 42
% 
% Selected by SInE selection: 26
% SInE iterations: 1
% 
% Duplicate literals: 1
% Trivial inequalities: 1
% 
% Simple tautologies: 13
% Equational tautologies: 1
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.numeral
% (82039)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 149
% Passive clauses: 1236
% Generated clauses: 20300
% Final active clauses: 123
% Final passive clauses: 624
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 375
% Trivial inequalities: 335
% Fw subsumption resolutions: 1859
% Fw demodulations: 1517
% Bw demodulations: 630
% Global subsumptions: 5
% 
% Simple tautologies: 398
% Equational tautologies: 61
% Forward subsumptions: 14390
% Fw demodulations to eq. taut.: 107
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 108
% Forward superposition: 7414
% Backward superposition: 8013
% Self superposition: 207
% Equality factoring: 90
% Equality resolution: 12
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 1274
% SAT solver unit clauses: 449
% SAT solver binary clauses: 30
% 
% Sat splits: 14
% Sat splitting refutations: 16
% 
% Memory used [KB]: 19957
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.numeral
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 32
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 32
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 1
% Trivial inequalities: 1
% 
% Simple tautologies: 13
% Equational tautologies: 1
% 
% SAT solver clauses: 3
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.numeral
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 32
% Passive clauses: 1
% Generated clauses: 2
% Final active clauses: 32
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 1
% Trivial inequalities: 1
% 
% Simple tautologies: 13
% Equational tautologies: 1
% 
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.026 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.numeral
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 32
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 32
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 1
% Trivial inequalities: 1
% 
% Simple tautologies: 13
% Equational tautologies: 1
% 
% SAT solver clauses: 1
% SAT solver unit clauses: 1
% 
% TWLsolver clauses: 1
% TWLsolver calls for satisfiability: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 115
lrs+11_28_cond=on:gs=on:gsaa=from_current:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1.7:stl=30:sac=on:sdd=off:sfr=on:ssfp=40000:ssfq=1.1:smm=off:ssnc=none:urr=on:updr=off_109 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f27790,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f27655])).
fof(f27655,plain,(
  s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))))) != s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1)))))),
  inference(superposition,[],[f83,f13487])).
fof(f13487,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X6))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X7)))))))) )),
  inference(forward_demodulation,[],[f13486,f786])).
fof(f786,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f785,f757])).
fof(f757,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f124,f139])).
fof(f139,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f137,f132])).
fof(f132,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f90,f86])).
fof(f86,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.|48|_0')),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.ALT_ZERO')).
fof(f90,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.ADD_|48|')).
fof(f137,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(backward_demodulation,[],[f124,f135])).
fof(f135,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f93,f86])).
fof(f93,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.BIT|49|')).
fof(f124,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f785,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f758,f124])).
fof(f758,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f124,f141])).
fof(f141,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f140,f132])).
fof(f140,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(forward_demodulation,[],[f138,f124])).
fof(f138,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(backward_demodulation,[],[f124,f136])).
fof(f136,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(forward_demodulation,[],[f94,f86])).
fof(f94,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.BIT2')).
fof(f13486,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X6))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X7)))))))) )),
  inference(forward_demodulation,[],[f13317,f757])).
fof(f13317,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X6))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X7)))))))) )),
  inference(superposition,[],[f2741,f12984])).
fof(f12984,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f128,f91])).
fof(f91,plain,(
  ( ! [X0] : (s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X8] : s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X8)))),
  file('thm.numeral.numeral_exp|split|2','thm.numeral.iSQR')).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9,X10,X6] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X9))),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X10)))))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.EXP_ADD')).
fof(f2741,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f2696,f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7,X6] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.EXP|split|1')).
fof(f2696,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(superposition,[],[f2512,f125])).
fof(f2512,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f127,f91])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X6,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X9)))),
  file('thm.numeral.numeral_exp|split|2','thm.arithmetic.MULT_ASSOC')).
fof(f83,plain,(
  s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',sK0))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))))),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',sK0))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK0),s('type.num.num',sK1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f62])).
fof(f62,plain,(
  ? [X0,X1] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X0))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X0))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X6))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iSQR_1'(s('type.num.num',X6))),s('type.num.num','const.numeral.iSQR_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num',X7)))))))),
  file('thm.numeral.numeral_exp|split|2','thm.numeral.numeral_exp|split|2')).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 182
% Passive clauses: 772
% Generated clauses: 27327
% Final active clauses: 176
% Final passive clauses: 522
% Input formulas: 30
% Initial clauses: 46
% 
% Duplicate literals: 52
% Trivial inequalities: 13
% Fw subsumption resolutions: 19
% Fw demodulations: 2992
% Bw demodulations: 74
% Condensations: 62
% 
% Simple tautologies: 47
% Equational tautologies: 25
% Forward subsumptions: 22948
% Fw demodulations to eq. taut.: 262
% 
% Binary resolution: 11
% Unit resulting resolution: 19
% Forward superposition: 11741
% Backward superposition: 12463
% Self superposition: 165
% Equality factoring: 26
% Equality resolution: 5
% 
% SAT solver clauses: 772
% SAT solver unit clauses: 744
% SAT solver binary clauses: 6
% 
% TWLsolver clauses: 772
% TWLsolver calls for satisfiability: 1312
% 
% Memory used [KB]: 33005
% Time elapsed: 0.654 s
% ------------------------------
% ------------------------------
% Success in time 13.119 s
