Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f370,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f365,f70,f109])).
fof(f109,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.transc_2|49|6_prove_ltb|split|1',aHL_FALSITY)).
fof(f365,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f359,f109])).
fof(f359,plain,(
  p(s(bool,'F_0'))),
  inference(superposition,[],[f355,f140])).
fof(f140,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK0))),s('type.realax.real',sK0)))),
  inference(resolution,[],[f139,f79])).
fof(f79,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.transc.transc_2|49|6_prove_ltb|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f139,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',sK0))),s('type.realax.real',sK0))))),
  inference(forward_demodulation,[],[f67,f71])).
fof(f71,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X13))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X13)))),
  file('thm.transc.transc_2|49|6_prove_ltb|split|1','ARITY_0_1_const.transc.exp')).
fof(f67,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',sK0))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',sK0))),s('type.realax.real',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f46])).
fof(f46,plain,(
  ? [X0] : ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f44])).
fof(f44,plain,(
  ? [X1] : ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X1))),s('type.realax.real',X1))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X1] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X1))),s('type.realax.real',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X7))),s('type.realax.real',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X7))),s('type.realax.real',X7))))),
  file('thm.transc.transc_2|49|6_prove_ltb|split|1','thm.transc.transc_2|49|6_prove_ltb|split|1')).
fof(f355,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) )),
  inference(forward_demodulation,[],[f72,f71])).
fof(f72,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X7))),s('type.realax.real',X7))))),
  file('thm.transc.transc_2|49|6_prove_ltb|split|1','thm.transc.DIFF_EXP')).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 16
% Passive clauses: 51
% Generated clauses: 285
% Final active clauses: 15
% Final passive clauses: 24
% Input formulas: 20
% Initial clauses: 36
% 
% Duplicate literals: 38
% Trivial inequalities: 10
% Fw subsumption resolutions: 21
% Fw demodulations: 3
% Bw demodulations: 1
% 
% Simple tautologies: 38
% Equational tautologies: 17
% Forward subsumptions: 106
% 
% Binary resolution: 15
% Forward superposition: 33
% Backward superposition: 76
% Self superposition: 23
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 13
% SAT solver unit clauses: 5
% SAT solver binary clauses: 7
% SAT solver learnt clauses: 1
% 
% Sat splits: 6
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 9
% TWLsolver calls for satisfiability: 4
% 
% Memory used [KB]: 639
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
% Success in time 0.017 s
