Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f11546,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f11545,f1695,f2592,f2590,f11441,f1696,f2593,f2591])).
fof(f2591,plain,(
  s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(cnf_transformation,[],[f2591_D])).
fof(f2591_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2593,plain,(
  s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f2593_D])).
fof(f2593_D,plain,(
  s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1696,plain,(
  s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f1401])).
fof(f1401,plain,(
  (s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f1400])).
fof(f1400,plain,(
  ? [X0] : ((s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(nnf_transformation,[],[f1138])).
fof(f1138,plain,(
  ? [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))) <~> s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))) <=> s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X6))) <=> s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X6))) <=> s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.neg_eq|48|','thm.extreal.neg_eq|48|')).
fof(f11441,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f11440,f4260])).
fof(f4260,plain,(
  s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2591,f2592])).
fof(f11440,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) | $spl2),
  inference(forward_demodulation,[],[f1722,f2592])).
fof(f1722,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.neg_eq|48|','thm.extreal.neg_|48|')).
fof(f2590,plain,(
  s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(cnf_transformation,[],[f2590_D])).
fof(f2590_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2592,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f2592_D])).
fof(f2592_D,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1695,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f1401])).
fof(f11545,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f11544,f11444])).
fof(f11444,plain,(
  s('type.extreal.extreal',sK19) != s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) | ($spl0 | $spl3)),
  inference(superposition,[],[f2593,f2590])).
fof(f11544,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) | $spl0),
  inference(forward_demodulation,[],[f11543,f1752])).
fof(f1752,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0)))))) )),
  inference(cnf_transformation,[],[f537])).
fof(f537,plain,(
  ! [X0] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X6)))))),
  file('thm.extreal.neg_eq|48|','thm.extreal.neg_neg')).
fof(f11543,plain,(
  s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))) = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK19))))) | $spl0),
  inference(forward_demodulation,[],[f1722,f2590])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 630
% Passive clauses: 1823
% Generated clauses: 9830
% Final active clauses: 605
% Final passive clauses: 1085
% Input formulas: 517
% Initial clauses: 885
% 
% Duplicate literals: 66
% Trivial inequalities: 25
% Fw subsumption resolutions: 438
% Fw demodulations: 386
% Bw demodulations: 21
% 
% Simple tautologies: 197
% Equational tautologies: 11
% Forward subsumptions: 6767
% 
% Binary resolution: 1823
% Factoring: 169
% Forward superposition: 2473
% Backward superposition: 3447
% Self superposition: 19
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 157
% Split components: 316
% SAT solver clauses: 227
% SAT solver unit clauses: 28
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 157
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 225
% TWLsolver calls for satisfiability: 163
% 
% Memory used [KB]: 5628
% Time elapsed: 0.132 s
% ------------------------------
% ------------------------------
% Success in time 0.15 s
