Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (61085)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 55
% Passive clauses: 1820
% Generated clauses: 10228
% Final active clauses: 53
% Final passive clauses: 1715
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 638
% Trivial inequalities: 90
% Fw subsumption resolutions: 946
% Fw demodulations: 144
% 
% Simple tautologies: 184
% Equational tautologies: 110
% Forward subsumptions: 4958
% 
% Binary resolution: 25
% Forward superposition: 4028
% Backward superposition: 3558
% Self superposition: 404
% Equality factoring: 328
% Equality resolution: 8
% 
% Split clauses: 391
% Split components: 785
% SAT solver clauses: 1235
% SAT solver unit clauses: 280
% SAT solver binary clauses: 412
% 
% Sat splits: 391
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1233
% TWLsolver calls for satisfiability: 34
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (61103)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 37
% Passive clauses: 229
% Generated clauses: 5860
% Final active clauses: 34
% Final passive clauses: 190
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 841
% Trivial inequalities: 55
% Fw subsumption resolutions: 141
% Fw demodulations: 36
% Bw demodulations: 1
% Condensations: 1688
% 
% Simple tautologies: 116
% Equational tautologies: 77
% Forward subsumptions: 2108
% 
% Binary resolution: 16
% Forward superposition: 1713
% Backward superposition: 1511
% Self superposition: 268
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 601
% SAT solver unit clauses: 46
% SAT solver binary clauses: 66
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 601
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 5884
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f75339,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f75317])).
fof(f75317,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))),
  inference(superposition,[],[f77,f15403])).
fof(f15403,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s('type.canonical.canonical_sum'(X11),'const.canonical.Cons_monom_3'(s(X11,X13),s('type.list.list'('type.quote.index'),X14),s('type.canonical.canonical_sum'(X11),X15))) = s('type.canonical.canonical_sum'(X11),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X11),X12),s('type.canonical.canonical_sum'(X11),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X11),'const.canonical.Cons_monom_3'(s(X11,X13),s('type.list.list'('type.quote.index'),X14),s('type.canonical.canonical_sum'(X11),X15)))))) )),
  inference(superposition,[],[f10553,f97])).
fof(f97,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X4),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X3),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X8,X9,X10,X15] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|5','thm.canonical.canonical_sum_merge_def|split|5')).
fof(f10553,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f10504,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|5','ARITY_1_3_const.ringNorm.r_canonical_sum_merge')).
fof(f10504,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f92,f81])).
fof(f81,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|5','thm.ringNorm.r_canonical_sum_merge_def')).
fof(f92,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|5','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f77,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_monom_3'(s(sK0,sK4),s('type.list.list'('type.quote.index'),sK3),s('type.canonical.canonical_sum'(sK0),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_monom_3'(s(X0,X4),s('type.list.list'('type.quote.index'),X3),s('type.canonical.canonical_sum'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_monom_3'(s(X6,X10),s('type.list.list'('type.quote.index'),X9),s('type.canonical.canonical_sum'(X6),X8)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|5','thm.ringNorm.canonical_sum_merge_def|split|5')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 212
% Passive clauses: 5300
% Generated clauses: 79462
% Final active clauses: 119
% Final passive clauses: 629
% Input formulas: 25
% Initial clauses: 26
% Discarded non-redundant clauses: 4551
% 
% Duplicate literals: 5338
% Trivial inequalities: 437
% Fw subsumption resolutions: 3502
% Fw demodulations: 6810
% Bw demodulations: 2831
% Condensations: 10708
% Global subsumptions: 33
% 
% Simple tautologies: 736
% Equational tautologies: 1699
% Forward subsumptions: 44495
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 312
% 
% Unit resulting resolution: 4374
% Forward superposition: 22522
% Backward superposition: 16187
% Self superposition: 522
% Equality factoring: 298
% Equality resolution: 17
% 
% Split clauses: 220
% Split components: 458
% SAT solver clauses: 5811
% SAT solver unit clauses: 79
% SAT solver binary clauses: 427
% 
% Sat splits: 220
% Sat splitting refutations: 36
% 
% Memory used [KB]: 52834
% Time elapsed: 8.246 s
% ------------------------------
% ------------------------------
% Success in time 9.57 s
