Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.quotient_list
% (109296)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 2233
% Generated clauses: 32526
% Final active clauses: 134
% Final passive clauses: 1629
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 1078
% Trivial inequalities: 516
% Fw subsumption resolutions: 4486
% Fw demodulations: 5648
% Bw demodulations: 28
% 
% Simple tautologies: 494
% Equational tautologies: 63
% Forward subsumptions: 14723
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 277
% Factoring: 3
% Forward superposition: 10374
% Backward superposition: 9532
% Self superposition: 284
% Equality factoring: 167
% Equality resolution: 41
% 
% Split clauses: 177
% Split components: 367
% SAT solver clauses: 478
% SAT solver unit clauses: 30
% SAT solver binary clauses: 104
% 
% Sat splits: 177
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 470
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 15991
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient_list
% (109320)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 587
% Generated clauses: 17553
% Final active clauses: 114
% Final passive clauses: 407
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 716
% Trivial inequalities: 269
% Fw subsumption resolutions: 1979
% Fw demodulations: 2429
% Bw demodulations: 66
% Condensations: 1094
% Global subsumptions: 5
% 
% Simple tautologies: 277
% Equational tautologies: 42
% Forward subsumptions: 9151
% 
% Binary resolution: 256
% Factoring: 1
% Forward superposition: 5401
% Backward superposition: 4849
% Self superposition: 217
% Equality factoring: 93
% Equality resolution: 28
% 
% Split clauses: 119
% Split components: 246
% SAT solver clauses: 1067
% SAT solver unit clauses: 51
% SAT solver binary clauses: 37
% SAT solver learnt clauses: 10
% 
% Sat splits: 119
% 
% TWLsolver clauses: 1047
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 11513
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quotient_list
% (109353)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 439
% Passive clauses: 4826
% Generated clauses: 137372
% Final active clauses: 91
% Final passive clauses: 585
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 7839
% Trivial inequalities: 1076
% Fw subsumption resolutions: 12530
% Fw demodulations: 13313
% Bw demodulations: 3149
% Condensations: 18012
% Global subsumptions: 419
% 
% Simple tautologies: 1104
% Equational tautologies: 2023
% Forward subsumptions: 75186
% Fw demodulations to eq. taut.: 233
% Bw demodulations to eq. taut.: 233
% 
% Unit resulting resolution: 6740
% Factoring: 1
% Forward superposition: 32775
% Backward superposition: 31280
% Self superposition: 958
% Equality factoring: 474
% Equality resolution: 85
% 
% Split clauses: 442
% Split components: 894
% SAT solver clauses: 6457
% SAT solver unit clauses: 185
% SAT solver binary clauses: 478
% 
% Sat splits: 442
% Sat splitting refutations: 326
% 
% Memory used [KB]: 61278
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quotient_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient_list
% SZS output start Proof for thm.quotient_list
fof(f7351,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f135,f124,f125,f126,f127,f128,f129,f130,f131,f132,f133,f134,f156,f171,f172,f176,f163,f164,f181,f239,f255,f299,f285,f301,f302,f303,f386,f387,f246,f245,f178,f174,f258,f256,f173,f389,f445,f139,f434,f143,f435,f436,f437,f438,f439,f440,f444,f448,f170,f137,f550,f552,f553,f562,f563,f566,f569,f636,f637,f718,f719,f443,f442,f441,f768,f755,f756,f757,f758,f759,f760,f761,f762,f763,f764,f849,f871,f872,f907,f908,f909,f910,f950,f951,f952,f1000,f987,f988,f989,f990,f991,f992,f993,f994,f995,f996,f1101,f1102,f388,f1137,f1149,f1155,f1156,f1153,f1151,f1186,f1204,f141,f1191,f1192,f146,f1193,f1194,f147,f1195,f1196,f1197,f1198,f144,f1199,f151,f153,f1200,f149,f155,f1315,f1339,f1340,f175,f145,f1372,f1148,f1385,f1525,f1537,f1542,f1539,f1557,f1535,f1559,f1538,f1562,f1566,f1596,f1757,f1769,f1774,f1773,f1727,f1534,f1789,f1767,f1766,f1792,f1796,f1799,f1983,f2007,f2005,f2293,f2292,f2314,f2316,f2324,f2325,f2334,f2336,f2337,f2312,f2346,f2298,f2394,f2406,f2407,f2412,f2411,f2409,f1911,f1770,f2416,f2446,f2405,f2404,f2450,f2408,f2455,f2458,f2494,f2512,f2510,f2520,f2539,f2537,f2608,f2612,f2613,f2616,f2610,f2621,f2625,f2572,f2623,f2626,f2556,f2535,f1536,f2655,f2658,f2661,f2664,f2670,f2672,f2675,f2678,f2680,f2682,f2685,f2641,f2642,f2690,f2545,f2645,f2519,f2647,f1146,f2464,f1771,f2462,f2692,f2461,f1768,f2694,f2698,f2709,f2710,f2714,f2712,f2716,f2744,f2708,f2747,f2711,f2751,f2851,f2870,f2876,f2943,f2970,f2971,f2609,f2973,f3004,f3007,f3006,f3009,f3012,f3014,f3077,f3080,f3079,f3082,f3084,f3085,f3037,f3088,f3160,f3187,f3161,f3188,f3171,f3087,f3192,f3504,f3503,f3527,f3528,f3539,f3540,f3549,f3550,f3551,f3526,f3560,f3508,f3621,f3626,f3623,f3765,f3625,f3767,f3862,f3865,f2618,f3791,f3864,f3866,f3868,f3870,f3871,f3876,f3877,f3879,f3881,f3883,f3885,f3873,f3929,f2615,f3863,f3931,f1541,f3932,f3875,f3933,f3016,f3934,f3878,f3935,f3880,f3936,f3882,f3937,f3884,f3938,f3886,f3872,f3958,f3987,f3985,f4055,f4085,f4083,f4088,f3011,f4091,f2622,f4092,f4090,f3869,f4101,f4112,f2707,f4102,f4116,f4114,f4119,f4123,f4121,f4157,f4179,f4177,f4180,f4185,f4187,f4188,f4190,f4191,f4193,f4194,f4184,f4182,f4334,f4356,f4354,f4357,f4362,f4364,f4365,f4367,f4370,f4361,f4359,f4591,f4603,f4601,f4606,f4610,f4608,f4882,f4881,f2313,f4891,f4892,f4899,f4902,f4900,f4915,f4916,f4923,f2335,f4925,f4926,f2315,f4890,f4935,f4889,f5052,f5053,f5054,f5055,f5056,f5057,f5058,f5074,f5076,f5078,f5079,f5080,f5091,f85,f5106,f5105,f5526,f2689,f6870,f6871,f2297,f7280,f7303,f5545,f7344])).
fof(f7344,plain,(
  $false | ($spl0 | $spl33 | $spl130)),
  inference(subsumption_resolution,[],[f7343,f92])).
fof(f92,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient_list.NULL_PRS','thm.bool.TRUTH')).
fof(f7343,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl33 | $spl130)),
  inference(forward_demodulation,[],[f7326,f316])).
fof(f316,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl33),
  inference(trivial_inequality_removal,[],[f315])).
fof(f315,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl33),
  inference(superposition,[],[f301,f79])).
fof(f79,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient_list.NULL_PRS',aHL_BOOL_CASES)).
fof(f7326,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl0 | $spl130)),
  inference(superposition,[],[f7273,f2712])).
fof(f7273,plain,(
  ( ! [X33,X34] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X33),'const.list.MAP_2'(s(fun(sK1,X33),X34),s('type.list.list'(sK1),sK5))))))) ) | $spl0),
  inference(superposition,[],[f5510,f137])).
fof(f5510,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X2),'const.list.MAP_2'(s(fun(X3,X2),X4),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X5),s('type.list.list'(X3),X6))))))))) )),
  inference(superposition,[],[f81,f80])).
fof(f80,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X3),s('type.list.list'(X1),X4))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))),s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X4)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X6] : ! [X15,X16,X14] : s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X15),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X16),s('type.list.list'(X6),X14))))) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,happ(s(fun(X6,X7),X15),s(X6,X16))),s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X15),s('type.list.list'(X6),X14)))))),
  file('thm.quotient_list.NULL_PRS','thm.list.MAP|split|1')).
fof(f81,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X16,X14] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X16),s('type.list.list'(X6),X14))))))),
  file('thm.quotient_list.NULL_PRS','thm.list.NULL|split|1')).
fof(f5545,plain,(
  ( ! [X14] : (s('type.list.list'(sK1),X14) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X14,sK1)),s('type.list.list'(sK1),sK7(X14,sK1))))) ) | $spl256),
  inference(cnf_transformation,[],[f5545_D])).
fof(f5545_D,plain,(
  ( ! [X14] : (s('type.list.list'(sK1),X14) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X14,sK1)),s('type.list.list'(sK1),sK7(X14,sK1))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f7303,plain,(
  ( ! [X50,X49] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X49),'const.list.MAP_2'(s(fun(sK1,X49),X50),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))))) ) | $spl260),
  inference(cnf_transformation,[],[f7303_D])).
fof(f7303_D,plain,(
  ( ! [X50,X49] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X49),'const.list.MAP_2'(s(fun(sK1,X49),X50),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f7280,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X49),'const.list.MAP_2'(s(fun(sK1,X49),X50),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))))) | s('type.list.list'(sK1),X48) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X48,sK1)),s('type.list.list'(sK1),sK7(X48,sK1))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f5510,f2895])).
fof(f2895,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f2686,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK6(X1,X0)),s('type.list.list'(X0),sK7(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK6(X1,X0)),s('type.list.list'(X0),sK7(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f28])).
fof(f28,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X11] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X11) | ? [X16,X14] : s('type.list.list'(X6),X11) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X16),s('type.list.list'(X6),X14))))),
  file('thm.quotient_list.NULL_PRS','thm.list.list_CASES')).
fof(f2686,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl44 | $spl84)),
  inference(forward_demodulation,[],[f1534,f2409])).
fof(f2297,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) | $spl68),
  inference(cnf_transformation,[],[f2297_D])).
fof(f2297_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,X24)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f6871,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,X41)) ) | $spl258),
  inference(cnf_transformation,[],[f6871_D])).
fof(f6871_D,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,X41)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f6870,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X40)) ) | $spl14),
  inference(subsumption_resolution,[],[f6855,f76])).
fof(f76,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.EQ_CLAUSES|split|0')).
fof(f6855,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,X41) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X40)) ) | $spl14),
  inference(superposition,[],[f5685,f6002])).
fof(f6002,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,X34)) ) | $spl14),
  inference(subsumption_resolution,[],[f5990,f86])).
fof(f86,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient_list.NULL_PRS',aHL_FALSITY)).
fof(f5990,plain,(
  ( ! [X34] : (p(s(bool,'F_0')) | s(bool,X34) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,X34)) ) | $spl14),
  inference(superposition,[],[f5211,f148])).
fof(f148,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl14),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5211,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f72,f79])).
fof(f72,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.EQ_CLAUSES|split|2')).
fof(f5685,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(resolution,[],[f5202,f72])).
fof(f5202,plain,(
  ( ! [X2,X3] : (~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f5164,f86])).
fof(f5164,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X3)) | p(s(bool,X2)) | p(s(bool,'F_0'))) )),
  inference(superposition,[],[f71,f66])).
fof(f66,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.IMP_ANTISYM_AX')).
fof(f71,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f2689,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl119),
  inference(cnf_transformation,[],[f2689_D])).
fof(f2689_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f5526,plain,(
  ( ! [X14] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | s('type.list.list'(sK1),X14) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X14,sK1)),s('type.list.list'(sK1),sK7(X14,sK1))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f81,f2895])).
fof(f5105,plain,(
  ~p(s(bool,'F_0')) | $spl253),
  inference(cnf_transformation,[],[f5105_D])).
fof(f5105_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f5106,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl254),
  inference(cnf_transformation,[],[f5106_D])).
fof(f5106_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f85,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14] : (p(s(bool,'F_0')) => p(s(bool,X14)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.FALSITY')).
fof(f5091,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f4968,f92])).
fof(f4968,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(backward_demodulation,[],[f2297,f81])).
fof(f5080,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f5069,f92])).
fof(f5069,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(duplicate_literal_removal,[],[f4945])).
fof(f4945,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl68),
  inference(backward_demodulation,[],[f2297,f90])).
fof(f90,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient_list.NULL_PRS','thm.bool.NOT_CLAUSES|split|2')).
fof(f5079,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f5071,f92])).
fof(f5071,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(duplicate_literal_removal,[],[f4943])).
fof(f4943,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl68),
  inference(backward_demodulation,[],[f2297,f88])).
fof(f88,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f11])).
fof(f11,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quotient_list.NULL_PRS','thm.bool.NOT_CLAUSES|split|1')).
fof(f5078,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f4941,f92])).
fof(f4941,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(backward_demodulation,[],[f2297,f86])).
fof(f5076,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f5075,f2297])).
fof(f5075,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl68),
  inference(subsumption_resolution,[],[f4938,f77])).
fof(f77,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.EQ_CLAUSES|split|1')).
fof(f4938,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl68),
  inference(backward_demodulation,[],[f2297,f73])).
fof(f73,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quotient_list.NULL_PRS','thm.bool.EQ_CLAUSES|split|3')).
fof(f5074,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f5073,f2297])).
fof(f5073,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl68),
  inference(subsumption_resolution,[],[f4936,f77])).
fof(f4936,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl68),
  inference(backward_demodulation,[],[f2297,f71])).
fof(f5058,plain,(
  $false | ($spl29 | $spl68)),
  inference(trivial_inequality_removal,[],[f5040])).
fof(f5040,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl29 | $spl68)),
  inference(superposition,[],[f178,f2297])).
fof(f5057,plain,(
  $false | ($spl68 | $spl130 | $spl162)),
  inference(trivial_inequality_removal,[],[f5041])).
fof(f5041,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl68 | $spl130 | $spl162)),
  inference(superposition,[],[f3631,f2297])).
fof(f3631,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) | ($spl130 | $spl162)),
  inference(resolution,[],[f3615,f75])).
fof(f75,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f3615,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))))) | ($spl130 | $spl162)),
  inference(superposition,[],[f81,f3587])).
fof(f3587,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl130 | $spl162)),
  inference(superposition,[],[f2857,f3526])).
fof(f2857,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X0))),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,X0),s('type.list.list'(sK1),sK5)))))) ) | $spl130),
  inference(superposition,[],[f80,f2712])).
fof(f5056,plain,(
  $false | ($spl17 | $spl68)),
  inference(trivial_inequality_removal,[],[f5042])).
fof(f5042,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl17 | $spl68)),
  inference(superposition,[],[f151,f2297])).
fof(f5055,plain,(
  $false | ($spl68 | $spl92)),
  inference(trivial_inequality_removal,[],[f5043])).
fof(f5043,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl68 | $spl92)),
  inference(superposition,[],[f4199,f2297])).
fof(f4199,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | $spl92),
  inference(resolution,[],[f4170,f75])).
fof(f4170,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))))) | $spl92),
  inference(superposition,[],[f81,f2510])).
fof(f5054,plain,(
  $false | ($spl68 | $spl96)),
  inference(trivial_inequality_removal,[],[f5044])).
fof(f5044,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl68 | $spl96)),
  inference(superposition,[],[f4375,f2297])).
fof(f4375,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl96),
  inference(resolution,[],[f4347,f75])).
fof(f4347,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))))) | $spl96),
  inference(superposition,[],[f81,f2537])).
fof(f5053,plain,(
  $false | ($spl36 | $spl68)),
  inference(trivial_inequality_removal,[],[f5045])).
fof(f5045,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl36 | $spl68)),
  inference(superposition,[],[f1571,f2297])).
fof(f1571,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl36),
  inference(resolution,[],[f1527,f75])).
fof(f1527,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK5,sK1))))) | $spl36),
  inference(superposition,[],[f81,f1146])).
fof(f5052,plain,(
  $false | ($spl44 | $spl68)),
  inference(trivial_inequality_removal,[],[f5046])).
fof(f5046,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl44 | $spl68)),
  inference(superposition,[],[f1804,f2297])).
fof(f1804,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl44),
  inference(resolution,[],[f1759,f75])).
fof(f1759,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f81,f1534])).
fof(f4889,plain,(
  ( ! [X28,X29] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X28,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X29,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X28))) | s('type.list.list'(sK1),X28) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X28,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X29))) | s('type.list.list'(sK1),X29) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X29,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')))) ) | $spl244),
  inference(cnf_transformation,[],[f4889_D])).
fof(f4889_D,plain,(
  ( ! [X28,X29] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X28,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X29,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X28))) | s('type.list.list'(sK1),X28) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X28,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X29))) | s('type.list.list'(sK1),X29) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X29,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f4935,plain,(
  ( ! [X80,X81,X82] : (s('type.list.list'(sK1),X82) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X82,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X80,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X82,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X82))) | s('type.list.list'(sK1),X80) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X80,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X80))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4860,f76])).
fof(f4860,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X81)) | s('type.list.list'(sK1),X82) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X82,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X80,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X82,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X82))) | s('type.list.list'(sK1),X80) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X80,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X80))) | s(bool,'T_0') = s(bool,X81)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f4520,f4637])).
fof(f4637,plain,(
  ( ! [X0,X1] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,X1) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X0,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl2 | $spl33)),
  inference(resolution,[],[f4520,f78])).
fof(f78,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f58])).
fof(f4520,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,X1) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X0,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) | ($spl2 | $spl33)),
  inference(resolution,[],[f4302,f66])).
fof(f4302,plain,(
  ( ! [X17] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X17,sK1))))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X17))) | s('type.list.list'(sK1),X17) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X17,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f81,f1741])).
fof(f1741,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(X0,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(X0,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(X0,sK1),sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f639,f70])).
fof(f639,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f638,f70])).
fof(f638,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl2 | $spl33)),
  inference(forward_demodulation,[],[f138,f316])).
fof(f138,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl2),
  inference(cnf_transformation,[],[f138_D])).
fof(f138_D,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4890,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) | $spl246),
  inference(cnf_transformation,[],[f4890_D])).
fof(f4890_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f2315,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl76),
  inference(cnf_transformation,[],[f2315_D])).
fof(f2315_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4926,plain,(
  ( ! [X19,X20,X18] : (s(bool,X19) = s(bool,X20) | p(s(bool,X20)) | s('type.list.list'(sK1),X18) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X18,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X18))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4833,f4302])).
fof(f4833,plain,(
  ( ! [X19,X20,X18] : (s(bool,X19) = s(bool,X20) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X18,sK1))))) | p(s(bool,X20)) | s('type.list.list'(sK1),X18) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X18,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X18))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f66,f4637])).
fof(f4925,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) = s(bool,X17) | p(s(bool,X17)) | s('type.list.list'(sK1),X15) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X15,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X15))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4832,f4302])).
fof(f4832,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) = s(bool,X17) | p(s(bool,X17)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X15,sK1))))) | s('type.list.list'(sK1),X15) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X15,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X15))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f66,f4637])).
fof(f2335,plain,(
  ( ! [X97] : (s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f2335_D])).
fof(f2335_D,plain,(
  ( ! [X97] : (s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4923,plain,(
  ( ! [X105,X104] : (s('type.list.list'(sK1),X105) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X105,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X105))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X104),'const.list.NIL_0')))) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4801,f4302])).
fof(f4801,plain,(
  ( ! [X105,X104] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X105,sK1))))) | s('type.list.list'(sK1),X105) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X105,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X105))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X104),'const.list.NIL_0')))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f82,f4637])).
fof(f82,plain,(
  ( ! [X0] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0'))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.NIL_0'))))),
  file('thm.quotient_list.NULL_PRS','thm.list.NULL|split|0')).
fof(f4916,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4773,f4302])).
fof(f4773,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X10,sK1))))) | s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f66,f4637])).
fof(f4915,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),X7) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X7,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4772,f4302])).
fof(f4772,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X7,sK1))))) | s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),X7) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X7,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f66,f4637])).
fof(f4900,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl248),
  inference(cnf_transformation,[],[f4900_D])).
fof(f4900_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f4902,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) | $spl250),
  inference(cnf_transformation,[],[f4902_D])).
fof(f4902_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f4899,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) | s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl2 | $spl17 | $spl33 | $spl99)),
  inference(subsumption_resolution,[],[f4898,f151])).
fof(f4898,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) | s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl2 | $spl33 | $spl99)),
  inference(forward_demodulation,[],[f4897,f137])).
fof(f4897,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) | s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl0 | $spl2 | $spl33 | $spl99)),
  inference(subsumption_resolution,[],[f4745,f137])).
fof(f4745,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.NIL_0',sK1)),s('type.list.list'(sK1),sK7('const.list.NIL_0',sK1)))) | s(bool,X2) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl2 | $spl33 | $spl99)),
  inference(superposition,[],[f3898,f4637])).
fof(f3898,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | $spl99),
  inference(superposition,[],[f2538,f70])).
fof(f2538,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl99),
  inference(cnf_transformation,[],[f2538_D])).
fof(f2538_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f4892,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | s('type.list.list'(sK1),X13) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X13,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X14))) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4733,f4302])).
fof(f4733,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | s('type.list.list'(sK1),X13) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X13,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X15) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X13,sK1))))) | p(s(bool,X14))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f4637,f66])).
fof(f4891,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,X11))) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f4732,f4302])).
fof(f4732,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,X11)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X10,sK1)))))) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f4637,f66])).
fof(f2313,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl74),
  inference(cnf_transformation,[],[f2313_D])).
fof(f2313_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4881,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | ($spl2 | $spl33)),
  inference(duplicate_literal_removal,[],[f4726])).
fof(f4726,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f4637,f4637])).
fof(f4882,plain,(
  ( ! [X28,X29,X27] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X28,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X29,sK1)))) | s('type.list.list'(sK1),X29) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X29,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X29))) | s(bool,'T_0') = s(bool,X27) | s('type.list.list'(sK1),X28) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X28,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X28)))) ) | ($spl2 | $spl33)),
  inference(duplicate_literal_removal,[],[f4715])).
fof(f4715,plain,(
  ( ! [X28,X29,X27] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X28,sK1)))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(X29,sK1)))) | s('type.list.list'(sK1),X29) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X29,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X29))) | s(bool,'T_0') = s(bool,X27) | s('type.list.list'(sK1),X28) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X28,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X28))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl2 | $spl33)),
  inference(superposition,[],[f4637,f4637])).
fof(f4608,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | $spl241),
  inference(cnf_transformation,[],[f4608_D])).
fof(f4608_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f4610,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl242),
  inference(cnf_transformation,[],[f4610_D])).
fof(f4610_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f4606,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | $spl137),
  inference(forward_demodulation,[],[f4593,f2916])).
fof(f2916,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl137),
  inference(trivial_inequality_removal,[],[f2915])).
fof(f2915,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.NIL_0') | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl137),
  inference(superposition,[],[f2876,f70])).
fof(f4593,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0))),sK0)))) | $spl137),
  inference(superposition,[],[f2913,f2916])).
fof(f2913,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),X0) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),X0))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(X0,sK0)),s('type.list.list'(sK0),sK7(X0,sK0))))) ) | $spl137),
  inference(superposition,[],[f2876,f70])).
fof(f4601,plain,(
  s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | $spl236),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f4603,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl238),
  inference(cnf_transformation,[],[f4603_D])).
fof(f4603_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f4591,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0),sK0)))) | $spl137),
  inference(superposition,[],[f2916,f70])).
fof(f4359,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) | $spl228),
  inference(cnf_transformation,[],[f4359_D])).
fof(f4359_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f4361,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | $spl230),
  inference(cnf_transformation,[],[f4361_D])).
fof(f4361_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f4370,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) | $spl96),
  inference(forward_demodulation,[],[f4350,f2537])).
fof(f4350,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | $spl96),
  inference(superposition,[],[f70,f2537])).
fof(f4367,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl235),
  inference(cnf_transformation,[],[f4367_D])).
fof(f4367_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f4365,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | ($spl47 | $spl96)),
  inference(forward_demodulation,[],[f4338,f2537])).
fof(f4338,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | ($spl47 | $spl96)),
  inference(superposition,[],[f2864,f2537])).
fof(f2864,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | $spl47),
  inference(superposition,[],[f1535,f70])).
fof(f4364,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl233),
  inference(cnf_transformation,[],[f4364_D])).
fof(f4364_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f4362,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | ($spl83 | $spl84 | $spl96)),
  inference(forward_demodulation,[],[f4337,f2537])).
fof(f4337,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | ($spl83 | $spl84 | $spl96)),
  inference(superposition,[],[f2479,f2537])).
fof(f2479,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | ($spl83 | $spl84)),
  inference(superposition,[],[f2478,f70])).
fof(f2478,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl83 | $spl84)),
  inference(forward_demodulation,[],[f2405,f2409])).
fof(f4357,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) | $spl96),
  inference(forward_demodulation,[],[f4335,f2537])).
fof(f4335,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1))),sK1)))) | $spl96),
  inference(superposition,[],[f2537,f70])).
fof(f4354,plain,(
  s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)))) | $spl224),
  inference(cnf_transformation,[],[f4354_D])).
fof(f4354_D,plain,(
  s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f4356,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl226),
  inference(cnf_transformation,[],[f4356_D])).
fof(f4356_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f4334,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1),sK1)))) | $spl96),
  inference(superposition,[],[f2537,f70])).
fof(f4182,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) | $spl214),
  inference(cnf_transformation,[],[f4182_D])).
fof(f4182_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f4184,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | $spl216),
  inference(cnf_transformation,[],[f4184_D])).
fof(f4184_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f4194,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) | $spl92),
  inference(forward_demodulation,[],[f4173,f2510])).
fof(f4173,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | $spl92),
  inference(superposition,[],[f70,f2510])).
fof(f4193,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | $spl223),
  inference(cnf_transformation,[],[f4193_D])).
fof(f4193_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f4191,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | ($spl92 | $spl99)),
  inference(forward_demodulation,[],[f4169,f2510])).
fof(f4169,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | ($spl92 | $spl99)),
  inference(superposition,[],[f3898,f2510])).
fof(f4190,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | $spl221),
  inference(cnf_transformation,[],[f4190_D])).
fof(f4190_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f4188,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | ($spl47 | $spl92)),
  inference(forward_demodulation,[],[f4161,f2510])).
fof(f4161,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | ($spl47 | $spl92)),
  inference(superposition,[],[f2864,f2510])).
fof(f4187,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | $spl219),
  inference(cnf_transformation,[],[f4187_D])).
fof(f4187_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f4185,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | ($spl83 | $spl84 | $spl92)),
  inference(forward_demodulation,[],[f4160,f2510])).
fof(f4160,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | ($spl83 | $spl84 | $spl92)),
  inference(superposition,[],[f2479,f2510])).
fof(f4180,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) | $spl92),
  inference(forward_demodulation,[],[f4158,f2510])).
fof(f4158,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1))),sK1)))) | $spl92),
  inference(superposition,[],[f2510,f70])).
fof(f4177,plain,(
  s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)))) | $spl210),
  inference(cnf_transformation,[],[f4177_D])).
fof(f4177_D,plain,(
  s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f4179,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl212),
  inference(cnf_transformation,[],[f4179_D])).
fof(f4179_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f4157,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1),sK1)))) | $spl92),
  inference(superposition,[],[f2510,f70])).
fof(f4121,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))))) | $spl207),
  inference(cnf_transformation,[],[f4121_D])).
fof(f4121_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))))) <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f4123,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)))) | $spl208),
  inference(cnf_transformation,[],[f4123_D])).
fof(f4123_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f4119,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))))) | ($spl134 | $spl137)),
  inference(forward_demodulation,[],[f4104,f2751])).
fof(f4104,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0))),sK0)))) | ($spl134 | $spl137)),
  inference(superposition,[],[f2913,f2751])).
fof(f4114,plain,(
  s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)))) | $spl202),
  inference(cnf_transformation,[],[f4114_D])).
fof(f4114_D,plain,(
  s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f4116,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl204),
  inference(cnf_transformation,[],[f4116_D])).
fof(f4116_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f4102,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0),sK0)))) | $spl134),
  inference(superposition,[],[f2751,f70])).
fof(f2707,plain,(
  s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) | $spl126),
  inference(cnf_transformation,[],[f2707_D])).
fof(f2707_D,plain,(
  s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f4112,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl200),
  inference(cnf_transformation,[],[f4112_D])).
fof(f4112_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f4101,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) | $spl134),
  inference(superposition,[],[f2751,f70])).
fof(f3869,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl175),
  inference(cnf_transformation,[],[f3869_D])).
fof(f3869_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f4090,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | $spl198),
  inference(cnf_transformation,[],[f4090_D])).
fof(f4090_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f4092,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl50 | $spl99)),
  inference(forward_demodulation,[],[f4074,f1541])).
fof(f4074,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | ($spl50 | $spl99)),
  inference(superposition,[],[f3898,f1541])).
fof(f2622,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl109),
  inference(cnf_transformation,[],[f2622_D])).
fof(f2622_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f4091,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl47 | $spl50)),
  inference(forward_demodulation,[],[f4066,f1541])).
fof(f4066,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | ($spl47 | $spl50)),
  inference(superposition,[],[f2864,f1541])).
fof(f3011,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl141),
  inference(cnf_transformation,[],[f3011_D])).
fof(f3011_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f4088,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl50 | $spl83 | $spl84)),
  inference(forward_demodulation,[],[f4065,f1541])).
fof(f4065,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1))),sK1)))) | ($spl50 | $spl83 | $spl84)),
  inference(superposition,[],[f2479,f1541])).
fof(f4083,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)))) | $spl194),
  inference(cnf_transformation,[],[f4083_D])).
fof(f4083_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f4085,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl196),
  inference(cnf_transformation,[],[f4085_D])).
fof(f4085_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f4055,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1),sK1)))) | $spl50),
  inference(superposition,[],[f1541,f70])).
fof(f3985,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)))) | $spl190),
  inference(cnf_transformation,[],[f3985_D])).
fof(f3985_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f3987,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl192),
  inference(cnf_transformation,[],[f3987_D])).
fof(f3987_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f3958,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1),sK1)))) | ($spl44 | $spl59 | $spl84)),
  inference(superposition,[],[f2906,f70])).
fof(f2906,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl59 | $spl84)),
  inference(forward_demodulation,[],[f2905,f2686])).
fof(f2905,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl59 | $spl84)),
  inference(subsumption_resolution,[],[f2897,f1770])).
fof(f2897,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl84)),
  inference(superposition,[],[f2686,f70])).
fof(f3872,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl177),
  inference(cnf_transformation,[],[f3872_D])).
fof(f3872_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f3886,plain,(
  ( ! [X10] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)))) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),sK7(X10,sK1))))) ) | $spl188),
  inference(cnf_transformation,[],[f3886_D])).
fof(f3886_D,plain,(
  ( ! [X10] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)))) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),sK7(X10,sK1))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f3938,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),X8) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X8,sK1)),s('type.list.list'(sK1),sK7(X8,sK1))))) ) | ($spl44 | $spl84 | $spl99 | $spl162)),
  inference(forward_demodulation,[],[f3914,f3526])).
fof(f3914,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s('type.list.list'(sK1),X8) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X8,sK1)),s('type.list.list'(sK1),sK7(X8,sK1))))) ) | ($spl44 | $spl84 | $spl99)),
  inference(superposition,[],[f3898,f2895])).
fof(f3884,plain,(
  ( ! [X9] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X9)))) ) | $spl186),
  inference(cnf_transformation,[],[f3884_D])).
fof(f3884_D,plain,(
  ( ! [X9] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X9)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f3937,plain,(
  ( ! [X7] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7)),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7)))) ) | ($spl44 | $spl84 | $spl99 | $spl162)),
  inference(forward_demodulation,[],[f3913,f3526])).
fof(f3913,plain,(
  ( ! [X7] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7)))) ) | ($spl44 | $spl84 | $spl99)),
  inference(superposition,[],[f3898,f3104])).
fof(f3104,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X4)))) ) | ($spl44 | $spl84)),
  inference(resolution,[],[f3067,f72])).
fof(f3067,plain,(
  ( ! [X6] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f81,f2895])).
fof(f3882,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X8)))) ) | $spl184),
  inference(cnf_transformation,[],[f3882_D])).
fof(f3882_D,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X8)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f3936,plain,(
  ( ! [X6] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) ) | ($spl44 | $spl84 | $spl99 | $spl162)),
  inference(forward_demodulation,[],[f3912,f3526])).
fof(f3912,plain,(
  ( ! [X6] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) ) | ($spl44 | $spl84 | $spl99)),
  inference(superposition,[],[f3898,f3106])).
fof(f3106,plain,(
  ( ! [X6] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) ) | ($spl44 | $spl84)),
  inference(resolution,[],[f3067,f75])).
fof(f3880,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X7) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) ) | $spl182),
  inference(cnf_transformation,[],[f3880_D])).
fof(f3880_D,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X7) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f3935,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4)),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,X5) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl44 | $spl84 | $spl99 | $spl162)),
  inference(forward_demodulation,[],[f3911,f3526])).
fof(f3911,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4)),sK1)))) | s(bool,X5) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl44 | $spl84 | $spl99)),
  inference(superposition,[],[f3898,f3265])).
fof(f3265,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X1))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl44 | $spl84)),
  inference(resolution,[],[f3102,f78])).
fof(f3102,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0)))) ) | ($spl44 | $spl84)),
  inference(resolution,[],[f3067,f66])).
fof(f3878,plain,(
  ( ! [X5] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1))))) ) | $spl180),
  inference(cnf_transformation,[],[f3878_D])).
fof(f3878_D,plain,(
  ( ! [X5] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f3934,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) ) | ($spl99 | $spl162)),
  inference(forward_demodulation,[],[f3910,f3526])).
fof(f3910,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1))))) ) | ($spl99 | $spl162)),
  inference(superposition,[],[f3898,f3526])).
fof(f3016,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl144),
  inference(cnf_transformation,[],[f3016_D])).
fof(f3016_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f3933,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl44 | $spl84 | $spl99)),
  inference(forward_demodulation,[],[f3909,f2686])).
fof(f3909,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl84 | $spl99)),
  inference(superposition,[],[f3898,f2686])).
fof(f3875,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl178),
  inference(cnf_transformation,[],[f3875_D])).
fof(f3875_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f3932,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl99 | $spl150)),
  inference(forward_demodulation,[],[f3908,f3087])).
fof(f3908,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl99 | $spl150)),
  inference(superposition,[],[f3898,f3087])).
fof(f1541,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl50),
  inference(cnf_transformation,[],[f1541_D])).
fof(f1541_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3931,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl36 | $spl99)),
  inference(forward_demodulation,[],[f3907,f1146])).
fof(f3907,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl36 | $spl99)),
  inference(superposition,[],[f3898,f1146])).
fof(f3863,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | $spl171),
  inference(cnf_transformation,[],[f3863_D])).
fof(f3863_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f2615,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | $spl104),
  inference(cnf_transformation,[],[f2615_D])).
fof(f2615_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3929,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl99)),
  inference(forward_demodulation,[],[f3905,f1153])).
fof(f3905,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl99)),
  inference(superposition,[],[f3898,f1153])).
fof(f3873,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl176),
  inference(cnf_transformation,[],[f3873_D])).
fof(f3873_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f3885,plain,(
  ( ! [X10] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),sK7(X10,sK1))))) ) | ($spl44 | $spl84 | $spl98 | $spl162)),
  inference(forward_demodulation,[],[f3800,f3526])).
fof(f3800,plain,(
  ( ! [X10] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10)),sK1)))) | s('type.list.list'(sK1),X10) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X10,sK1)),s('type.list.list'(sK1),sK7(X10,sK1))))) ) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2546,f2895])).
fof(f2546,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | $spl98),
  inference(superposition,[],[f2539,f70])).
fof(f3883,plain,(
  ( ! [X9] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X9)))) ) | ($spl44 | $spl84 | $spl98 | $spl162)),
  inference(forward_demodulation,[],[f3799,f3526])).
fof(f3799,plain,(
  ( ! [X9] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X9)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X9)))) ) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2546,f3104])).
fof(f3881,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X8)))) ) | ($spl44 | $spl84 | $spl98 | $spl162)),
  inference(forward_demodulation,[],[f3798,f3526])).
fof(f3798,plain,(
  ( ! [X8] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X8)),sK1)))) | s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X8)))) ) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2546,f3106])).
fof(f3879,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,X7) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl44 | $spl84 | $spl98 | $spl162)),
  inference(forward_demodulation,[],[f3797,f3526])).
fof(f3797,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X6)),sK1)))) | s(bool,X7) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X6))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2546,f3265])).
fof(f3877,plain,(
  ( ! [X5] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) ) | ($spl98 | $spl162)),
  inference(forward_demodulation,[],[f3796,f3526])).
fof(f3796,plain,(
  ( ! [X5] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X5)),sK1))))) ) | ($spl98 | $spl162)),
  inference(superposition,[],[f2546,f3526])).
fof(f3876,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl44 | $spl84 | $spl98)),
  inference(forward_demodulation,[],[f3795,f2686])).
fof(f3795,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2546,f2686])).
fof(f3871,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl98 | $spl150)),
  inference(forward_demodulation,[],[f3794,f3087])).
fof(f3794,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl98 | $spl150)),
  inference(superposition,[],[f2546,f3087])).
fof(f3870,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl174),
  inference(cnf_transformation,[],[f3870_D])).
fof(f3870_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f3868,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl36 | $spl98)),
  inference(forward_demodulation,[],[f3793,f1146])).
fof(f3793,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl36 | $spl98)),
  inference(superposition,[],[f2546,f1146])).
fof(f3866,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) | $spl172),
  inference(cnf_transformation,[],[f3866_D])).
fof(f3866_D,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f3864,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | $spl170),
  inference(cnf_transformation,[],[f3864_D])).
fof(f3864_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f3791,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),X4)),sK1)))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) | $spl98),
  inference(superposition,[],[f2546,f2546])).
fof(f2618,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl106),
  inference(cnf_transformation,[],[f2618_D])).
fof(f2618_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3865,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | $spl98),
  inference(forward_demodulation,[],[f3790,f2539])).
fof(f3790,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl98),
  inference(superposition,[],[f2546,f2539])).
fof(f3862,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl98)),
  inference(forward_demodulation,[],[f3789,f1153])).
fof(f3789,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl98)),
  inference(superposition,[],[f2546,f1153])).
fof(f3767,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) | $spl169),
  inference(cnf_transformation,[],[f3767_D])).
fof(f3767_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f3625,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | $spl166),
  inference(cnf_transformation,[],[f3625_D])).
fof(f3625_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f3765,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) | ($spl130 | $spl137 | $spl162)),
  inference(forward_demodulation,[],[f3753,f3587])).
fof(f3753,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl130 | $spl137 | $spl162)),
  inference(superposition,[],[f2913,f3587])).
fof(f3623,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl164),
  inference(cnf_transformation,[],[f3623_D])).
fof(f3623_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f3626,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl130 | $spl162)),
  inference(forward_demodulation,[],[f3618,f3587])).
fof(f3618,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl130 | $spl162)),
  inference(superposition,[],[f70,f3587])).
fof(f3621,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl130 | $spl162)),
  inference(forward_demodulation,[],[f3614,f3587])).
fof(f3614,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)))),s('type.list.list'(sK0),'const.list.NIL_0')),sK0)))) | ($spl130 | $spl162)),
  inference(superposition,[],[f3587,f70])).
fof(f3508,plain,(
  ( ! [X26,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X25))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X26)))) ) | $spl160),
  inference(cnf_transformation,[],[f3508_D])).
fof(f3508_D,plain,(
  ( ! [X26,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X25))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X26)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f3560,plain,(
  ( ! [X76,X77,X75] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X75))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X75))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3471,f76])).
fof(f3471,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X75))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X75))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3102,f3265])).
fof(f3526,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0)))) ) | $spl162),
  inference(cnf_transformation,[],[f3526_D])).
fof(f3526_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f3551,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) = s(bool,X18) | p(s(bool,X18)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X16))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3445,f3067])).
fof(f3445,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) = s(bool,X18) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X16)))) | p(s(bool,X18)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X16))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f66,f3265])).
fof(f3550,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | p(s(bool,X15)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3444,f3067])).
fof(f3444,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X13)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f66,f3265])).
fof(f3549,plain,(
  ( ! [X97,X98] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X98))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3419,f3067])).
fof(f3419,plain,(
  ( ! [X97,X98] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X98)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X98))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f82,f3265])).
fof(f3540,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3392,f3067])).
fof(f3392,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10)))) | s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f66,f3265])).
fof(f3539,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3391,f3067])).
fof(f3391,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7)))) | s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f66,f3265])).
fof(f3528,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,X11))) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3359,f3067])).
fof(f3359,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10)))) | p(s(bool,X11))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3265,f66])).
fof(f3527,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) = s(bool,X9) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X9) | p(s(bool,X8))) ) | ($spl44 | $spl84)),
  inference(subsumption_resolution,[],[f3358,f3067])).
fof(f3358,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) = s(bool,X9) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X9) | p(s(bool,X8)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3265,f66])).
fof(f3503,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | ($spl44 | $spl84)),
  inference(duplicate_literal_removal,[],[f3353])).
fof(f3353,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3265,f3265])).
fof(f3504,plain,(
  ( ! [X26,X24,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X26))) | s(bool,'T_0') = s(bool,X24) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X25)))) ) | ($spl44 | $spl84)),
  inference(duplicate_literal_removal,[],[f3327])).
fof(f3327,plain,(
  ( ! [X26,X24,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X26))) | s(bool,'T_0') = s(bool,X24) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X25))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3265,f3265])).
fof(f3192,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X2)))) ) | $spl158),
  inference(cnf_transformation,[],[f3192_D])).
fof(f3192_D,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X2)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f3087,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl150),
  inference(cnf_transformation,[],[f3087_D])).
fof(f3087_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f3171,plain,(
  ( ! [X2] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X2)),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X2)))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f2895,f3104])).
fof(f3188,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4))))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) | $spl156),
  inference(cnf_transformation,[],[f3188_D])).
fof(f3188_D,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4))))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f3161,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X4))))) | s('type.list.list'(sK1),X4) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X4,sK1)),s('type.list.list'(sK1),sK7(X4,sK1))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3104,f2895])).
fof(f3187,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X3)))) ) | $spl154),
  inference(cnf_transformation,[],[f3187_D])).
fof(f3187_D,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X3)))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f3160,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X3)))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f3104,f3104])).
fof(f3088,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) | $spl152),
  inference(cnf_transformation,[],[f3088_D])).
fof(f3088_D,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f3037,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) | ($spl44 | $spl84)),
  inference(superposition,[],[f2895,f2895])).
fof(f3085,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl44 | $spl84)),
  inference(forward_demodulation,[],[f3036,f2686])).
fof(f3036,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl84)),
  inference(superposition,[],[f2895,f2686])).
fof(f3084,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl148),
  inference(cnf_transformation,[],[f3084_D])).
fof(f3084_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3082,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl36 | $spl44 | $spl84)),
  inference(forward_demodulation,[],[f3035,f1146])).
fof(f3035,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl36 | $spl44 | $spl84)),
  inference(superposition,[],[f2895,f1146])).
fof(f3079,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | $spl146),
  inference(cnf_transformation,[],[f3079_D])).
fof(f3079_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f3080,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl44 | $spl84 | $spl98)),
  inference(forward_demodulation,[],[f3033,f2539])).
fof(f3033,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl84 | $spl98)),
  inference(superposition,[],[f2895,f2539])).
fof(f3077,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl44 | $spl84)),
  inference(forward_demodulation,[],[f3032,f1153])).
fof(f3032,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl44 | $spl84)),
  inference(superposition,[],[f2895,f1153])).
fof(f3014,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl143),
  inference(cnf_transformation,[],[f3014_D])).
fof(f3014_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f3012,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl44 | $spl83 | $spl84)),
  inference(forward_demodulation,[],[f2992,f2686])).
fof(f2992,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl44 | $spl83 | $spl84)),
  inference(superposition,[],[f2479,f2686])).
fof(f3009,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl36 | $spl83 | $spl84)),
  inference(forward_demodulation,[],[f2991,f1146])).
fof(f2991,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl36 | $spl83 | $spl84)),
  inference(superposition,[],[f2479,f1146])).
fof(f3006,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | $spl139),
  inference(cnf_transformation,[],[f3006_D])).
fof(f3006_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f3007,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl83 | $spl84 | $spl98)),
  inference(forward_demodulation,[],[f2989,f2539])).
fof(f2989,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl83 | $spl84 | $spl98)),
  inference(superposition,[],[f2479,f2539])).
fof(f3004,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl83 | $spl84)),
  inference(forward_demodulation,[],[f2988,f1153])).
fof(f2988,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl83 | $spl84)),
  inference(superposition,[],[f2479,f1153])).
fof(f2973,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | ($spl36 | $spl47)),
  inference(forward_demodulation,[],[f2952,f1146])).
fof(f2952,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | ($spl36 | $spl47)),
  inference(superposition,[],[f2864,f1146])).
fof(f2609,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | $spl101),
  inference(cnf_transformation,[],[f2609_D])).
fof(f2609_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f2971,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl47 | $spl98)),
  inference(forward_demodulation,[],[f2950,f2539])).
fof(f2950,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl47 | $spl98)),
  inference(superposition,[],[f2864,f2539])).
fof(f2970,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl47)),
  inference(forward_demodulation,[],[f2949,f1153])).
fof(f2949,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl47)),
  inference(superposition,[],[f2864,f1153])).
fof(f2943,plain,(
  $false | ($spl66 | $spl130)),
  inference(subsumption_resolution,[],[f2938,f82])).
fof(f2938,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl66 | $spl130)),
  inference(superposition,[],[f81,f2930])).
fof(f2930,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK5,sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl66 | $spl130)),
  inference(forward_demodulation,[],[f2921,f2712])).
fof(f2921,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK6(sK5,sK1)))),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl66 | $spl130)),
  inference(superposition,[],[f2857,f2007])).
fof(f2876,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl137),
  inference(cnf_transformation,[],[f2876_D])).
fof(f2876_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f2870,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) | ($spl5 | $spl130)),
  inference(superposition,[],[f2852,f70])).
fof(f2852,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | ($spl5 | $spl130)),
  inference(backward_demodulation,[],[f2712,f140])).
fof(f140,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl5),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f2851,plain,(
  $false | ($spl5 | $spl131)),
  inference(subsumption_resolution,[],[f2850,f140])).
fof(f2850,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl131),
  inference(trivial_inequality_removal,[],[f2849])).
fof(f2849,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.NIL_0') | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl131),
  inference(superposition,[],[f2711,f70])).
fof(f2751,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) | $spl134),
  inference(cnf_transformation,[],[f2751_D])).
fof(f2751_D,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f2711,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) | $spl131),
  inference(cnf_transformation,[],[f2711_D])).
fof(f2711_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f2747,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) | $spl5),
  inference(superposition,[],[f140,f70])).
fof(f2708,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl129),
  inference(cnf_transformation,[],[f2708_D])).
fof(f2708_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f2744,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) | $spl5),
  inference(superposition,[],[f140,f70])).
fof(f2716,plain,(
  $false | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f2715,f92])).
fof(f2715,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl20)),
  inference(forward_demodulation,[],[f2700,f155])).
fof(f2700,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | $spl4),
  inference(superposition,[],[f81,f141])).
fof(f2712,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) | $spl130),
  inference(cnf_transformation,[],[f2712_D])).
fof(f2712_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2714,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) | $spl132),
  inference(cnf_transformation,[],[f2714_D])).
fof(f2714_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f2710,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) | $spl4),
  inference(forward_demodulation,[],[f2699,f141])).
fof(f2699,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)),s('type.list.list'(sK0),sK7('const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))),sK0)))) | $spl4),
  inference(superposition,[],[f141,f70])).
fof(f2709,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl128),
  inference(cnf_transformation,[],[f2709_D])).
fof(f2709_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f2698,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)),s('type.list.list'(sK0),sK7(sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0),sK0)))) | $spl4),
  inference(superposition,[],[f141,f70])).
fof(f2694,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl125),
  inference(cnf_transformation,[],[f2694_D])).
fof(f2694_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f1768,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl56),
  inference(cnf_transformation,[],[f1768_D])).
fof(f1768_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2461,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl39 | $spl84)),
  inference(superposition,[],[f2350,f2409])).
fof(f2350,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1)))) | s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0)))) ) | $spl39),
  inference(superposition,[],[f1147,f70])).
fof(f1147,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl39),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  s('type.list.list'(sK1),sK5) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f2692,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl122),
  inference(cnf_transformation,[],[f2692_D])).
fof(f2692_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2462,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | ($spl2 | $spl33 | $spl84)),
  inference(superposition,[],[f639,f2409])).
fof(f1771,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | $spl58),
  inference(cnf_transformation,[],[f1771_D])).
fof(f1771_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2464,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | $spl84),
  inference(superposition,[],[f70,f2409])).
fof(f1146,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl36),
  inference(cnf_transformation,[],[f1146_D])).
fof(f1146_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2647,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl116),
  inference(cnf_transformation,[],[f2647_D])).
fof(f2647_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2519,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl42),
  inference(superposition,[],[f1153,f70])).
fof(f2645,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl114),
  inference(cnf_transformation,[],[f2645_D])).
fof(f2645_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2545,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl98),
  inference(superposition,[],[f2539,f70])).
fof(f2690,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl120),
  inference(cnf_transformation,[],[f2690_D])).
fof(f2690_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2642,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f1811])).
fof(f1811,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) ) | $spl44),
  inference(superposition,[],[f1759,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f50])).
fof(f2641,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f1810])).
fof(f1810,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | ~p(s(bool,X2))) ) | $spl44),
  inference(superposition,[],[f1759,f69])).
fof(f2685,plain,(
  $false | ($spl29 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2684,f82])).
fof(f2684,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl29 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2683,f1771])).
fof(f2683,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | ($spl29 | $spl44)),
  inference(subsumption_resolution,[],[f1851,f178])).
fof(f1851,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f78,f1802])).
fof(f1802,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl44),
  inference(resolution,[],[f1759,f72])).
fof(f2682,plain,(
  $false | ($spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2681,f82])).
fof(f2681,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f2637,f1771])).
fof(f2637,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(trivial_inequality_removal,[],[f1894])).
fof(f1894,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f1804,f78])).
fof(f2680,plain,(
  $false | ($spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2679,f82])).
fof(f2679,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f2638,f1771])).
fof(f2638,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(trivial_inequality_removal,[],[f1839])).
fof(f1839,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f71,f1802])).
fof(f2678,plain,(
  $false | ($spl2 | $spl29 | $spl33 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2677,f178])).
fof(f2677,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl33 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2676,f638])).
fof(f2676,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f1802,f1771])).
fof(f2675,plain,(
  $false | ($spl2 | $spl29 | $spl33 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2674,f178])).
fof(f2674,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl33 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2673,f638])).
fof(f2673,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f2639,f1771])).
fof(f2639,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl44),
  inference(trivial_inequality_removal,[],[f1895])).
fof(f1895,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl44),
  inference(superposition,[],[f1804,f79])).
fof(f2672,plain,(
  $false | ($spl2 | $spl33 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2671,f638])).
fof(f2671,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f1804,f1771])).
fof(f2670,plain,(
  $false | ($spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2669,f82])).
fof(f2669,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f2640,f1771])).
fof(f2640,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(trivial_inequality_removal,[],[f1842])).
fof(f1842,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f73,f1802])).
fof(f2664,plain,(
  $false | ($spl29 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2663,f82])).
fof(f2663,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl29 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2662,f1771])).
fof(f2662,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | ($spl29 | $spl44)),
  inference(subsumption_resolution,[],[f1828,f178])).
fof(f1828,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f1802,f76])).
fof(f2661,plain,(
  $false | ($spl29 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2660,f82])).
fof(f2660,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl29 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2659,f1771])).
fof(f2659,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | ($spl29 | $spl44)),
  inference(subsumption_resolution,[],[f1829,f178])).
fof(f1829,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f1802,f78])).
fof(f2658,plain,(
  $false | ($spl29 | $spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2657,f82])).
fof(f2657,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl29 | $spl44 | $spl58)),
  inference(forward_demodulation,[],[f2656,f1771])).
fof(f2656,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | ($spl29 | $spl44)),
  inference(subsumption_resolution,[],[f1847,f178])).
fof(f1847,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f76,f1802])).
fof(f2655,plain,(
  $false | ($spl44 | $spl58)),
  inference(subsumption_resolution,[],[f2654,f82])).
fof(f2654,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f2643,f1771])).
fof(f2643,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(trivial_inequality_removal,[],[f1893])).
fof(f1893,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))))) | $spl44),
  inference(superposition,[],[f1804,f76])).
fof(f1536,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl46),
  inference(cnf_transformation,[],[f1536_D])).
fof(f1536_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2535,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl42 | $spl58)),
  inference(forward_demodulation,[],[f2519,f1771])).
fof(f2556,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl58 | $spl98)),
  inference(forward_demodulation,[],[f2545,f1771])).
fof(f2626,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) | $spl112),
  inference(cnf_transformation,[],[f2626_D])).
fof(f2626_D,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2623,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl108),
  inference(cnf_transformation,[],[f2623_D])).
fof(f2623_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2572,plain,(
  ( ! [X3] : (s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X3)),sK1)))) | s('type.list.list'(sK1),X3) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X3,sK1)),s('type.list.list'(sK1),sK7(X3,sK1))))) ) | $spl46),
  inference(superposition,[],[f1775,f1775])).
fof(f1775,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1))))) ) | $spl46),
  inference(superposition,[],[f1536,f70])).
fof(f2625,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl110),
  inference(cnf_transformation,[],[f2625_D])).
fof(f2625_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2621,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl46),
  inference(forward_demodulation,[],[f2571,f1536])).
fof(f2571,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl46),
  inference(superposition,[],[f1775,f1536])).
fof(f2610,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | $spl100),
  inference(cnf_transformation,[],[f2610_D])).
fof(f2610_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2616,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl46 | $spl98)),
  inference(forward_demodulation,[],[f2568,f2539])).
fof(f2568,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | ($spl46 | $spl98)),
  inference(superposition,[],[f1775,f2539])).
fof(f2613,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl42 | $spl46)),
  inference(forward_demodulation,[],[f2567,f1153])).
fof(f2567,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1))),sK1)))) | ($spl42 | $spl46)),
  inference(superposition,[],[f1775,f1153])).
fof(f2612,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)))) | $spl102),
  inference(cnf_transformation,[],[f2612_D])).
fof(f2612_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2608,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | ($spl41 | $spl46 | $spl66)),
  inference(forward_demodulation,[],[f2566,f2026])).
fof(f2026,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | ($spl41 | $spl66)),
  inference(forward_demodulation,[],[f2025,f2007])).
fof(f2025,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | ($spl41 | $spl66)),
  inference(subsumption_resolution,[],[f2016,f1150])).
fof(f1150,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK5) | $spl41),
  inference(cnf_transformation,[],[f1150_D])).
fof(f1150_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK5) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f2016,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) | $spl66),
  inference(superposition,[],[f2007,f70])).
fof(f2566,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1))),sK1)))) | ($spl41 | $spl46 | $spl66)),
  inference(superposition,[],[f1775,f2026])).
fof(f2537,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) | $spl96),
  inference(cnf_transformation,[],[f2537_D])).
fof(f2537_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f2539,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl98),
  inference(cnf_transformation,[],[f2539_D])).
fof(f2539_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f2520,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1),sK1)))) | $spl42),
  inference(superposition,[],[f1153,f70])).
fof(f2510,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) | $spl92),
  inference(cnf_transformation,[],[f2510_D])).
fof(f2510_D,plain,(
  s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2512,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl94),
  inference(cnf_transformation,[],[f2512_D])).
fof(f2512_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2494,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5)),sK1),sK1)))) | ($spl41 | $spl66)),
  inference(superposition,[],[f2026,f70])).
fof(f2458,plain,(
  $false | ($spl55 | $spl85)),
  inference(subsumption_resolution,[],[f2457,f1765])).
fof(f1765,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl55),
  inference(cnf_transformation,[],[f1765_D])).
fof(f1765_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f2457,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl85),
  inference(trivial_inequality_removal,[],[f2456])).
fof(f2456,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),'const.list.NIL_0') | s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl85),
  inference(superposition,[],[f2408,f70])).
fof(f2455,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | $spl90),
  inference(cnf_transformation,[],[f2455_D])).
fof(f2455_D,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2408,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | $spl85),
  inference(cnf_transformation,[],[f2408_D])).
fof(f2408_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f2450,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | $spl55),
  inference(superposition,[],[f1765,f70])).
fof(f2404,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)))) | $spl80),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2405,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl83),
  inference(cnf_transformation,[],[f2405_D])).
fof(f2405_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f2446,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)))) | $spl55),
  inference(superposition,[],[f1765,f70])).
fof(f2416,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl88),
  inference(cnf_transformation,[],[f2416_D])).
fof(f2416_D,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1770,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | $spl59),
  inference(cnf_transformation,[],[f1770_D])).
fof(f1770_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f1911,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0')),sK1)))) | $spl57),
  inference(superposition,[],[f1767,f70])).
fof(f2409,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | $spl84),
  inference(cnf_transformation,[],[f2409_D])).
fof(f2409_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2411,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | $spl86),
  inference(cnf_transformation,[],[f2411_D])).
fof(f2411_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2412,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | $spl54),
  inference(forward_demodulation,[],[f2400,f1766])).
fof(f2400,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | $spl54),
  inference(superposition,[],[f70,f1766])).
fof(f2407,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | $spl54),
  inference(forward_demodulation,[],[f2395,f1766])).
fof(f2395,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1))),sK1)))) | $spl54),
  inference(superposition,[],[f1766,f70])).
fof(f2406,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl82),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2394,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1),sK1)))) | $spl54),
  inference(superposition,[],[f1766,f70])).
fof(f2298,plain,(
  ( ! [X26,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X25))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X26)))) ) | $spl70),
  inference(cnf_transformation,[],[f2298_D])).
fof(f2298_D,plain,(
  ( ! [X26,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X25))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X26)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2346,plain,(
  ( ! [X76,X77,X75] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X75))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X75))) | s(bool,'T_0') = s(bool,X76)) ) | $spl38),
  inference(subsumption_resolution,[],[f2265,f76])).
fof(f2265,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X76)) | s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X75))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X77))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X75))) | s(bool,'T_0') = s(bool,X76)) ) | $spl38),
  inference(superposition,[],[f1937,f2064])).
fof(f2064,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X1))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(resolution,[],[f1937,f78])).
fof(f1937,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0)))) ) | $spl38),
  inference(resolution,[],[f1930,f66])).
fof(f1930,plain,(
  ( ! [X2] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X2)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X2)))) ) | $spl38),
  inference(superposition,[],[f81,f1543])).
fof(f1543,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),X0) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(X0,sK1)),s('type.list.list'(sK1),sK7(X0,sK1)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0)))) ) | $spl38),
  inference(superposition,[],[f1148,f70])).
fof(f2312,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0)))) ) | $spl72),
  inference(cnf_transformation,[],[f2312_D])).
fof(f2312_D,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2337,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) = s(bool,X18) | p(s(bool,X18)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X16))) | s(bool,'T_0') = s(bool,X17)) ) | $spl38),
  inference(subsumption_resolution,[],[f2239,f1930])).
fof(f2239,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) = s(bool,X18) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X16)))) | p(s(bool,X18)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X16))) | s(bool,'T_0') = s(bool,X17)) ) | $spl38),
  inference(superposition,[],[f66,f2064])).
fof(f2336,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | p(s(bool,X15)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X14)) ) | $spl38),
  inference(subsumption_resolution,[],[f2238,f1930])).
fof(f2238,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X13)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X13))) | s(bool,'T_0') = s(bool,X14)) ) | $spl38),
  inference(superposition,[],[f66,f2064])).
fof(f2334,plain,(
  ( ! [X97,X98] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X98))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) | $spl38),
  inference(subsumption_resolution,[],[f2212,f1930])).
fof(f2212,plain,(
  ( ! [X97,X98] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X98)))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X98))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X97),'const.list.NIL_0')))) ) | $spl38),
  inference(superposition,[],[f82,f2064])).
fof(f2325,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl38),
  inference(subsumption_resolution,[],[f2185,f1930])).
fof(f2185,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10)))) | s(bool,X9) = s(bool,X11) | p(s(bool,X11)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X9)) ) | $spl38),
  inference(superposition,[],[f66,f2064])).
fof(f2324,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | $spl38),
  inference(subsumption_resolution,[],[f2184,f1930])).
fof(f2184,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7)))) | s(bool,X6) = s(bool,X8) | p(s(bool,X8)) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X6)) ) | $spl38),
  inference(superposition,[],[f66,f2064])).
fof(f2316,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,X11))) ) | $spl38),
  inference(subsumption_resolution,[],[f2155,f1930])).
fof(f2155,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X10))) | s(bool,'T_0') = s(bool,X12) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X10)))) | p(s(bool,X11))) ) | $spl38),
  inference(superposition,[],[f2064,f66])).
fof(f2314,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) = s(bool,X9) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X9) | p(s(bool,X8))) ) | $spl38),
  inference(subsumption_resolution,[],[f2154,f1930])).
fof(f2154,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) = s(bool,X9) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X7))) | s(bool,'T_0') = s(bool,X9) | p(s(bool,X8)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X7))))) ) | $spl38),
  inference(superposition,[],[f2064,f66])).
fof(f2292,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | $spl38),
  inference(duplicate_literal_removal,[],[f2149])).
fof(f2149,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X2) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl38),
  inference(superposition,[],[f2064,f2064])).
fof(f2293,plain,(
  ( ! [X26,X24,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X26))) | s(bool,'T_0') = s(bool,X24) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X25)))) ) | $spl38),
  inference(duplicate_literal_removal,[],[f2125])).
fof(f2125,plain,(
  ( ! [X26,X24,X25] : (s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X25))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X26))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X26))) | s(bool,'T_0') = s(bool,X24) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X25))) | s(bool,'T_0') = s(bool,X24)) ) | $spl38),
  inference(superposition,[],[f2064,f2064])).
fof(f2005,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) | $spl64),
  inference(cnf_transformation,[],[f2005_D])).
fof(f2005_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2007,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5))) | $spl66),
  inference(cnf_transformation,[],[f2007_D])).
fof(f2007_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1983,plain,(
  ( ! [X0] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X0))))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X0)))) ) | $spl38),
  inference(superposition,[],[f1939,f1939])).
fof(f1939,plain,(
  ( ! [X4] : (s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),X4))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),X4)))) ) | $spl38),
  inference(resolution,[],[f1930,f72])).
fof(f1799,plain,(
  $false | ($spl45 | $spl59)),
  inference(subsumption_resolution,[],[f1798,f1533])).
fof(f1533,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl45),
  inference(cnf_transformation,[],[f1533_D])).
fof(f1533_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1798,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl59),
  inference(trivial_inequality_removal,[],[f1797])).
fof(f1797,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),'const.list.NIL_0') | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl59),
  inference(superposition,[],[f1770,f70])).
fof(f1796,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | $spl62),
  inference(cnf_transformation,[],[f1796_D])).
fof(f1796_D,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1792,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | $spl45),
  inference(superposition,[],[f1533,f70])).
fof(f1766,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl54),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1767,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl57),
  inference(cnf_transformation,[],[f1767_D])).
fof(f1767_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f1789,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl45),
  inference(superposition,[],[f1533,f70])).
fof(f1534,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl44),
  inference(cnf_transformation,[],[f1534_D])).
fof(f1534_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1727,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl36),
  inference(superposition,[],[f1146,f70])).
fof(f1773,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | $spl60),
  inference(cnf_transformation,[],[f1773_D])).
fof(f1773_D,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1774,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | $spl44),
  inference(forward_demodulation,[],[f1762,f1534])).
fof(f1762,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | $spl44),
  inference(superposition,[],[f70,f1534])).
fof(f1769,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | $spl44),
  inference(forward_demodulation,[],[f1758,f1534])).
fof(f1758,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1))),sK1)))) | $spl44),
  inference(superposition,[],[f1534,f70])).
fof(f1757,plain,(
  s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK7(sK5,sK1),sK1),sK1),sK1)))) | $spl44),
  inference(superposition,[],[f1534,f70])).
fof(f1596,plain,(
  $false | ($spl36 | $spl48)),
  inference(subsumption_resolution,[],[f1589,f82])).
fof(f1589,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl36 | $spl48)),
  inference(backward_demodulation,[],[f1539,f1527])).
fof(f1566,plain,(
  $false | ($spl37 | $spl49)),
  inference(subsumption_resolution,[],[f1565,f1145])).
fof(f1145,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl37),
  inference(cnf_transformation,[],[f1145_D])).
fof(f1145_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1565,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl49),
  inference(trivial_inequality_removal,[],[f1564])).
fof(f1564,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),'const.list.NIL_0') | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl49),
  inference(superposition,[],[f1538,f70])).
fof(f1562,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl52),
  inference(cnf_transformation,[],[f1562_D])).
fof(f1562_D,plain,(
  s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1538,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK5,sK1)) | $spl49),
  inference(cnf_transformation,[],[f1538_D])).
fof(f1538_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK5,sK1)) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1559,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') != s('type.list.list'(sK1),sK7(sK5,sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl37),
  inference(superposition,[],[f1145,f70])).
fof(f1535,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl47),
  inference(cnf_transformation,[],[f1535_D])).
fof(f1535_D,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f1557,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) != s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl37),
  inference(superposition,[],[f1145,f70])).
fof(f1539,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) | $spl48),
  inference(cnf_transformation,[],[f1539_D])).
fof(f1539_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1542,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) | $spl36),
  inference(forward_demodulation,[],[f1530,f1146])).
fof(f1530,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl36),
  inference(superposition,[],[f70,f1146])).
fof(f1537,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) | $spl36),
  inference(forward_demodulation,[],[f1526,f1146])).
fof(f1526,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK7(sK5,sK1)) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1))),sK1)))) | $spl36),
  inference(superposition,[],[f1146,f70])).
fof(f1525,plain,(
  s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK7(sK5,sK1),sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK7(sK5,sK1),sK1),sK1)))) | $spl36),
  inference(superposition,[],[f1146,f70])).
fof(f1385,plain,(
  $false | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f153,f1230])).
fof(f1230,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl14),
  inference(backward_demodulation,[],[f148,f64])).
fof(f64,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,sK1),sK3),s(fun(sK1,sK0),sK4)))) & s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) & ? [X5] : s(bool,'const.list.NULL_1'(s('type.list.list'(X1),X5))) != s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X5))))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5] : s(bool,'const.list.NULL_1'(s('type.list.list'(X1),X5))) = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X5))))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5] : s(bool,'const.list.NULL_1'(s('type.list.list'(X1),X5))) = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X4),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,X7),X9),s(fun(X7,X6),X10)))) => ! [X11] : s(bool,'const.list.NULL_1'(s('type.list.list'(X7),X11))) = s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,X7),X9),s(fun(X7,X6),X10)))) => ! [X11] : s(bool,'const.list.NULL_1'(s('type.list.list'(X7),X11))) = s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X10),s('type.list.list'(X7),X11))))))),
  file('thm.quotient_list.NULL_PRS','thm.quotient_list.NULL_PRS')).
fof(f1148,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl38),
  inference(cnf_transformation,[],[f1148_D])).
fof(f1148_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1372,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl0),
  inference(superposition,[],[f137,f70])).
fof(f145,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | $spl8),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f175,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1340,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f1324,f71])).
fof(f1324,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f1230,f66])).
fof(f1339,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f1323,f71])).
fof(f1323,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(superposition,[],[f1230,f66])).
fof(f1315,plain,(
  $false | ($spl6 | $spl17 | $spl33)),
  inference(subsumption_resolution,[],[f921,f151])).
fof(f921,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ($spl6 | $spl33)),
  inference(forward_demodulation,[],[f142,f316])).
fof(f142,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl6),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f155,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl20),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f149,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl15),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1200,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f153,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl18),
  inference(cnf_transformation,[],[f153_D])).
fof(f153_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f151,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl17),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1199,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f144,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | $spl9),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1198,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f78])).
fof(f1197,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f76])).
fof(f1196,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f74])).
fof(f74,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f56])).
fof(f1195,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f72])).
fof(f147,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1194,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f69])).
fof(f1193,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f64,f69])).
fof(f146,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0))) ) | $spl10),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1192,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f64,f66])).
fof(f1191,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f66])).
fof(f141,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl4),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1204,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl33),
  inference(forward_demodulation,[],[f1190,f316])).
fof(f1190,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))))),
  inference(superposition,[],[f64,f70])).
fof(f1186,plain,(
  $false | ($spl25 | $spl40)),
  inference(subsumption_resolution,[],[f1180,f82])).
fof(f1180,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl25 | $spl40)),
  inference(backward_demodulation,[],[f1151,f173])).
fof(f1151,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | $spl40),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1153,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl42),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1156,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | $spl0),
  inference(forward_demodulation,[],[f1142,f137])).
fof(f1142,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl0),
  inference(superposition,[],[f70,f137])).
fof(f1155,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f1154,f92])).
fof(f1154,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f1139,f150])).
fof(f150,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl16),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1139,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl0),
  inference(superposition,[],[f81,f137])).
fof(f1149,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | $spl0),
  inference(forward_demodulation,[],[f1138,f137])).
fof(f1138,plain,(
  s('type.list.list'(sK1),'const.list.NIL_0') = s('type.list.list'(sK1),sK5) | s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)),s('type.list.list'(sK1),sK7('const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))),sK1)))) | $spl0),
  inference(superposition,[],[f137,f70])).
fof(f1137,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK7(sK5,sK1)) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK7(sK5,sK1),sK1)),s('type.list.list'(sK1),sK7(sK7(sK5,sK1),sK1)))) | $spl0),
  inference(superposition,[],[f137,f70])).
fof(f388,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl34),
  inference(cnf_transformation,[],[f388_D])).
fof(f388_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1102,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1088,f75])).
fof(f1088,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(superposition,[],[f1017,f69])).
fof(f1017,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl16),
  inference(backward_demodulation,[],[f150,f64])).
fof(f1101,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1087,f75])).
fof(f1087,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) ) | $spl16),
  inference(superposition,[],[f1017,f69])).
fof(f996,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f995,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f994,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f78])).
fof(f993,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f76])).
fof(f992,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f74])).
fof(f991,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f72])).
fof(f990,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f69])).
fof(f989,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f64,f69])).
fof(f988,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f64,f66])).
fof(f987,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f66])).
fof(f1000,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl33),
  inference(forward_demodulation,[],[f986,f316])).
fof(f986,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))))),
  inference(superposition,[],[f64,f70])).
fof(f952,plain,(
  $false | ($spl2 | $spl33 | $spl34)),
  inference(trivial_inequality_removal,[],[f948])).
fof(f948,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl33 | $spl34)),
  inference(superposition,[],[f388,f638])).
fof(f951,plain,(
  $false | ($spl33 | $spl34)),
  inference(trivial_inequality_removal,[],[f949])).
fof(f949,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl33 | $spl34)),
  inference(superposition,[],[f388,f316])).
fof(f950,plain,(
  $false | $spl34),
  inference(equality_resolution,[],[f388])).
fof(f910,plain,(
  $false | ($spl2 | $spl12 | $spl16 | $spl33)),
  inference(trivial_inequality_removal,[],[f904])).
fof(f904,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl12 | $spl16 | $spl33)),
  inference(superposition,[],[f877,f638])).
fof(f877,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f876,f77])).
fof(f876,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f147,f150])).
fof(f909,plain,(
  $false | ($spl12 | $spl16)),
  inference(trivial_inequality_removal,[],[f905])).
fof(f905,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl16)),
  inference(superposition,[],[f877,f150])).
fof(f908,plain,(
  $false | ($spl12 | $spl16 | $spl33)),
  inference(trivial_inequality_removal,[],[f906])).
fof(f906,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl16 | $spl33)),
  inference(superposition,[],[f877,f316])).
fof(f907,plain,(
  $false | ($spl12 | $spl16)),
  inference(equality_resolution,[],[f877])).
fof(f872,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f858,f75])).
fof(f858,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(superposition,[],[f785,f69])).
fof(f785,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl16),
  inference(backward_demodulation,[],[f150,f64])).
fof(f871,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f857,f75])).
fof(f857,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) ) | $spl16),
  inference(superposition,[],[f785,f69])).
fof(f849,plain,(
  $false | ($spl7 | $spl16 | $spl33)),
  inference(subsumption_resolution,[],[f728,f150])).
fof(f728,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ($spl7 | $spl33)),
  inference(forward_demodulation,[],[f143,f316])).
fof(f764,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f763,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f762,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f78])).
fof(f761,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f76])).
fof(f760,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f74])).
fof(f759,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f72])).
fof(f758,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f69])).
fof(f757,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f64,f69])).
fof(f756,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f64,f66])).
fof(f755,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f66])).
fof(f768,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl33),
  inference(forward_demodulation,[],[f754,f316])).
fof(f754,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))))),
  inference(superposition,[],[f64,f70])).
fof(f441,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f76])).
fof(f442,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f78])).
fof(f443,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f719,plain,(
  $false | ($spl10 | $spl14)),
  inference(trivial_inequality_removal,[],[f715])).
fof(f715,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl14)),
  inference(superposition,[],[f688,f148])).
fof(f688,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f687,f73])).
fof(f687,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl10 | $spl14)),
  inference(forward_demodulation,[],[f146,f148])).
fof(f718,plain,(
  $false | ($spl10 | $spl14)),
  inference(equality_resolution,[],[f688])).
fof(f637,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f621,f71])).
fof(f621,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f446,f66])).
fof(f446,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl14),
  inference(backward_demodulation,[],[f148,f64])).
fof(f636,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f620,f71])).
fof(f620,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(superposition,[],[f446,f66])).
fof(f569,plain,(
  $false | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f568,f82])).
fof(f568,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f537,f178])).
fof(f537,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f78,f300])).
fof(f300,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl32),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f566,plain,(
  $false | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f565,f82])).
fof(f565,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f533,f178])).
fof(f533,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f76,f300])).
fof(f563,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f543,f82])).
fof(f543,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(trivial_inequality_removal,[],[f528])).
fof(f528,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f73,f300])).
fof(f562,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f544,f82])).
fof(f544,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(trivial_inequality_removal,[],[f525])).
fof(f525,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f71,f300])).
fof(f553,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f516,f86])).
fof(f516,plain,(
  p(s(bool,'F_0')) | $spl32),
  inference(superposition,[],[f82,f300])).
fof(f552,plain,(
  $false | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f551,f82])).
fof(f551,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f514,f178])).
fof(f514,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f300,f78])).
fof(f550,plain,(
  $false | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f549,f82])).
fof(f549,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl29 | $spl32)),
  inference(subsumption_resolution,[],[f513,f178])).
fof(f513,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl32),
  inference(superposition,[],[f300,f76])).
fof(f137,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl0),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f170,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl22),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f448,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl14),
  inference(superposition,[],[f148,f70])).
fof(f444,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f440,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f74])).
fof(f439,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f72])).
fof(f438,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f69])).
fof(f437,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f64,f69])).
fof(f436,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f64,f66])).
fof(f435,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f66])).
fof(f143,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl7),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f434,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))))),
  inference(superposition,[],[f64,f70])).
fof(f139,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl3),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f445,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))))),
  inference(forward_demodulation,[],[f433,f65])).
fof(f65,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : ! [X15] : s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X15),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.quotient_list.NULL_PRS','thm.list.MAP|split|0')).
fof(f433,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),'const.list.NIL_0'))))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))))),
  inference(superposition,[],[f64,f70])).
fof(f389,plain,(
  $false | ($spl15 | $spl17)),
  inference(subsumption_resolution,[],[f167,f151])).
fof(f167,plain,(
  s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl15),
  inference(trivial_inequality_removal,[],[f166])).
fof(f166,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | $spl15),
  inference(superposition,[],[f149,f79])).
fof(f173,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl25),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f256,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl17),
  inference(subsumption_resolution,[],[f243,f75])).
fof(f243,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | ~p(s(bool,X2))) ) | $spl17),
  inference(superposition,[],[f151,f69])).
fof(f258,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl17),
  inference(subsumption_resolution,[],[f244,f75])).
fof(f244,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl17),
  inference(superposition,[],[f151,f69])).
fof(f174,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl24),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f178,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f245,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl17),
  inference(superposition,[],[f151,f72])).
fof(f246,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl17),
  inference(superposition,[],[f151,f74])).
fof(f387,plain,(
  $false | ($spl3 | $spl33)),
  inference(subsumption_resolution,[],[f380,f82])).
fof(f380,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl3 | $spl33)),
  inference(trivial_inequality_removal,[],[f376])).
fof(f376,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl3 | $spl33)),
  inference(superposition,[],[f366,f78])).
fof(f366,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | ($spl3 | $spl33)),
  inference(forward_demodulation,[],[f139,f316])).
fof(f386,plain,(
  $false | ($spl3 | $spl33)),
  inference(subsumption_resolution,[],[f381,f82])).
fof(f381,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl3 | $spl33)),
  inference(trivial_inequality_removal,[],[f375])).
fof(f375,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0')))) | ($spl3 | $spl33)),
  inference(superposition,[],[f366,f76])).
fof(f303,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f287,f71])).
fof(f287,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) ) | $spl14),
  inference(superposition,[],[f180,f66])).
fof(f180,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))) | $spl14),
  inference(backward_demodulation,[],[f148,f64])).
fof(f302,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f286,f71])).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) ) | $spl14),
  inference(superposition,[],[f180,f66])).
fof(f301,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl33),
  inference(cnf_transformation,[],[f301_D])).
fof(f301_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f285,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)))) | $spl14),
  inference(superposition,[],[f180,f70])).
fof(f299,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl14),
  inference(forward_demodulation,[],[f284,f65])).
fof(f284,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),'const.list.NIL_0'))))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl14),
  inference(superposition,[],[f180,f70])).
fof(f255,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl31),
  inference(cnf_transformation,[],[f255_D])).
fof(f255_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f239,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl17),
  inference(superposition,[],[f151,f70])).
fof(f181,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl14),
  inference(superposition,[],[f148,f70])).
fof(f164,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl15),
  inference(superposition,[],[f149,f78])).
fof(f163,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | $spl15),
  inference(superposition,[],[f149,f76])).
fof(f176,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl15),
  inference(subsumption_resolution,[],[f158,f71])).
fof(f158,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5)))) | p(s(bool,X1))) ) | $spl15),
  inference(superposition,[],[f149,f66])).
fof(f172,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl15),
  inference(subsumption_resolution,[],[f157,f71])).
fof(f157,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))))) ) | $spl15),
  inference(superposition,[],[f149,f66])).
fof(f171,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | $spl23),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f156,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1)))) | $spl15),
  inference(superposition,[],[f149,f70])).
fof(f134,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f133,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))),
  inference(superposition,[],[f64,f79])).
fof(f132,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f78])).
fof(f131,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f76])).
fof(f130,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f74])).
fof(f129,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))),
  inference(superposition,[],[f64,f72])).
fof(f128,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f69])).
fof(f127,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f64,f69])).
fof(f126,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f64,f66])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | p(s(bool,X0)) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(superposition,[],[f64,f66])).
fof(f124,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),sK5))) | s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0)),s('type.list.list'(sK0),sK7('const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),sK5)),sK0))))),
  inference(superposition,[],[f64,f70])).
fof(f135,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))))),
  inference(forward_demodulation,[],[f123,f65])).
fof(f123,plain,(
  s(bool,'const.list.NULL_1'(s('type.list.list'(sK1),'const.list.NIL_0'))) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK1,sK0),sK4),s('type.list.list'(sK1),'const.list.NIL_0'))))) | s('type.list.list'(sK1),sK5) = s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK6(sK5,sK1)),s('type.list.list'(sK1),sK7(sK5,sK1))))),
  inference(superposition,[],[f64,f70])).
% SZS output end Proof for thm.quotient_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 246
% Passive clauses: 1071
% Generated clauses: 7330
% Final active clauses: 109
% Final passive clauses: 428
% Input formulas: 20
% Initial clauses: 31
% 
% Selected by SInE selection: 20
% SInE iterations: 1
% 
% Duplicate literals: 229
% Trivial inequalities: 221
% Fw subsumption resolutions: 818
% Fw demodulations: 883
% Bw demodulations: 76
% 
% Simple tautologies: 210
% Equational tautologies: 7
% Forward subsumptions: 3810
% Bw demodulations to eq. taut.: 10
% 
% Binary resolution: 288
% Forward superposition: 2614
% Backward superposition: 1812
% Self superposition: 44
% Equality factoring: 12
% Equality resolution: 12
% 
% Split clauses: 218
% Split components: 436
% SAT solver clauses: 275
% SAT solver unit clauses: 10
% SAT solver binary clauses: 87
% 
% Sat splits: 218
% Sat splitting refutations: 56
% 
% Memory used [KB]: 4477
% Time elapsed: 0.154 s
% ------------------------------
% ------------------------------
% Success in time 11.883 s
