Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.canonical
% (113805)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 150
% Passive clauses: 3897
% Generated clauses: 20635
% Final active clauses: 141
% Final passive clauses: 3590
% Input formulas: 68
% Initial clauses: 122
% 
% Duplicate literals: 585
% Trivial inequalities: 285
% Fw subsumption resolutions: 2021
% Fw demodulations: 3584
% Bw demodulations: 8
% 
% Simple tautologies: 306
% Equational tautologies: 51
% Forward subsumptions: 9437
% 
% Binary resolution: 500
% Factoring: 2
% Forward superposition: 6510
% Backward superposition: 6627
% Self superposition: 185
% Equality factoring: 92
% Equality resolution: 30
% 
% Split clauses: 254
% Split components: 512
% SAT solver clauses: 621
% SAT solver unit clauses: 10
% SAT solver binary clauses: 128
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 254
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 620
% TWLsolver calls for satisfiability: 138
% 
% Memory used [KB]: 19061
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.canonical
% (113826)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 1222
% Generated clauses: 14276
% Final active clauses: 99
% Final passive clauses: 815
% Input formulas: 68
% Initial clauses: 122
% 
% Duplicate literals: 523
% Trivial inequalities: 189
% Fw subsumption resolutions: 810
% Fw demodulations: 1459
% Bw demodulations: 281
% Condensations: 1162
% Global subsumptions: 8
% 
% Simple tautologies: 189
% Equational tautologies: 44
% Forward subsumptions: 6884
% Bw demodulations to eq. taut.: 19
% 
% Binary resolution: 338
% Factoring: 1
% Forward superposition: 4515
% Backward superposition: 3378
% Self superposition: 132
% Equality factoring: 74
% Equality resolution: 21
% 
% Split clauses: 81
% Split components: 163
% SAT solver clauses: 1837
% SAT solver unit clauses: 147
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 1
% 
% Sat splits: 81
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1804
% TWLsolver calls for satisfiability: 419
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.canonical
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.canonical
% SZS output start Proof for thm.canonical
fof(f31189,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f243,f374,f372,f370,f242,f373,f241,f371,f369,f378,f1262,f1272,f1271,f2424,f2438,f2436,f5064,f5066,f5065,f5069,f5071,f5073,f5077,f5079,f5081,f5083,f5085,f5087,f5089,f5091,f5093,f5095,f5097,f5099,f5101,f5103,f5105,f5107,f5109,f5111,f5115,f5117,f5119,f5121,f5123,f5125,f5127,f5129,f5131,f5133,f5135,f5137,f5139,f5141,f5143,f5145,f5147,f5149,f5173,f5175,f5177,f5179,f5181,f5183,f5185,f5187,f5189,f5192,f5194,f5196,f5198,f5201,f5240,f5241,f5242,f5243,f5244,f5245,f5246,f5247,f5248,f5249,f5250,f5251,f5252,f5253,f5254,f5255,f5256,f5257,f5258,f5259,f5260,f5261,f5262,f5263,f5264,f5265,f5266,f5267,f5268,f5269,f5270,f5271,f5272,f5273,f5274,f5275,f5276,f5277,f5278,f5279,f5280,f5281,f19723,f19776,f19724,f19725,f19726,f19780,f19774,f19727,f19779,f19777,f24108,f24134,f24132,f29432,f29433,f29452,f29058,f29465,f29059,f29464,f29525,f31092,f31093,f31110,f30724,f29467,f31123,f30725,f29466,f31122,f31177])).
fof(f31177,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f31176,f29542])).
fof(f29542,plain,(
  ( ! [X0,X1] : (s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),X0))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f372,f309])).
fof(f309,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4))) & (s('type.canonical.spolynom'(X0),X1) != s('type.canonical.spolynom'(X0),X3) | s('type.canonical.spolynom'(X0),X2) != s('type.canonical.spolynom'(X0),X4) | s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4)))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4))) & ((s('type.canonical.spolynom'(X0),X1) != s('type.canonical.spolynom'(X0),X3) | s('type.canonical.spolynom'(X0),X2) != s('type.canonical.spolynom'(X0),X4)) | s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4)))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) <=> (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) <=> (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X7),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))) <=> (s('type.canonical.spolynom'(X6),X7) = s('type.canonical.spolynom'(X6),X9) & s('type.canonical.spolynom'(X6),X8) = s('type.canonical.spolynom'(X6),X10)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.canonical.spolynom_|49||49||split|2')).
fof(f31176,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | $spl0),
  inference(forward_demodulation,[],[f30845,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X11),s('type.canonical.spolynom'(X6),X12))) = s('type.canonical.spolynom'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6)),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X6),X11))),s('type.canonical.spolynom'(X6),X12)))),
  file('thm.canonical.spolynom_|49||49||split|3','ARITY_0_2_const.canonical.SPplus')).
fof(f30845,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) = s('type.canonical.spolynom'(sK0),happ(s(fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0)),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) | $spl0),
  inference(superposition,[],[f295,f29955])).
fof(f29955,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun('type.canonical.spolynom'(sK0),X8),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X8)),X9),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) = s(X8,happ(s(fun('type.canonical.spolynom'(sK0),X8),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X8)),X9),s('type.canonical.spolynom'(sK0),sK3))),s('type.canonical.spolynom'(sK0),sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f29735,f333])).
fof(f333,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X6] : ! [X7,X8,X14,X15,X16,X17] : s(X13,happ(s(fun('type.canonical.spolynom'(X6),X13),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X13)),X17),s('type.canonical.spolynom'(X6),X7))),s('type.canonical.spolynom'(X6),X8))) = s(X13,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X7),s('type.canonical.spolynom'(X6),X8))),s(fun('type.quote.index',X13),X14),s(fun(X6,X13),X15),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X13)),X16),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X13)),X17)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.canonical.spolynom_case_def|split|3')).
fof(f29735,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,happ(s(fun('type.canonical.spolynom'(sK0),X8),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X8)),X9),s('type.canonical.spolynom'(sK0),sK3))),s('type.canonical.spolynom'(sK0),sK4))) = s(X8,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s(fun('type.quote.index',X8),X10),s(fun(sK0,X8),X11),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X8)),X12),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X8)),X9)))) ) | $spl0),
  inference(superposition,[],[f333,f369])).
fof(f31122,plain,(
  s('type.canonical.spolynom'(sK0),sK3) = s('type.canonical.spolynom'(sK0),sK4) | $spl36),
  inference(cnf_transformation,[],[f31122_D])).
fof(f31122_D,plain,(
  s('type.canonical.spolynom'(sK0),sK3) = s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f29466,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | $spl35),
  inference(cnf_transformation,[],[f29466_D])).
fof(f29466_D,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f30725,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | s('type.canonical.spolynom'(sK0),sK3) = s('type.canonical.spolynom'(sK0),sK4) | $spl0),
  inference(superposition,[],[f334,f29955])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X47,X48,X49] : (p(s(bool,happ(s(fun(X47,bool),happ(s(fun(X47,fun(X47,bool)),'=_0'),s(X47,X48))),s(X47,X49)))) <=> s(X47,X48) = s(X47,X49))),
  file('thm.canonical.spolynom_|49||49||split|3',arity_0_HL_EQ)).
fof(f31123,plain,(
  s('type.canonical.spolynom'(sK0),sK3) != s('type.canonical.spolynom'(sK0),sK4) | $spl37),
  inference(cnf_transformation,[],[f31123_D])).
fof(f31123_D,plain,(
  s('type.canonical.spolynom'(sK0),sK3) != s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f29467,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | $spl34),
  inference(cnf_transformation,[],[f29467_D])).
fof(f29467_D,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f30724,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | s('type.canonical.spolynom'(sK0),sK3) != s('type.canonical.spolynom'(sK0),sK4) | $spl0),
  inference(superposition,[],[f335,f29955])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f31110,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f31109,f29542])).
fof(f31109,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | $spl0),
  inference(forward_demodulation,[],[f30678,f295])).
fof(f30678,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) = s('type.canonical.spolynom'(sK0),happ(s(fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0)),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) | $spl0),
  inference(superposition,[],[f29955,f295])).
fof(f31093,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f31072,f22484])).
fof(f22484,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f244,f2220,f316])).
fof(f316,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X18,X19] : ((p(s(bool,X18)) => p(s(bool,X19))) => ((p(s(bool,X19)) => p(s(bool,X18))) => s(bool,X18) = s(bool,X19)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.IMP_ANTISYM_AX')).
fof(f2220,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2158,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X13,X6] : ! [X22,X30] : s(X6,X22) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X13),'const.pair.,_2'(s(X6,X22),s(X13,X30)))))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.pair.FST')).
fof(f2158,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f325,f335])).
fof(f244,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  p(s(bool,'T_0'))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.TRUTH')).
fof(f31072,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl5)),
  inference(backward_demodulation,[],[f31069,f2125])).
fof(f2125,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2))) | $spl5),
  inference(unit_resulting_resolution,[],[f2004,f248])).
fof(f248,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X32] : (s(bool,'T_0') = s(bool,X32) <=> p(s(bool,X32)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f2004,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2)))) | $spl5),
  inference(unit_resulting_resolution,[],[f374,f334])).
fof(f31069,plain,(
  ( ! [X62,X63] : (s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),X63),s('type.canonical.spolynom'(sK0),sK2))) = s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),X63),s('type.canonical.spolynom'(sK0),sK4)))) ) | $spl0),
  inference(forward_demodulation,[],[f30657,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X25,X23] : ! [X26,X28] : s(X23,X26) = s(X23,happ(s(fun(X25,X23),'_dst_x0_1'(s(X23,X26))),s(X25,X28)))),
  file('thm.canonical.spolynom_|49||49||split|3','_dst_x0_ax')).
fof(f30657,plain,(
  ( ! [X62,X63] : (s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),X63),s('type.canonical.spolynom'(sK0),sK4))) = s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X62)),'_dst_x0_1'(s(fun('type.canonical.spolynom'(sK0),X62),X63))),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) ) | $spl0),
  inference(superposition,[],[f29955,f324])).
fof(f31092,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f31091,f380])).
fof(f380,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f246,f248])).
fof(f246,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.canonical.spolynom_|49||49||split|3',aHL_FALSITY)).
fof(f31091,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f31071,f22484])).
fof(f31071,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl5)),
  inference(backward_demodulation,[],[f31069,f2124])).
fof(f2124,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2))) | $spl5),
  inference(unit_resulting_resolution,[],[f2004,f253])).
fof(f253,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X32] : (s(bool,'F_0') = s(bool,X32) <=> ~p(s(bool,X32)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f29525,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f29524,f3405])).
fof(f3405,plain,(
  ( ! [X0,X1] : (s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),X0),s('type.canonical.spolynom'(sK0),sK4))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),X1),s('type.canonical.spolynom'(sK0),sK2)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f374,f310])).
fof(f310,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)) )),
  inference(cnf_transformation,[],[f231])).
fof(f29524,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f29180,f295])).
fof(f29180,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK4))) = s('type.canonical.spolynom'(sK0),happ(s(fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0)),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f295,f16222])).
fof(f16222,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(sK0),X0),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) = s(X0,happ(s(fun('type.canonical.spolynom'(sK0),X0),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK4)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f15798,f333])).
fof(f15798,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(sK0),X0),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK4))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X4),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f333,f379])).
fof(f379,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f369,f371])).
fof(f29464,plain,(
  s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK4) | $spl32),
  inference(cnf_transformation,[],[f29464_D])).
fof(f29464_D,plain,(
  s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f29059,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK4) | ($spl0 | $spl2)),
  inference(superposition,[],[f334,f16222])).
fof(f29465,plain,(
  s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK4) | $spl33),
  inference(cnf_transformation,[],[f29465_D])).
fof(f29465_D,plain,(
  s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f29058,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2)))) | s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK4) | ($spl0 | $spl2)),
  inference(superposition,[],[f335,f16222])).
fof(f29452,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f29451,f3405])).
fof(f29451,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK4))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f29012,f295])).
fof(f29012,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK4))) = s('type.canonical.spolynom'(sK0),happ(s(fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0)),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),'type.canonical.spolynom'(sK0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(sK0),sK1))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl2)),
  inference(superposition,[],[f16222,f295])).
fof(f29433,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f29411,f22484])).
fof(f29411,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f28990,f2125])).
fof(f28990,plain,(
  ( ! [X62,X63] : (s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),X63),s('type.canonical.spolynom'(sK0),sK2))) = s(X62,happ(s(fun('type.canonical.spolynom'(sK0),X62),X63),s('type.canonical.spolynom'(sK0),sK4)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f16222,f324])).
fof(f29432,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f29431,f380])).
fof(f29431,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f29410,f22484])).
fof(f29410,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f28990,f2124])).
fof(f24132,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl28),
  inference(cnf_transformation,[],[f24132_D])).
fof(f24132_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f24134,plain,(
  p(s(bool,sK7('_i_0_0','_i_0_0',bool))) | $spl30),
  inference(cnf_transformation,[],[f24134_D])).
fof(f24134_D,plain,(
  p(s(bool,sK7('_i_0_0','_i_0_0',bool))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f24108,plain,(
  p(s(bool,sK7('_i_0_0','_i_0_0',bool))) | p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl16),
  inference(superposition,[],[f5066,f286])).
fof(f286,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X25] : ! [X27] : s(X25,X27) = s(X25,happ(s(fun(X25,X25),'_i_0_0'),s(X25,X27)))),
  file('thm.canonical.spolynom_|49||49||split|3','_i0_ax')).
fof(f19777,plain,(
  ( ! [X0] : (s(bool,'const.bool.ARB_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f19777_D])).
fof(f19777_D,plain,(
  ( ! [X0] : (s(bool,'const.bool.ARB_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f19779,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'T_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | $spl24),
  inference(cnf_transformation,[],[f19779_D])).
fof(f19779_D,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'T_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f19727,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'T_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'const.bool.ARB_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f312,f389])).
fof(f389,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f247,f247])).
fof(f247,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.canonical.spolynom_|49||49||split|3',aHL_BOOL_CASES)).
fof(f312,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',X0)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s(X0,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(X0,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun(X0,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(X0),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),sK8(X0)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',X0)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s(X0,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(X0,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun(X0,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(X0),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),sK8(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f128])).
fof(f128,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',X0)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s(X0,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),'_dst_x0x1_2'(s(fun(fun(X0,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(X0,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun(X0,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),happ(s(fun('type.pair.prod'('type.quote.index',X0),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',X0),happ(s(fun(X0,'type.pair.prod'('type.quote.index',X0)),happ(s(fun('type.quote.index',fun(X0,'type.pair.prod'('type.quote.index',X0))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X0,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(X0),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X0))),X1))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ? [X41] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X6),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',X6)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(X6,'type.pair.prod'('type.quote.index',X6))),'const.pair.,_0'),s(X6,'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),'_dst_x0x1_2'(s(fun(fun(X6,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(X6,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun(X6,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',X6),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,'type.pair.prod'('type.quote.index',X6)),happ(s(fun('type.quote.index',fun(X6,'type.pair.prod'('type.quote.index',X6))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),happ(s(fun('type.pair.prod'('type.quote.index',X6),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',X6),happ(s(fun(X6,'type.pair.prod'('type.quote.index',X6)),happ(s(fun('type.quote.index',fun(X6,'type.pair.prod'('type.quote.index',X6))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X6,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),happ(s(fun('type.pair.prod'('type.quote.index',X6),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',X6),happ(s(fun(X6,'type.pair.prod'('type.quote.index',X6)),happ(s(fun('type.quote.index',fun(X6,'type.pair.prod'('type.quote.index',X6))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(X6,'const.bool.ARB_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(X6),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',X6))),X41))))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.canonical.spolynom_TY_DEF')).
fof(f19774,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') | $spl18),
  inference(cnf_transformation,[],[f19774_D])).
fof(f19774_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ARB_0') <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f19780,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,X1))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f19780_D])).
fof(f19780_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,X1))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f19726,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,X2))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X2))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X2))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'T_0') = s(bool,'const.bool.ARB_0') | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f312,f389])).
fof(f19725,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,X1))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,X1))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.ARB_0')) )),
  inference(superposition,[],[f312,f389])).
fof(f19724,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'T_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'T_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'const.bool.ARB_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f312,f389])).
fof(f19776,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'F_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | $spl20),
  inference(cnf_transformation,[],[f19776_D])).
fof(f19776_D,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'F_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f19723,plain,(
  p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun('type.quote.index',bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.quote.index',bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.quote.index','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.quote.index','type.pair.prod'('type.quote.index',bool)),'_dst_x1x0_2'(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s(bool,'F_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),'_dst_x0x1_2'(s(fun(fun(bool,bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(bool,bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun(bool,'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11x11_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_dst_x00x11_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x111_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool),'?_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)))),'_dst_x0000x0111x1011_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x00x11_2'(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),happ(s(fun('type.pair.prod'('type.quote.index',bool),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.pair.prod'('type.quote.index',bool),happ(s(fun(bool,'type.pair.prod'('type.quote.index',bool)),happ(s(fun('type.quote.index',fun(bool,'type.pair.prod'('type.quote.index',bool))),'const.pair.,_0'),s('type.quote.index','const.bool.ARB_0'))),s(bool,'F_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)))),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))))),'const.ind_type.FCONS_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'const.ind_type.BOTTOM_0'))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool))),'_dst_x000x110x101_3'(s(fun(bool,fun(bool,bool)),'|slash||bslash|_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool)),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),'_i_0_0'))))))),s(fun('type.canonical.spolynom'(bool),'type.ind_type.recspace'('type.pair.prod'('type.quote.index',bool))),sK8(bool))))) | s(bool,'T_0') = s(bool,'const.bool.ARB_0')),
  inference(superposition,[],[f312,f247])).
fof(f5281,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5072,f5200])).
fof(f5200,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f335])).
fof(f5072,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f253])).
fof(f5280,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5074,f5200])).
fof(f5074,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f246,f5065,f313])).
fof(f313,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f5279,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5075,f5200])).
fof(f5075,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f5065,f313])).
fof(f5278,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5094,f246])).
fof(f5094,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f380,f5065,f313])).
fof(f5277,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5096,f5065])).
fof(f5096,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2105,f5065,f313])).
fof(f2105,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4))) | $spl5),
  inference(unit_resulting_resolution,[],[f1998,f248])).
fof(f1998,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4)))) | $spl5),
  inference(unit_resulting_resolution,[],[f374,f334])).
fof(f5276,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5098,f5065])).
fof(f5098,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2125,f5065,f313])).
fof(f5275,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5100,f5065])).
fof(f5100,plain,(
  p(s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1')))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3002,f5065,f313])).
fof(f3002,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f2952,f248])).
fof(f2952,plain,(
  ~p(s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1')))) | $spl5),
  inference(unit_resulting_resolution,[],[f2097,f317])).
fof(f317,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f234])).
fof(f234,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X47,X50] : (p(s(bool,happ(s(fun(fun(X47,bool),bool),'!_1'),s(fun(X47,bool),X50)))) <=> ! [X48] : p(s(bool,happ(s(fun(X47,bool),X50),s(X47,X48)))))),
  file('thm.canonical.spolynom_|49||49||split|3','!_HL_QUANT')).
fof(f2097,plain,(
  ~p(s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f1998,f317])).
fof(f5274,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5102,f5065])).
fof(f5102,plain,(
  p(s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1')))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3365,f5065,f313])).
fof(f3365,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'))) | $spl5),
  inference(unit_resulting_resolution,[],[f3001,f248])).
fof(f3001,plain,(
  ~p(s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1')))) | $spl5),
  inference(unit_resulting_resolution,[],[f2952,f317])).
fof(f5273,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5104,f5065])).
fof(f5104,plain,(
  p(s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3037,f5065,f313])).
fof(f3037,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))))) | $spl5),
  inference(unit_resulting_resolution,[],[f3000,f248])).
fof(f3000,plain,(
  ~p(s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f2952,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f5272,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5106,f5065])).
fof(f5106,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0)))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3767,f5065,f313])).
fof(f3767,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0))))) | $spl5),
  inference(unit_resulting_resolution,[],[f3035,f248])).
fof(f3035,plain,(
  ~p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0)))))) | $spl5),
  inference(unit_resulting_resolution,[],[f3000,f318])).
fof(f5271,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f5108,f5065])).
fof(f5108,plain,(
  ( ! [X0] : (p(s(bool,sK5(X0,'_i_0_0',bool)))) ) | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f5065,f313])).
fof(f1279,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK5(X0,'_i_0_0',bool))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1272,f248])).
fof(f5270,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f5110,f5065])).
fof(f5110,plain,(
  p(s(bool,sK9('_i_0_0',bool))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f2447,f5065,f313])).
fof(f2447,plain,(
  s(bool,'T_0') != s(bool,sK9('_i_0_0',bool)) | $spl13),
  inference(unit_resulting_resolution,[],[f2438,f248])).
fof(f5269,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5112,f5200])).
fof(f5112,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f246,f5065,f313])).
fof(f5268,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5113,f5200])).
fof(f5113,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f5065,f313])).
fof(f5267,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5132,f246])).
fof(f5132,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f380,f5065,f313])).
fof(f5266,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5134,f5065])).
fof(f5134,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2105,f5065,f313])).
fof(f5265,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5136,f5065])).
fof(f5136,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2125,f5065,f313])).
fof(f5264,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5138,f5065])).
fof(f5138,plain,(
  p(s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1')))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3002,f5065,f313])).
fof(f5263,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5140,f5065])).
fof(f5140,plain,(
  p(s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1')))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3365,f5065,f313])).
fof(f5262,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5142,f5065])).
fof(f5142,plain,(
  p(s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3037,f5065,f313])).
fof(f5261,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5144,f5065])).
fof(f5144,plain,(
  p(s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0)))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3767,f5065,f313])).
fof(f5260,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f5146,f5065])).
fof(f5146,plain,(
  ( ! [X0] : (p(s(bool,sK5(X0,'_i_0_0',bool)))) ) | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f5065,f313])).
fof(f5259,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f5148,f5065])).
fof(f5148,plain,(
  p(s(bool,sK9('_i_0_0',bool))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f2447,f5065,f313])).
fof(f5258,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5163,f380])).
fof(f5163,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f380,f5065,f439])).
fof(f439,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f434,f246])).
fof(f434,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f313,f247])).
fof(f5257,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5164,f2105])).
fof(f5164,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2105,f5065,f439])).
fof(f5256,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5165,f2125])).
fof(f5165,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2125,f5065,f439])).
fof(f5255,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5166,f3002])).
fof(f5166,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3002,f5065,f439])).
fof(f5254,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5167,f3365])).
fof(f5167,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3365,f5065,f439])).
fof(f5253,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5168,f3037])).
fof(f5168,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3037,f5065,f439])).
fof(f5252,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5169,f3767])).
fof(f5169,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3767,f5065,f439])).
fof(f5251,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f5170,f1279])).
fof(f5170,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5(X0,'_i_0_0',bool))) ) | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f5065,f439])).
fof(f5250,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f5171,f2447])).
fof(f5171,plain,(
  s(bool,'T_0') = s(bool,sK9('_i_0_0',bool)) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f2447,f5065,f439])).
fof(f5249,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5172,f5200])).
fof(f5172,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f380,f5065,f439])).
fof(f5248,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5174,f5200])).
fof(f5174,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK4)))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2105,f5065,f439])).
fof(f5247,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5176,f5200])).
fof(f5176,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),bool)),'=_0'),s('type.canonical.spolynom'(sK0),sK4))),s('type.canonical.spolynom'(sK0),sK2)))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2125,f5065,f439])).
fof(f5246,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5178,f5200])).
fof(f5178,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1'),s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1')))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3002,f5065,f439])).
fof(f5245,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5180,f5200])).
fof(f5180,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),bool),'!_1'),s(fun(fun(fun('type.canonical.spolynom'(sK0),bool),bool),bool),'!_1')))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3365,f5065,f439])).
fof(f5244,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5182,f5200])).
fof(f5182,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(fun('type.canonical.spolynom'(sK0),bool),bool),'!_1'),s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)))))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3037,f5065,f439])).
fof(f5243,plain,(
  $false | ($spl5 | $spl14)),
  inference(subsumption_resolution,[],[f5184,f5200])).
fof(f5184,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.canonical.spolynom'(sK0),bool),sK9('!_1',fun('type.canonical.spolynom'(sK0),bool))),s('type.canonical.spolynom'(sK0),sK9(sK9('!_1',fun('type.canonical.spolynom'(sK0),bool)),'type.canonical.spolynom'(sK0)))))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3767,f5065,f439])).
fof(f5242,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f5186,f5200])).
fof(f5186,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,sK5(X1,'_i_0_0',bool))) ) | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f5065,f439])).
fof(f5241,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f5188,f5200])).
fof(f5188,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK9('_i_0_0',bool))) ) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f2447,f5065,f439])).
fof(f5240,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f5191,f5200])).
fof(f5191,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK5(X1,X2,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f296])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK5(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK5(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X21,X32] : (! [X22] : (s(X6,X22) = s(X6,X32) => p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X22))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X21)))))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.JRH_INDUCT_UTIL')).
fof(f5201,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f286,f5065,f335])).
fof(f5198,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f244,f5065,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X48,X49] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X48))),s(bool,X49)))) <=> (p(s(bool,X48)) | p(s(bool,X49))))),
  file('thm.canonical.spolynom_|49||49||split|3','|bslash||slash|_HL_0')).
fof(f5196,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f244,f5065,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f5194,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f244,f5065,f1263])).
fof(f1263,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),'_dst_x0_1'(s(bool,X21))))))) )),
  inference(superposition,[],[f297,f324])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f5192,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f5065,f296])).
fof(f5189,plain,(
  $false | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f2447,f2447,f5065,f439])).
fof(f5187,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1279,f1279,f5065,f439])).
fof(f5185,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3767,f3767,f5065,f439])).
fof(f5183,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3037,f3037,f5065,f439])).
fof(f5181,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3365,f3365,f5065,f439])).
fof(f5179,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f3002,f3002,f5065,f439])).
fof(f5177,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2125,f2125,f5065,f439])).
fof(f5175,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f2105,f2105,f5065,f439])).
fof(f5173,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f380,f380,f5065,f439])).
fof(f5149,plain,(
  $false | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2447,f5065,f313])).
fof(f5147,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f1279,f5065,f313])).
fof(f5145,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3767,f5065,f313])).
fof(f5143,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3037,f5065,f313])).
fof(f5141,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3365,f5065,f313])).
fof(f5139,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3002,f5065,f313])).
fof(f5137,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2125,f5065,f313])).
fof(f5135,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2105,f5065,f313])).
fof(f5133,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f246,f380,f5065,f313])).
fof(f5131,plain,(
  $false | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2447,f5065,f313])).
fof(f5129,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f1279,f5065,f313])).
fof(f5127,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3767,f5065,f313])).
fof(f5125,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3037,f5065,f313])).
fof(f5123,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3365,f5065,f313])).
fof(f5121,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3002,f5065,f313])).
fof(f5119,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2125,f5065,f313])).
fof(f5117,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2105,f5065,f313])).
fof(f5115,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f380,f5065,f313])).
fof(f5111,plain,(
  $false | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2447,f5065,f313])).
fof(f5109,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f1279,f5065,f313])).
fof(f5107,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3767,f5065,f313])).
fof(f5105,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3037,f5065,f313])).
fof(f5103,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3365,f5065,f313])).
fof(f5101,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3002,f5065,f313])).
fof(f5099,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2125,f5065,f313])).
fof(f5097,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2105,f5065,f313])).
fof(f5095,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f246,f380,f5065,f313])).
fof(f5093,plain,(
  $false | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2447,f5065,f313])).
fof(f5091,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f1279,f5065,f313])).
fof(f5089,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3767,f5065,f313])).
fof(f5087,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3037,f5065,f313])).
fof(f5085,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3365,f5065,f313])).
fof(f5083,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f3002,f5065,f313])).
fof(f5081,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2125,f5065,f313])).
fof(f5079,plain,(
  $false | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f5065,f2105,f5065,f313])).
fof(f5077,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5065,f380,f5065,f313])).
fof(f5073,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f380,f5065,f253])).
fof(f5071,plain,(
  $false | ($spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1419,f5065,f248])).
fof(f1419,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl8),
  inference(unit_resulting_resolution,[],[f1375,f249])).
fof(f249,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f1375,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(unit_resulting_resolution,[],[f1279,f296])).
fof(f5069,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f244,f5065])).
fof(f5065,plain,(
  ( ! [X39] : (~p(s(bool,X39))) ) | $spl14),
  inference(cnf_transformation,[],[f5065_D])).
fof(f5065_D,plain,(
  ( ! [X39] : (~p(s(bool,X39))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5066,plain,(
  ( ! [X40] : (p(s(bool,happ(s(fun(bool,bool),X40),s(bool,sK7('_i_0_0',X40,bool))))) | p(s(bool,sK6('_i_0_0',X40,bool)))) ) | $spl16),
  inference(cnf_transformation,[],[f5066_D])).
fof(f5066_D,plain,(
  ( ! [X40] : (p(s(bool,happ(s(fun(bool,bool),X40),s(bool,sK7('_i_0_0',X40,bool))))) | p(s(bool,sK6('_i_0_0',X40,bool)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5064,plain,(
  ( ! [X39,X40] : (p(s(bool,sK6('_i_0_0',X40,bool))) | ~p(s(bool,X39)) | p(s(bool,happ(s(fun(bool,bool),X40),s(bool,sK7('_i_0_0',X40,bool)))))) )),
  inference(forward_demodulation,[],[f5045,f286])).
fof(f5045,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | p(s(bool,happ(s(fun(bool,bool),X40),s(bool,sK7('_i_0_0',X40,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0',X40,bool)))))) )),
  inference(superposition,[],[f298,f286])).
fof(f298,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK6(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X20,X21] : (! [X22] : (p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X22)))) => p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X22))))) => (? [X22] : p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X22)))) => ? [X22] : p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X22))))))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.bool.MONO_EXISTS')).
fof(f2436,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2436_D])).
fof(f2436_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2438,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | $spl13),
  inference(cnf_transformation,[],[f2438_D])).
fof(f2438_D,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f2424,plain,(
  ~p(s(bool,sK9('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f318,f286])).
fof(f1271,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1272,plain,(
  ( ! [X19] : (~p(s(bool,sK5(X19,'_i_0_0',bool)))) ) | $spl8),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  ( ! [X19] : (~p(s(bool,sK5(X19,'_i_0_0',bool)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1262,plain,(
  ( ! [X19] : (~p(s(bool,sK5(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f297,f286])).
fof(f378,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f377])).
fof(f377,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f376,f371])).
fof(f376,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK2))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f370,f373])).
fof(f369,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | $spl0),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f371,plain,(
  s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK3) | $spl2),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f241,plain,(
  s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK3) | s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4)))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | (s('type.canonical.spolynom'(sK0),sK1) = s('type.canonical.spolynom'(sK0),sK3) & s('type.canonical.spolynom'(sK0),sK2) = s('type.canonical.spolynom'(sK0),sK4))) & (s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK3) | s('type.canonical.spolynom'(sK0),sK2) != s('type.canonical.spolynom'(sK0),sK4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f198])).
fof(f198,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4))) & (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | s('type.canonical.spolynom'(X0),X1) != s('type.canonical.spolynom'(X0),X3) | s('type.canonical.spolynom'(X0),X2) != s('type.canonical.spolynom'(X0),X4)))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4))) & (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) | (s('type.canonical.spolynom'(X0),X1) != s('type.canonical.spolynom'(X0),X3) | s('type.canonical.spolynom'(X0),X2) != s('type.canonical.spolynom'(X0),X4))))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) <~> (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) <=> (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))) <=> (s('type.canonical.spolynom'(X0),X1) = s('type.canonical.spolynom'(X0),X3) & s('type.canonical.spolynom'(X0),X2) = s('type.canonical.spolynom'(X0),X4)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X7),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))) <=> (s('type.canonical.spolynom'(X6),X7) = s('type.canonical.spolynom'(X6),X9) & s('type.canonical.spolynom'(X6),X8) = s('type.canonical.spolynom'(X6),X10)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X7),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))) <=> (s('type.canonical.spolynom'(X6),X7) = s('type.canonical.spolynom'(X6),X9) & s('type.canonical.spolynom'(X6),X8) = s('type.canonical.spolynom'(X6),X10)))),
  file('thm.canonical.spolynom_|49||49||split|3','thm.canonical.spolynom_|49||49||split|3')).
fof(f373,plain,(
  s('type.canonical.spolynom'(sK0),sK2) = s('type.canonical.spolynom'(sK0),sK4) | $spl4),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  s('type.canonical.spolynom'(sK0),sK2) = s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f242,plain,(
  s('type.canonical.spolynom'(sK0),sK2) = s('type.canonical.spolynom'(sK0),sK4) | s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4)))),
  inference(cnf_transformation,[],[f199])).
fof(f370,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) | $spl1),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f372,plain,(
  s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK3) | $spl3),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f374,plain,(
  s('type.canonical.spolynom'(sK0),sK2) != s('type.canonical.spolynom'(sK0),sK4) | $spl5),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  s('type.canonical.spolynom'(sK0),sK2) != s('type.canonical.spolynom'(sK0),sK4) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f243,plain,(
  s('type.canonical.spolynom'(sK0),sK2) != s('type.canonical.spolynom'(sK0),sK4) | s('type.canonical.spolynom'(sK0),sK1) != s('type.canonical.spolynom'(sK0),sK3) | s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4)))),
  inference(cnf_transformation,[],[f199])).
% SZS output end Proof for thm.canonical
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 203
% Passive clauses: 1905
% Generated clauses: 30916
% Final active clauses: 154
% Final passive clauses: 1504
% Input formulas: 68
% Initial clauses: 122
% 
% Duplicate literals: 473
% Trivial inequalities: 80
% Fw subsumption resolutions: 755
% Fw demodulations: 7523
% Bw demodulations: 193
% Condensations: 582
% 
% Simple tautologies: 147
% Equational tautologies: 269
% Forward subsumptions: 17004
% Fw demodulations to eq. taut.: 2058
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 5453
% Factoring: 1
% Forward superposition: 7352
% Backward superposition: 8069
% Self superposition: 212
% Equality factoring: 56
% Equality resolution: 11
% 
% Split clauses: 16
% Split components: 33
% SAT solver clauses: 2043
% SAT solver unit clauses: 502
% SAT solver binary clauses: 834
% 
% Sat splits: 16
% Sat splitting refutations: 104
% 
% Memory used [KB]: 48357
% Time elapsed: 1.373 s
% ------------------------------
% ------------------------------
% Success in time 2.71 s
