Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (125159)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 207
% Passive clauses: 3799
% Generated clauses: 24612
% Final active clauses: 153
% Final passive clauses: 3247
% Input formulas: 132
% Initial clauses: 168
% 
% Duplicate literals: 882
% Trivial inequalities: 482
% Fw subsumption resolutions: 2314
% Fw demodulations: 3167
% Bw demodulations: 47
% 
% Simple tautologies: 402
% Equational tautologies: 54
% Forward subsumptions: 12517
% 
% Binary resolution: 591
% Factoring: 1
% Forward superposition: 8892
% Backward superposition: 7619
% Self superposition: 187
% Equality factoring: 106
% Equality resolution: 31
% 
% Split clauses: 451
% Split components: 917
% SAT solver clauses: 1007
% SAT solver unit clauses: 28
% SAT solver binary clauses: 220
% 
% Sat splits: 451
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 994
% TWLsolver calls for satisfiability: 206
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (125184)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 180
% Passive clauses: 1703
% Generated clauses: 22052
% Final active clauses: 127
% Final passive clauses: 934
% Input formulas: 132
% Initial clauses: 168
% 
% Duplicate literals: 760
% Trivial inequalities: 252
% Fw subsumption resolutions: 1197
% Fw demodulations: 2197
% Bw demodulations: 610
% Condensations: 1586
% Global subsumptions: 11
% 
% Simple tautologies: 252
% Equational tautologies: 42
% Forward subsumptions: 10789
% Bw demodulations to eq. taut.: 19
% 
% Binary resolution: 492
% Factoring: 1
% Forward superposition: 7275
% Backward superposition: 5464
% Self superposition: 156
% Equality factoring: 76
% Equality resolution: 18
% 
% Split clauses: 118
% Split components: 236
% SAT solver clauses: 2229
% SAT solver unit clauses: 281
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 22
% 
% Sat splits: 118
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 2185
% TWLsolver calls for satisfiability: 465
% 
% Memory used [KB]: 18677
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f78284,plain,(
  $false),
  inference(subsumption_resolution,[],[f78283,f428])).
fof(f428,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|7',aHL_TRUTH)).
fof(f78283,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f78281,f9111])).
fof(f9111,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f606,f2871,f1304])).
fof(f1304,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1303,f472])).
fof(f472,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X33] : ! [X35] : s(X33,X35) = s(X33,happ(s(fun(X33,X33),'_i_0_0'),s(X33,X35)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','_i0_ax')).
fof(f1303,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1279,f428])).
fof(f1279,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f533,f641])).
fof(f641,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f472,f436])).
fof(f436,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_fupdfupds_comp|split|7',aHL_BOOL_CASES)).
fof(f533,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X21,X22] : ((p(s(bool,X21)) => p(s(bool,X22))) => ((p(s(bool,X22)) => p(s(bool,X21))) => s(bool,X21) = s(bool,X22)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.bool.IMP_ANTISYM_AX')).
fof(f2871,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2786,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X7] : ! [X21,X22] : s(X7,X22) = s(X7,'const.bool.COND_3'(s(bool,'F_0'),s(X7,X21),s(X7,X22)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.bool.COND_CLAUSES|split|1')).
fof(f2786,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f484,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X43,X44,X45] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X44))),s(X43,X45)))) <=> s(X43,X44) = s(X43,X45))),
  file('thm.ring.ring_fupdfupds_comp|split|7',arity_0_HL_EQ)).
fof(f606,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f601,f438])).
fof(f438,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.bool.EQ_CLAUSES|split|0')).
fof(f601,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f431,f428])).
fof(f431,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f116])).
fof(f116,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.bool.NOT_CLAUSES|split|1')).
fof(f78281,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))))),
  inference(backward_demodulation,[],[f78243,f31876])).
fof(f31876,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))))),
  inference(unit_resulting_resolution,[],[f427,f566])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f426])).
fof(f427,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),'const.combin.o_2'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4),s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(sK1,fun(sK1,sK1)),fun(sK1,fun(sK1,sK1))),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f372])).
fof(f372,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),'const.combin.o_2'(s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X4),s(fun(fun(X1,fun(X1,X1)),fun(X1,fun(X1,X1))),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.ring.ring_fupdfupds_comp|split|7')).
fof(f78243,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(fun(X19,'type.ring.ring'(X16)),'const.combin.o_2'(s(fun('type.ring.ring'(X16),'type.ring.ring'(X16)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X16,fun(X16,X16)),fun(X16,fun(X16,X16))),X17))),s(fun(X19,'type.ring.ring'(X16)),'const.combin.o_2'(s(fun('type.ring.ring'(X16),'type.ring.ring'(X16)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X16,fun(X16,X16)),fun(X16,fun(X16,X16))),X18))),s(fun(X19,'type.ring.ring'(X16)),X20))))) = s(fun(X19,'type.ring.ring'(X16)),'const.combin.o_2'(s(fun('type.ring.ring'(X16),'type.ring.ring'(X16)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X16,fun(X16,X16)),fun(X16,fun(X16,X16))),'const.combin.o_2'(s(fun(fun(X16,fun(X16,X16)),fun(X16,fun(X16,X16))),X17),s(fun(fun(X16,fun(X16,X16)),fun(X16,fun(X16,X16))),X18))))),s(fun(X19,'type.ring.ring'(X16)),X20)))) )),
  inference(superposition,[],[f574,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.o_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),'const.combin.o_2'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10),s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9))))) = s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X10))),s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_RM_fupd_1'(s(fun(fun(X7,fun(X7,X7)),fun(X7,fun(X7,X7))),X9)))))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.ring.ring_fupdfupds_comp|split|6')).
fof(f574,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7,X16,X11] : ! [X10,X9,X8] : s(fun(X16,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X16,X7),'const.combin.o_2'(s(fun(X11,X7),X9),s(fun(X16,X11),X8))))) = s(fun(X16,X6),'const.combin.o_2'(s(fun(X11,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X11,X7),X9))),s(fun(X16,X11),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|7','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 371
% Passive clauses: 4744
% Generated clauses: 78737
% Final active clauses: 235
% Final passive clauses: 2576
% Input formulas: 132
% Initial clauses: 168
% Discarded non-redundant clauses: 2123
% 
% Duplicate literals: 1959
% Trivial inequalities: 544
% Fw subsumption resolutions: 2949
% Fw demodulations: 10870
% Bw demodulations: 1270
% Condensations: 3902
% Global subsumptions: 81
% 
% Simple tautologies: 530
% Equational tautologies: 399
% Forward subsumptions: 52183
% Fw demodulations to eq. taut.: 61
% Bw demodulations to eq. taut.: 40
% 
% Unit resulting resolution: 4233
% Factoring: 1
% Forward superposition: 27917
% Backward superposition: 22703
% Self superposition: 465
% Equality factoring: 142
% Equality resolution: 38
% 
% Split clauses: 45
% Split components: 90
% SAT solver clauses: 5089
% SAT solver unit clauses: 1458
% SAT solver binary clauses: 1800
% 
% Sat splits: 45
% Sat splitting refutations: 172
% 
% Memory used [KB]: 86864
% Time elapsed: 3.997 s
% ------------------------------
% ------------------------------
% Success in time 5.34 s
