Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.hreal
% (53103)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1383
% Passive clauses: 10618
% Generated clauses: 25280
% Final active clauses: 1347
% Final passive clauses: 8963
% Input formulas: 30
% Initial clauses: 93
% 
% Duplicate literals: 596
% Trivial inequalities: 153
% Fw subsumption resolutions: 2857
% Fw demodulations: 1035
% Bw demodulations: 35
% 
% Simple tautologies: 162
% Equational tautologies: 25
% Forward subsumptions: 9025
% 
% Binary resolution: 11852
% Factoring: 1163
% Forward superposition: 3488
% Backward superposition: 3716
% Self superposition: 26
% Equality factoring: 32
% Equality resolution: 10
% 
% Split clauses: 669
% Split components: 1343
% SAT solver clauses: 1078
% SAT solver unit clauses: 61
% SAT solver binary clauses: 374
% 
% Sat splits: 669
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1078
% TWLsolver calls for satisfiability: 409
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.hreal
% (53158)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 537
% Passive clauses: 3326
% Generated clauses: 16254
% Final active clauses: 477
% Final passive clauses: 2360
% Input formulas: 30
% Initial clauses: 93
% 
% Duplicate literals: 330
% Trivial inequalities: 73
% Fw subsumption resolutions: 1495
% Fw demodulations: 802
% Bw demodulations: 484
% Condensations: 334
% Global subsumptions: 5
% 
% Simple tautologies: 218
% Equational tautologies: 27
% Forward subsumptions: 5769
% 
% Binary resolution: 5119
% Factoring: 415
% Forward superposition: 2011
% Backward superposition: 1252
% Self superposition: 28
% Equality factoring: 30
% Equality resolution: 5
% 
% Split clauses: 128
% Split components: 261
% SAT solver clauses: 3889
% SAT solver unit clauses: 90
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 10
% SAT solver learnt literals: 10
% 
% Sat splits: 128
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 3869
% TWLsolver calls for satisfiability: 988
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.hreal
% (53214)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1514
% Passive clauses: 8029
% Generated clauses: 256539
% Final active clauses: 650
% Final passive clauses: 1621
% Input formulas: 30
% Initial clauses: 93
% 
% Duplicate literals: 7560
% Trivial inequalities: 487
% Fw subsumption resolutions: 11665
% Fw demodulations: 76730
% Bw demodulations: 4294
% Condensations: 11311
% Global subsumptions: 1021
% 
% Simple tautologies: 463
% Equational tautologies: 4862
% Forward subsumptions: 130991
% Fw demodulations to eq. taut.: 248
% Bw demodulations to eq. taut.: 203
% 
% Unit resulting resolution: 68614
% Forward superposition: 31449
% Backward superposition: 28580
% Self superposition: 396
% Equality factoring: 238
% Equality resolution: 13
% 
% Split clauses: 471
% Split components: 952
% SAT solver clauses: 12855
% SAT solver unit clauses: 2420
% SAT solver binary clauses: 1576
% 
% Sat splits: 471
% Sat splitting refutations: 2861
% 
% Memory used [KB]: 93004
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.hreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.hreal
% SZS output start Proof for thm.hreal
fof(f7431,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f274,f286,f296,f276,f306,f313,f321,f380,f378,f419,f513,f522,f531,f515,f2726,f2754,f298,f2744,f288,f2736,f278,f7424,f280,f279,f307,f7426])).
fof(f7426,plain,(
  $false | ($spl2 | $spl8 | $spl11)),
  inference(subsumption_resolution,[],[f7425,f203])).
fof(f203,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.hreal.ISACUT_HRAT',aHL_FALSITY)).
fof(f7425,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl8 | $spl11)),
  inference(forward_demodulation,[],[f7378,f2829])).
fof(f2829,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))) | $spl11),
  inference(resolution,[],[f298,f194])).
fof(f194,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.hreal.ISACUT_HRAT','thm.bool.EQ_CLAUSES|split|3')).
fof(f7378,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ($spl2 | $spl8)),
  inference(resolution,[],[f2757,f288])).
fof(f2757,plain,(
  ( ! [X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5))))) ) | $spl2),
  inference(resolution,[],[f278,f152])).
fof(f152,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))) => p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9,X12,X13] : ((p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X12)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',X13))))) => p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X13)))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.HRAT_LT_TRANS')).
fof(f307,plain,(
  ( ! [X6] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5))))) ) | $spl12),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  ( ! [X6] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f279,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5))))) ) | $spl4),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f280,plain,(
  ( ! [X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5))))) ) | $spl6),
  inference(cnf_transformation,[],[f280_D])).
fof(f280_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f7424,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7423,f219])).
fof(f219,plain,(
  ( ! [X0,X1] : (s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0))) = s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0))) = s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9,X5] : s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',X9))) = s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X5)))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.cut_of_hrat')).
fof(f7423,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X7)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7422,f217])).
fof(f217,plain,(
  ( ! [X0,X1] : (s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))) = s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X0))),s('type.hrat.hrat',X1)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))) = s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X0))),s('type.hrat.hrat',X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X19,X20] : s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',X19),s('type.hrat.hrat',X20))) = s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X19))),s('type.hrat.hrat',X20)))),
  file('thm.hreal.ISACUT_HRAT','ARITY_1_2_const.hreal.cut_of_hrat')).
fof(f7422,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7421,f219])).
fof(f7421,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X6)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7420,f217])).
fof(f7420,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7419,f219])).
fof(f7419,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7418,f217])).
fof(f7418,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(subsumption_resolution,[],[f7417,f203])).
fof(f7417,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'F_0')) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | ($spl2 | $spl11)),
  inference(forward_demodulation,[],[f7416,f2829])).
fof(f7416,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | $spl2),
  inference(subsumption_resolution,[],[f7377,f222])).
fof(f222,plain,(
  ~sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),
  inference(resolution,[],[f134,f216])).
fof(f216,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) | ~sP4(X0)) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((~p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) | sP4(X0)) & (~sP4(X0) | p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0))))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) <=> sP4(X0))),
  inference(definition_folding,[],[f89,f98])).
fof(f98,plain,(
  ! [X0] : (sP4(X0) <=> (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f89,plain,(
  ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) <=> (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) <=> (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : ((~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) <=> (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : ((p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3))))) => p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X0)))) <=> (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : ((p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3))))) => p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X14] : (p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),X14)))) <=> (? [X9] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X9)))) & ? [X9] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X9)))) & ! [X9,X12] : ((p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X9)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',X9))))) => p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X12))))) & ! [X9] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X9)))) => ? [X12] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X14),s('type.hrat.hrat',X12)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X12))))))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.isacut')).
fof(f134,plain,(
  ~p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ~p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f70])).
fof(f70,plain,(
  ? [X0] : ~p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X0))))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0] : p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.hreal.isacut_1'(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',X6))))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.ISACUT_HRAT')).
fof(f7377,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7))))) ) | $spl2),
  inference(resolution,[],[f2757,f213])).
fof(f213,plain,(
  ( ! [X6,X2,X0,X5] : (sP4(X0) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14(X0)),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13(X0)),s('type.hrat.hrat',sK12(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~sP4(X0) | (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK9(X0))))) & ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK10(X0))))) & ! [X9,X10] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X9)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X10),s('type.hrat.hrat',X9)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X10))))) & ! [X7] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X7)))) | (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK11(X7,X0))))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',sK11(X7,X0))))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) | ! [X5] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK12(X0))))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13(X0)),s('type.hrat.hrat',sK12(X0))))) & ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK13(X0)))))) | (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK14(X0))))) & ! [X2] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14(X0)),s('type.hrat.hrat',X2)))))) | sP4(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10,sK11,sK12,sK13,sK14])],[f130])).
fof(f130,plain,(
  ! [X0] : ((~sP4(X0) | (? [X12] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X12)))) & ? [X11] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X11)))) & ! [X9,X10] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X9)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X10),s('type.hrat.hrat',X9)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X10))))) & ! [X7] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X7)))) | ? [X8] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X8)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',X8)))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) | ! [X5] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ? [X3,X4] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) & ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) | ? [X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))))) | sP4(X0)))),
  inference(rectify,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~sP4(X0) | (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) | ! [X5] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ? [X3,X4] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) & ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) | ? [X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))))) | sP4(X0)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~sP4(X0) | (? [X6] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) & ? [X5] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) & ! [X3,X4] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) & ! [X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2)))))))) & ((! [X6] : ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6)))) | ! [X5] : p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ? [X3,X4] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X3)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',X3)))) & ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X4))))) | ? [X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X1)))) & ! [X2] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X2))))))) | sP4(X0)))),
  inference(nnf_transformation,[],[f98])).
fof(f278,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | $spl2),
  inference(cnf_transformation,[],[f278_D])).
fof(f278_D,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2736,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2735,f219])).
fof(f2735,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X1)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2734,f217])).
fof(f2734,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2733,f219])).
fof(f2733,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X0)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2732,f217])).
fof(f2732,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2731,f219])).
fof(f2731,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2730,f217])).
fof(f2730,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2729,f471])).
fof(f471,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))),s('type.hrat.hrat',sK5)))) | $spl0),
  inference(resolution,[],[f276,f186])).
fof(f186,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(X1,X0)),s('type.hrat.hrat',X1)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK6(X1,X0))))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(X1,X0)),s('type.hrat.hrat',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | ? [X2] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X1))))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) => ? [X2] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X2)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',X1))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9,X12] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X12)))) => ? [X13] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X13)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X13),s('type.hrat.hrat',X12))))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.HRAT_MEAN')).
fof(f2729,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2728,f219])).
fof(f2728,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f2727,f217])).
fof(f2727,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2684,f222])).
fof(f2684,plain,(
  ( ! [X0,X1] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(resolution,[],[f470,f212])).
fof(f212,plain,(
  ( ! [X6,X2,X0,X5] : (sP4(X0) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14(X0)),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK12(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f470,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | $spl0),
  inference(resolution,[],[f276,f185])).
fof(f185,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK6(X1,X0))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f288,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | $spl8),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2744,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X3),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f2743,f219])).
fof(f2743,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X3)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f2742,f217])).
fof(f2742,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2741,f219])).
fof(f2741,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2740,f217])).
fof(f2740,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2739,f471])).
fof(f2739,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2738,f219])).
fof(f2738,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f2737,f217])).
fof(f2737,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2685,f222])).
fof(f2685,plain,(
  ( ! [X2,X3] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(resolution,[],[f470,f213])).
fof(f298,plain,(
  ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | $spl11),
  inference(cnf_transformation,[],[f298_D])).
fof(f298_D,plain,(
  ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2754,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2753,f219])).
fof(f2753,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2752,f217])).
fof(f2752,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2751,f219])).
fof(f2751,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X4)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2750,f217])).
fof(f2750,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2749,f219])).
fof(f2749,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2748,f217])).
fof(f2748,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2747,f471])).
fof(f2747,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2746,f219])).
fof(f2746,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f2745,f217])).
fof(f2745,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2686,f222])).
fof(f2686,plain,(
  ( ! [X4,X5] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(resolution,[],[f470,f214])).
fof(f214,plain,(
  ( ! [X6,X2,X0,X5] : (sP4(X0) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14(X0)),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK13(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f2726,plain,(
  $false | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f2683,f471])).
fof(f2683,plain,(
  ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK6(sK5,sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))),s('type.hrat.hrat',sK5)))) | ($spl0 | $spl12)),
  inference(resolution,[],[f470,f307])).
fof(f515,plain,(
  ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | $spl15),
  inference(cnf_transformation,[],[f515_D])).
fof(f515_D,plain,(
  ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f531,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f530,f219])).
fof(f530,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f529,f217])).
fof(f529,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f528,f219])).
fof(f528,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X4)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f527,f217])).
fof(f527,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f526,f219])).
fof(f526,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f525,f217])).
fof(f525,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f524,f219])).
fof(f524,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(forward_demodulation,[],[f523,f217])).
fof(f523,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(subsumption_resolution,[],[f465,f222])).
fof(f465,plain,(
  ( ! [X4,X5] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) ) | $spl0),
  inference(resolution,[],[f276,f214])).
fof(f522,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X3),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f521,f219])).
fof(f521,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X3)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f520,f217])).
fof(f520,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f519,f219])).
fof(f519,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X2)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f518,f217])).
fof(f518,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f517,f219])).
fof(f517,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(forward_demodulation,[],[f516,f217])).
fof(f516,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f464,f222])).
fof(f464,plain,(
  ( ! [X2,X3] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) ) | $spl0),
  inference(resolution,[],[f276,f213])).
fof(f513,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f512,f219])).
fof(f512,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X1)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5))))) ) | $spl0),
  inference(forward_demodulation,[],[f511,f217])).
fof(f511,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f510,f219])).
fof(f510,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X0)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f509,f217])).
fof(f509,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f508,f219])).
fof(f508,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f507,f217])).
fof(f507,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f506,f219])).
fof(f506,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f505,f217])).
fof(f505,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f463,f222])).
fof(f463,plain,(
  ( ! [X0,X1] : (sP4('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) ) | $spl0),
  inference(resolution,[],[f276,f212])).
fof(f419,plain,(
  $false | $spl6),
  inference(resolution,[],[f280,f187])).
fof(f187,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK7(X0)),s('type.hrat.hrat',X0))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK7(X0)),s('type.hrat.hrat',X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f48])).
fof(f48,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ? [X12] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',X9))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.HRAT_DOWN')).
fof(f378,plain,(
  $false | $spl4),
  inference(resolution,[],[f324,f279])).
fof(f324,plain,(
  ( ! [X4] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X4))))) ) | $spl4),
  inference(resolution,[],[f279,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9,X12] : ~(p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X12)))) & p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',X9)))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.HRAT_LT_ANTISYM')).
fof(f380,plain,(
  $false | $spl4),
  inference(resolution,[],[f324,f188])).
fof(f188,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK8(X0)))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK8(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f49])).
fof(f49,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',X1))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ? [X12] : p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',X12))))),
  file('thm.hreal.ISACUT_HRAT','thm.hreal.HRAT_UP')).
fof(f321,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X14),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X13),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12))))) )),
  inference(forward_demodulation,[],[f320,f219])).
fof(f320,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X14)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X13),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12))))) )),
  inference(forward_demodulation,[],[f319,f217])).
fof(f319,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X13),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f318,f219])).
fof(f318,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X13)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f317,f217])).
fof(f317,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X13)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f316,f219])).
fof(f316,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X13)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f315,f217])).
fof(f315,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X12),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X13)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f314,f219])).
fof(f314,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X12)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X13)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(forward_demodulation,[],[f265,f217])).
fof(f265,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X12)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X13)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X14))))) )),
  inference(resolution,[],[f222,f214])).
fof(f313,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X11),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X10),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) )),
  inference(forward_demodulation,[],[f312,f219])).
fof(f312,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X11)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X10),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) )),
  inference(forward_demodulation,[],[f311,f217])).
fof(f311,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X10),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X11))))) )),
  inference(forward_demodulation,[],[f310,f219])).
fof(f310,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X10)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X11))))) )),
  inference(forward_demodulation,[],[f309,f217])).
fof(f309,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X9),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X10)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X11))))) )),
  inference(forward_demodulation,[],[f308,f219])).
fof(f308,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X9)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X10)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X11))))) )),
  inference(forward_demodulation,[],[f264,f217])).
fof(f264,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X9)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X9)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X10)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X11))))) )),
  inference(resolution,[],[f222,f213])).
fof(f306,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X8),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6))))) )),
  inference(forward_demodulation,[],[f305,f219])).
fof(f305,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X8)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6))))) )),
  inference(forward_demodulation,[],[f304,f217])).
fof(f304,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X7),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f303,f219])).
fof(f303,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X7)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f302,f217])).
fof(f302,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f301,f219])).
fof(f301,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f300,f217])).
fof(f300,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X6),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f299,f219])).
fof(f299,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X6)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(forward_demodulation,[],[f263,f217])).
fof(f263,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',X6)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X6)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X7)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X8))))) )),
  inference(resolution,[],[f222,f212])).
fof(f276,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | $spl0),
  inference(cnf_transformation,[],[f276_D])).
fof(f276_D,plain,(
  p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f296,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X5),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) )),
  inference(forward_demodulation,[],[f295,f219])).
fof(f295,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) )),
  inference(forward_demodulation,[],[f294,f217])).
fof(f294,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X4),s('type.hrat.hrat',sK5)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f293,f219])).
fof(f293,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X4)))) | ~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f292,f217])).
fof(f292,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f291,f219])).
fof(f291,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f290,f217])).
fof(f290,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f289,f219])).
fof(f289,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(forward_demodulation,[],[f262,f217])).
fof(f262,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X4)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X5))))) )),
  inference(resolution,[],[f222,f211])).
fof(f211,plain,(
  ( ! [X6,X0,X5] : (sP4(X0) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK14(X0))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK13(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f286,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X3),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) )),
  inference(forward_demodulation,[],[f285,f219])).
fof(f285,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X3)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))))))) )),
  inference(forward_demodulation,[],[f284,f217])).
fof(f284,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X2),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) )),
  inference(forward_demodulation,[],[f283,f219])).
fof(f283,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X2)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) )),
  inference(forward_demodulation,[],[f282,f217])).
fof(f282,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) )),
  inference(forward_demodulation,[],[f281,f219])).
fof(f281,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) )),
  inference(forward_demodulation,[],[f261,f217])).
fof(f261,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X2)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X3))))) )),
  inference(resolution,[],[f222,f210])).
fof(f210,plain,(
  ( ! [X6,X0,X5] : (sP4(X0) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK14(X0))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK13(X0)),s('type.hrat.hrat',sK12(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f274,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X1),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) )),
  inference(forward_demodulation,[],[f273,f219])).
fof(f273,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X1)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5))))) )),
  inference(forward_demodulation,[],[f272,f217])).
fof(f272,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',X0),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f271,f219])).
fof(f271,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',X0)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f270,f217])).
fof(f270,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f269,f219])).
fof(f269,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f268,f217])).
fof(f268,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hrat_lt_2'(s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5)))),s('type.hrat.hrat',sK5)))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f267,f219])).
fof(f267,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.cut_of_hrat_2'(s('type.hrat.hrat',sK5),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(forward_demodulation,[],[f260,f217])).
fof(f260,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK14('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',sK12('const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X0)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),'const.hreal.cut_of_hrat_1'(s('type.hrat.hrat',sK5))),s('type.hrat.hrat',X1))))) )),
  inference(resolution,[],[f222,f209])).
fof(f209,plain,(
  ( ! [X6,X0,X5] : (sP4(X0) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK14(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',sK12(X0))))) | p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X5)))) | ~p(s(bool,happ(s(fun('type.hrat.hrat',bool),X0),s('type.hrat.hrat',X6))))) )),
  inference(cnf_transformation,[],[f131])).
% SZS output end Proof for thm.hreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 393
% Passive clauses: 1512
% Generated clauses: 7243
% Final active clauses: 363
% Final passive clauses: 791
% Input formulas: 30
% Initial clauses: 86
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 402
% Trivial inequalities: 36
% Fw subsumption resolutions: 722
% Fw demodulations: 967
% Bw demodulations: 9
% 
% Simple tautologies: 747
% Forward subsumptions: 3063
% 
% Binary resolution: 4233
% Forward superposition: 226
% Backward superposition: 615
% 
% Split clauses: 13
% Split components: 48
% SAT solver clauses: 19
% SAT solver unit clauses: 3
% SAT solver binary clauses: 1
% 
% Sat splits: 13
% Sat splitting refutations: 5
% 
% Memory used [KB]: 3070
% Time elapsed: 0.088 s
% ------------------------------
% ------------------------------
% Success in time 11.81 s
