Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f6226,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6192,f1048,f1047,f146,f208,f5812,f5804,f203,f206])).
fof(f206,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f206_D])).
fof(f206_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f203,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f173])).
fof(f173,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X10] : (s(X6,X10) = s(X6,X10) <=> p(s(bool,'T_0')))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.bool.REFL_CLAUSE')).
fof(f5804,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X5)))) ) | $spl58),
  inference(cnf_transformation,[],[f5804_D])).
fof(f5804_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X5)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5812,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X17)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5797,f175])).
fof(f175,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X5] : (p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X5)))) <=> ! [X27] : s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X5),s(X6,X27))))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.bool.FORALL_DEF')).
fof(f5797,plain,(
  ( ! [X17] : (p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X17)))) ) | $spl2),
  inference(superposition,[],[f5739,f664])).
fof(f664,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(trivial_inequality_removal,[],[f625])).
fof(f625,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X33) = s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f212,f241])).
fof(f241,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f147,f147])).
fof(f147,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR',aHL_BOOL_CASES)).
fof(f212,plain,(
  s(bool,'T_0') != s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1)))),
  inference(resolution,[],[f148,f143])).
fof(f143,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4)))) & (~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X4)))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X4)))))) & p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) & ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f99])).
fof(f99,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) => (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) => p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) => (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) => p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) & (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))))) => (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) => p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) & (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))))) => (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) => p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))))))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.quotient.LEFT_RES_FORALL_REGULAR')).
fof(f148,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.bool.EQ_CLAUSES|split|0')).
fof(f5739,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X0))))) ) | $spl2),
  inference(resolution,[],[f5725,f141])).
fof(f141,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X4)))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X4))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f5725,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5724,f206])).
fof(f5724,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0))))) ) | $spl2),
  inference(forward_demodulation,[],[f5723,f230])).
fof(f230,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X1)))) )),
  inference(resolution,[],[f140,f149])).
fof(f149,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f140,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X4))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f5723,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5693,f206])).
fof(f5693,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X0)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,X0))))) )),
  inference(superposition,[],[f5460,f233])).
fof(f233,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),
  inference(resolution,[],[f142,f149])).
fof(f142,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))),
  inference(cnf_transformation,[],[f111])).
fof(f5460,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(forward_demodulation,[],[f188,f183])).
fof(f183,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X10] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))) = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.pred_set.SPECIFICATION')).
fof(f188,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK7(X2,X1,X0)))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7,X15] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X15)))) <=> ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X10))))))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.bool.RES_FORALL_THM')).
fof(f208,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f146,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR',aHL_FALSITY)).
fof(f1047,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK1,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f1047_D])).
fof(f1047_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK1,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1048,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK1,sK0)))) | $spl5),
  inference(subsumption_resolution,[],[f1036,f208])).
fof(f1036,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK1,sK0))))),
  inference(superposition,[],[f176,f224])).
fof(f224,plain,(
  s(bool,'F_0') = s(bool,'|dollar|forall_1'(s(fun(sK0,bool),sK1)))),
  inference(resolution,[],[f153,f143])).
fof(f153,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.quotient.LEFT_RES_FORALL_REGULAR','thm.bool.EQ_CLAUSES|split|3')).
fof(f176,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f6192,plain,(
  $false | ($spl13 | $spl58)),
  inference(trivial_inequality_removal,[],[f6132])).
fof(f6132,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl13 | $spl58)),
  inference(superposition,[],[f1047,f5804])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 90
% Passive clauses: 700
% Generated clauses: 5984
% Final active clauses: 76
% Final passive clauses: 533
% Input formulas: 37
% Initial clauses: 63
% 
% Duplicate literals: 315
% Trivial inequalities: 148
% Fw subsumption resolutions: 702
% Fw demodulations: 965
% Bw demodulations: 14
% 
% Simple tautologies: 136
% Equational tautologies: 23
% Forward subsumptions: 2843
% 
% Binary resolution: 74
% Forward superposition: 1844
% Backward superposition: 1686
% Self superposition: 83
% Equality factoring: 49
% Equality resolution: 12
% 
% Split clauses: 109
% Split components: 219
% SAT solver clauses: 230
% SAT solver unit clauses: 34
% SAT solver binary clauses: 96
% SAT solver learnt clauses: 1
% 
% Sat splits: 109
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 229
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 3454
% Time elapsed: 0.126 s
% ------------------------------
% ------------------------------
% Success in time 0.13 s
