Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (73157)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 152
% Passive clauses: 3696
% Generated clauses: 18570
% Final active clauses: 144
% Final passive clauses: 3412
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 540
% Trivial inequalities: 304
% Fw subsumption resolutions: 885
% Fw demodulations: 1047
% Bw demodulations: 4
% 
% Simple tautologies: 310
% Equational tautologies: 22
% Forward subsumptions: 11442
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 8452
% Backward superposition: 6863
% Self superposition: 205
% Equality factoring: 54
% Equality resolution: 24
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 426
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 219
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 426
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 19701
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (73183)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 1450
% Generated clauses: 16514
% Final active clauses: 122
% Final passive clauses: 1172
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 599
% Trivial inequalities: 203
% Fw subsumption resolutions: 564
% Fw demodulations: 1040
% Bw demodulations: 121
% Condensations: 903
% 
% Simple tautologies: 175
% Equational tautologies: 25
% Forward subsumptions: 9329
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6489
% Backward superposition: 4687
% Self superposition: 215
% Equality factoring: 74
% Equality resolution: 19
% 
% Split clauses: 108
% Split components: 216
% SAT solver clauses: 2181
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 108
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2181
% TWLsolver calls for satisfiability: 268
% 
% Memory used [KB]: 23411
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f120413,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f120289])).
fof(f120289,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f307,f120095])).
fof(f120095,plain,(
  ( ! [X2] : (s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X2),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f24835,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7] : ! [X12,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X7),X12),s(X7,X6),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|25','thm.canonical.canonical_sum_scalar_def|split|2')).
fof(f24835,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f24801,f418])).
fof(f418,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0'),s('type.num.num',X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0'),s('type.num.num',X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X45,X46] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X45),s('type.canonical.canonical_sum'('type.num.num'),X46))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0'),s('type.num.num',X45))),s('type.canonical.canonical_sum'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|25','ARITY_0_2_const.numRing.num_canonical_sum_scalar')).
fof(f24801,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0'),s('type.num.num',X0))),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(superposition,[],[f364,f315])).
fof(f315,plain,(
  s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0') = s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f71])).
fof(f71,axiom,(
  s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_scalar_0') = s(fun('type.num.num',fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|25','thm.numRing.num_canonical_sum_scalar_def')).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun(X0,fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X7] : ! [X45,X46,X47] : s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_scalar_3'(s('type.semi_ring.semi_ring'(X7),X45),s(X7,X46),s('type.canonical.canonical_sum'(X7),X47))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7)),happ(s(fun(X7,fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7))),'const.canonical.canonical_sum_scalar_1'(s('type.semi_ring.semi_ring'(X7),X45))),s(X7,X46))),s('type.canonical.canonical_sum'(X7),X47)))),
  file('thm.numRing.num_ring_thms|split|25','ARITY_1_3_const.canonical.canonical_sum_scalar')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_scalar_2'(s('type.num.num',X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|25','thm.numRing.num_ring_thms|split|25')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 537
% Passive clauses: 9515
% Generated clauses: 123428
% Final active clauses: 220
% Final passive clauses: 1907
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 3252
% 
% Duplicate literals: 3462
% Trivial inequalities: 1175
% Fw subsumption resolutions: 3142
% Fw demodulations: 13956
% Bw demodulations: 1815
% Condensations: 4747
% Global subsumptions: 207
% 
% Simple tautologies: 599
% Equational tautologies: 264
% Forward subsumptions: 85281
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2729
% Forward superposition: 45806
% Backward superposition: 40938
% Self superposition: 716
% Equality factoring: 216
% Equality resolution: 48
% 
% Split clauses: 338
% Split components: 681
% SAT solver clauses: 10634
% SAT solver unit clauses: 1204
% SAT solver binary clauses: 2732
% 
% Sat splits: 338
% Sat splitting refutations: 234
% 
% Memory used [KB]: 156841
% Time elapsed: 6.938 s
% ------------------------------
% ------------------------------
% Success in time 8.256 s
