Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.measure
% (87701)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 2913
% Generated clauses: 19082
% Final active clauses: 162
% Final passive clauses: 2499
% Input formulas: 30
% Initial clauses: 94
% 
% Duplicate literals: 761
% Trivial inequalities: 237
% Fw subsumption resolutions: 3376
% Fw demodulations: 2044
% Bw demodulations: 1
% 
% Simple tautologies: 356
% Equational tautologies: 50
% Forward subsumptions: 7319
% 
% Binary resolution: 1521
% Factoring: 4
% Forward superposition: 6003
% Backward superposition: 4545
% Self superposition: 163
% Equality factoring: 124
% Equality resolution: 23
% 
% Split clauses: 882
% Split components: 1831
% SAT solver clauses: 1775
% SAT solver unit clauses: 36
% SAT solver binary clauses: 253
% 
% Sat splits: 882
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1770
% TWLsolver calls for satisfiability: 282
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (87719)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 460
% Generated clauses: 10104
% Final active clauses: 98
% Final passive clauses: 214
% Input formulas: 30
% Initial clauses: 94
% 
% Duplicate literals: 845
% Trivial inequalities: 143
% Fw subsumption resolutions: 964
% Fw demodulations: 309
% Bw demodulations: 134
% Condensations: 1411
% Global subsumptions: 13
% 
% Simple tautologies: 234
% Equational tautologies: 59
% Forward subsumptions: 4369
% Bw demodulations to eq. taut.: 9
% 
% Binary resolution: 194
% Factoring: 1
% Forward superposition: 2232
% Backward superposition: 2680
% Self superposition: 326
% Equality factoring: 164
% Equality resolution: 15
% 
% Split clauses: 74
% Split components: 148
% SAT solver clauses: 799
% SAT solver unit clauses: 109
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 26
% 
% Sat splits: 74
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 747
% TWLsolver calls for satisfiability: 750
% 
% Memory used [KB]: 6780
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (87736)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 267
% Passive clauses: 2738
% Generated clauses: 121446
% Final active clauses: 53
% Final passive clauses: 172
% Input formulas: 30
% Initial clauses: 94
% Discarded non-redundant clauses: 853
% 
% Duplicate literals: 7023
% Trivial inequalities: 957
% Fw subsumption resolutions: 11505
% Fw demodulations: 9010
% Bw demodulations: 1423
% Condensations: 15649
% Global subsumptions: 492
% 
% Simple tautologies: 1198
% Equational tautologies: 876
% Forward subsumptions: 70530
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 29
% 
% Unit resulting resolution: 3345
% Factoring: 1
% Forward superposition: 21762
% Backward superposition: 24395
% Self superposition: 647
% Equality factoring: 450
% Equality resolution: 51
% 
% Split clauses: 262
% Split components: 543
% SAT solver clauses: 4340
% SAT solver unit clauses: 47
% SAT solver binary clauses: 474
% 
% Sat splits: 262
% Sat splitting refutations: 584
% 
% Memory used [KB]: 42728
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f1277,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f228,f229,f230,f231,f232,f234,f262,f235,f247,f260,f273,f274,f275,f276,f277,f279,f321,f280,f292,f320,f368,f413,f414,f415,f458,f457,f462,f464,f468,f469,f470,f471,f506,f505,f509,f463,f510,f461,f557,f564,f566,f567,f563,f569,f571,f575,f573,f576,f579,f536,f524,f523,f521,f520,f519,f518,f263,f572,f517,f585,f636,f625,f620,f619,f664,f665,f666,f667,f668,f669,f813,f561,f559,f767,f769,f770,f772,f774,f788,f816,f1002,f1003,f1176,f1178,f1179,f261,f1180,f265,f1183,f264,f1195,f266,f1212,f1268,f1276])).
fof(f1276,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f563,f134])).
fof(f134,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X3))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3,sK4])],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (~s('type.num.num',X3) = s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X9,X10] : (~s('type.num.num',X9) = s('type.num.num',X10) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X9))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X10)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X9,X10] : (~s('type.num.num',X9) = s('type.num.num',X10) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X9))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X10)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.measure.CLOSED_CDI_DISJOINT')).
fof(f1268,plain,(
  $false | ($spl31 | $spl33)),
  inference(subsumption_resolution,[],[f1218,f566])).
fof(f1218,plain,(
  s('type.num.num',sK5(sK4,sK3,sK2)) = s('type.num.num',sK6(sK4,sK3,sK2)) | $spl33),
  inference(resolution,[],[f569,f133])).
fof(f133,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X3))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X4)))))) | s('type.num.num',X3) = s('type.num.num',X4)) )),
  inference(cnf_transformation,[],[f87])).
fof(f1212,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) ) | $spl42),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f266,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl12),
  inference(cnf_transformation,[],[f266_D])).
fof(f266_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1195,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f328,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f157])).
fof(f157,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X14,X17,X15,X18] : ((s(bool,X14) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CONG')).
fof(f328,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))),
  inference(resolution,[],[f227,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X9),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X9))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | (s('type.num.num',sK5(X6,X1,X0)) != s('type.num.num',sK6(X6,X1,X0)) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK5(X6,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK6(X6,X1,X0)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK7(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK8(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK8(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK8(X1,X0)),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK8(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7,sK8])],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X9),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X9))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X7,X8] : (s('type.num.num',X7) != s('type.num.num',X8) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',X7))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(rectify,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f227,plain,(
  sP1(sK2,sK3)),
  inference(resolution,[],[f131,f151])).
fof(f151,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | ~p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | sP1(X0,X1)) & (~sP1(X0,X1) | p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> sP1(X0,X1))),
  inference(definition_folding,[],[f70,f85,f84])).
fof(f84,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) | ? [X6] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f70,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) | ? [X6] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5] : ((~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) | ? [X6] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) & ! [X6] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) & ! [X6] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (~s('type.num.num',X3) = s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & ! [X19] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X19),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(X6,bool),X19))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X8] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num','const.num.|48|_0'))) & ! [X10] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X10))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X8] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X9,X10] : (~s('type.num.num',X9) = s('type.num.num',X10) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X9))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X8),s('type.num.num',X10)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.measure.closed_cdi_def')).
fof(f131,plain,(
  p(s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))),
  inference(cnf_transformation,[],[f87])).
fof(f264,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl8),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1183,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f328,f160])).
fof(f160,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f265,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl10),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1180,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f328,f159])).
fof(f159,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f261,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f261_D])).
fof(f261_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1179,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f328,f158])).
fof(f158,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f1178,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f328,f156])).
fof(f156,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f1176,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK2,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f328,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f1003,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f990])).
fof(f990,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f262,f240])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))),
  inference(resolution,[],[f131,f172])).
fof(f172,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.EQ_CLAUSES|split|0')).
fof(f1002,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f262])).
fof(f816,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) ) | $spl40),
  inference(cnf_transformation,[],[f816_D])).
fof(f816_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f788,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f134,f225])).
fof(f774,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f134,f160])).
fof(f772,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f134,f159])).
fof(f770,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f134,f158])).
fof(f769,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f134,f156])).
fof(f767,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) )),
  inference(resolution,[],[f134,f154])).
fof(f559,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0'))) | $spl25),
  inference(cnf_transformation,[],[f559_D])).
fof(f559_D,plain,(
  s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f561,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | $spl27),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f813,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f812,f330])).
fof(f330,plain,(
  sP0(sK2,sK3)),
  inference(resolution,[],[f227,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f812,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0'))) | ~sP0(sK2,sK3)),
  inference(subsumption_resolution,[],[f764,f132])).
fof(f132,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))))),
  inference(cnf_transformation,[],[f87])).
fof(f764,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) | ~sP0(sK2,sK3)),
  inference(resolution,[],[f134,f146])).
fof(f146,plain,(
  ( ! [X4,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num',sK9(X4,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(X4,X1,X0))))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num',sK9(X4,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(X4,X1,X0))))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK10(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK10(X1,X0)),s('type.num.num','const.num.|48|_0'))) & ! [X3] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK10(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK10(X1,X0)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK10(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f93])).
fof(f93,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.|48|_0'))) | ? [X5] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num',X5))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & (? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.|48|_0'))) & ! [X3] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) | ? [X6] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & (? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.|48|_0'))) & ! [X6] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f84])).
fof(f669,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f618,f265])).
fof(f618,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f265,f220])).
fof(f220,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.NOT_CLAUSES|split|2')).
fof(f668,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f617,f265])).
fof(f617,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f265,f218])).
fof(f218,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.NOT_CLAUSES|split|1')).
fof(f667,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f614,f265])).
fof(f614,plain,(
  ( ! [X48] : (p(s(bool,X48))) ) | $spl10),
  inference(resolution,[],[f265,f198])).
fof(f198,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CLAUSES|split|0')).
fof(f666,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f612,f265])).
fof(f612,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f265,f195])).
fof(f195,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CLAUSES|split|2')).
fof(f665,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f610,f265])).
fof(f610,plain,(
  ( ! [X45] : (p(s(bool,X45))) ) | $spl10),
  inference(resolution,[],[f265,f192])).
fof(f192,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CLAUSES|split|3')).
fof(f664,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f609,f265])).
fof(f609,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | $spl10),
  inference(resolution,[],[f265,f189])).
fof(f189,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CLAUSES|split|1')).
fof(f619,plain,(
  $false | $spl10),
  inference(resolution,[],[f265,f222])).
fof(f222,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.CLOSED_CDI_DISJOINT',aHL_TRUTH)).
fof(f620,plain,(
  $false | $spl10),
  inference(resolution,[],[f265,f223])).
fof(f223,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.TRUTH')).
fof(f625,plain,(
  $false | $spl10),
  inference(resolution,[],[f265,f131])).
fof(f636,plain,(
  $false | $spl10),
  inference(resolution,[],[f265,f132])).
fof(f585,plain,(
  ( ! [X27] : (p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl38),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  ( ! [X27] : (p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f154])).
fof(f572,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl34),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f263,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f263_D])).
fof(f263_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f518,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f156])).
fof(f519,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f132,f156])).
fof(f520,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f158])).
fof(f521,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f159])).
fof(f523,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f160])).
fof(f524,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))))))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f132,f160])).
fof(f536,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) )),
  inference(resolution,[],[f132,f225])).
fof(f579,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f536,f264])).
fof(f576,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f521,f264])).
fof(f573,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f573_D])).
fof(f573_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f575,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f574,f264])).
fof(f574,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f520,f264])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f570,f264])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun(fun('type.num.num',fun(sK2,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f517,f264])).
fof(f569,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK5(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK6(sK4,sK3,sK2))))))) | $spl33),
  inference(cnf_transformation,[],[f569_D])).
fof(f569_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK5(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK6(sK4,sK3,sK2))))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f563,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | $spl28),
  inference(cnf_transformation,[],[f563_D])).
fof(f563_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f567,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK5(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK6(sK4,sK3,sK2)))))))),
  inference(subsumption_resolution,[],[f516,f227])).
fof(f516,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK5(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK6(sK4,sK3,sK2))))))) | ~sP1(sK2,sK3)),
  inference(resolution,[],[f132,f139])).
fof(f139,plain,(
  ( ! [X6,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK5(X6,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK6(X6,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f566,plain,(
  s('type.num.num',sK5(sK4,sK3,sK2)) != s('type.num.num',sK6(sK4,sK3,sK2)) | $spl31),
  inference(cnf_transformation,[],[f566_D])).
fof(f566_D,plain,(
  s('type.num.num',sK5(sK4,sK3,sK2)) != s('type.num.num',sK6(sK4,sK3,sK2)) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f564,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | s('type.num.num',sK5(sK4,sK3,sK2)) != s('type.num.num',sK6(sK4,sK3,sK2))),
  inference(subsumption_resolution,[],[f515,f227])).
fof(f515,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | s('type.num.num',sK5(sK4,sK3,sK2)) != s('type.num.num',sK6(sK4,sK3,sK2)) | ~sP1(sK2,sK3)),
  inference(resolution,[],[f132,f138])).
fof(f138,plain,(
  ( ! [X6,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | s('type.num.num',sK5(X6,X1,X0)) != s('type.num.num',sK6(X6,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f91])).
fof(f557,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f514,f330])).
fof(f514,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK2,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK2,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK2,bool)),sK4),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK2,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',sK9(sK4,sK3,sK2)))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK9(sK4,sK3,sK2))))))))) | s(fun(sK2,bool),'const.pred_set.EMPTY_0') != s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num','const.num.|48|_0'))) | ~sP0(sK2,sK3)),
  inference(resolution,[],[f132,f146])).
fof(f461,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13)) ) | $spl18),
  inference(cnf_transformation,[],[f461_D])).
fof(f461_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f510,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X22))) ) | $spl8),
  inference(subsumption_resolution,[],[f481,f264])).
fof(f481,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X22)) | s(bool,'T_0') = s(bool,X22)) ) | $spl8),
  inference(superposition,[],[f264,f175])).
fof(f175,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.CLOSED_CDI_DISJOINT',aHL_BOOL_CASES)).
fof(f463,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) | $spl20),
  inference(cnf_transformation,[],[f463_D])).
fof(f463_D,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f509,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X20))) ) | $spl8),
  inference(subsumption_resolution,[],[f480,f264])).
fof(f480,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl8),
  inference(superposition,[],[f264,f175])).
fof(f505,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X18))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f479])).
fof(f479,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X18)) | ~p(s(bool,X18))) ) | $spl8),
  inference(superposition,[],[f264,f174])).
fof(f174,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.EQ_CLAUSES|split|1')).
fof(f506,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f478])).
fof(f478,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16)) | ~p(s(bool,X16))) ) | $spl8),
  inference(superposition,[],[f264,f172])).
fof(f471,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) | $spl22),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f470,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl2),
  inference(subsumption_resolution,[],[f449,f167])).
fof(f167,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.EQ_CLAUSES|split|3')).
fof(f449,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,X22)) ) | $spl2),
  inference(superposition,[],[f261,f175])).
fof(f469,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(subsumption_resolution,[],[f445,f167])).
fof(f445,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X15)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f261,f170])).
fof(f170,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.EQ_CLAUSES|split|2')).
fof(f468,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12))) ) | $spl2),
  inference(subsumption_resolution,[],[f444,f167])).
fof(f444,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl2),
  inference(superposition,[],[f261,f168])).
fof(f168,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f99])).
fof(f464,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22))) ) | $spl2),
  inference(subsumption_resolution,[],[f433,f261])).
fof(f433,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | s(bool,'T_0') = s(bool,X22)) ) | $spl2),
  inference(superposition,[],[f261,f175])).
fof(f462,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X20))) ) | $spl2),
  inference(subsumption_resolution,[],[f432,f261])).
fof(f432,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl2),
  inference(superposition,[],[f261,f175])).
fof(f457,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f429])).
fof(f429,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f261,f170])).
fof(f458,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f428])).
fof(f428,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | p(s(bool,X12))) ) | $spl2),
  inference(superposition,[],[f261,f168])).
fof(f415,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f370,f266])).
fof(f370,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f266,f221])).
fof(f221,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f129])).
fof(f414,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f369,f266])).
fof(f369,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f266,f219])).
fof(f219,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f128])).
fof(f413,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f363,f266])).
fof(f363,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl12),
  inference(resolution,[],[f266,f187])).
fof(f187,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.measure.CLOSED_CDI_DISJOINT','thm.bool.IMP_CLAUSES|split|4')).
fof(f368,plain,(
  $false | $spl12),
  inference(resolution,[],[f266,f217])).
fof(f217,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.CLOSED_CDI_DISJOINT',aHL_FALSITY)).
fof(f320,plain,(
  ( ! [X4,X2,X1] : (s('type.num.num',X1) = s('type.num.num',X2) | s(bool,X4) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X1))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X2))))) | p(s(bool,X4))) ) | $spl14),
  inference(cnf_transformation,[],[f320_D])).
fof(f320_D,plain,(
  ( ! [X4,X2,X1] : (s('type.num.num',X1) = s('type.num.num',X2) | s(bool,X4) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X1))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X2))))) | p(s(bool,X4))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f292,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.num.num',X65) = s('type.num.num',X66) | p(s(bool,X67)) | p(s(bool,X68)) | s(bool,X68) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X65))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X66)))))) )),
  inference(resolution,[],[f133,f225])).
fof(f280,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.num.num',X36) = s('type.num.num',X37) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X36))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X37))))) | s(bool,X39) != s(bool,X40)) )),
  inference(resolution,[],[f133,f160])).
fof(f321,plain,(
  ( ! [X6,X10,X7] : (s('type.num.num',X6) = s('type.num.num',X7) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X6))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X7)))))) ) | $spl16),
  inference(cnf_transformation,[],[f321_D])).
fof(f321_D,plain,(
  ( ! [X6,X10,X7] : (s('type.num.num',X6) = s('type.num.num',X7) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X6))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X7)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f279,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.num.num',X31) = s('type.num.num',X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X34) | s(bool,X35) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X31))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X32)))))) )),
  inference(resolution,[],[f133,f160])).
fof(f277,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.num.num',X21) = s('type.num.num',X22) | p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X25) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X21))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X22)))))) )),
  inference(resolution,[],[f133,f159])).
fof(f276,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.num.num',X16) = s('type.num.num',X17) | p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X16))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X17)))))) )),
  inference(resolution,[],[f133,f158])).
fof(f275,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s('type.num.num',X11) = s('type.num.num',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X11))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X12))))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f133,f156])).
fof(f274,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s('type.num.num',X6) = s('type.num.num',X7) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X6))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X7)))))) )),
  inference(resolution,[],[f133,f156])).
fof(f273,plain,(
  ( ! [X4,X2,X5,X3,X1] : (s('type.num.num',X1) = s('type.num.num',X2) | p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X1))),s(fun(sK2,bool),happ(s(fun('type.num.num',fun(sK2,bool)),sK4),s('type.num.num',X2)))))) )),
  inference(resolution,[],[f133,f154])).
fof(f260,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) ) | $spl0),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f247,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f225])).
fof(f235,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f131,f160])).
fof(f262,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) ) | $spl4),
  inference(cnf_transformation,[],[f262_D])).
fof(f262_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f234,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f160])).
fof(f232,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f159])).
fof(f231,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f158])).
fof(f230,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f131,f156])).
fof(f229,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f156])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.closed_cdi_1'(s('type.pair.prod'(fun(sK2,bool),fun(fun(sK2,bool),bool)),sK3)))) )),
  inference(resolution,[],[f131,f154])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 119
% Generated clauses: 1069
% Final active clauses: 69
% Final passive clauses: 78
% Input formulas: 30
% Initial clauses: 94
% 
% Selected by SInE selection: 30
% SInE iterations: 1
% 
% Duplicate literals: 53
% Trivial inequalities: 8
% Fw subsumption resolutions: 182
% Fw demodulations: 39
% Bw demodulations: 2
% 
% Simple tautologies: 100
% Forward subsumptions: 459
% 
% Binary resolution: 260
% Forward superposition: 383
% Backward superposition: 78
% Equality resolution: 7
% 
% Split clauses: 55
% Split components: 117
% SAT solver clauses: 74
% SAT solver unit clauses: 17
% SAT solver binary clauses: 35
% 
% Sat splits: 55
% Sat splitting refutations: 18
% 
% Memory used [KB]: 1663
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
% Success in time 11.753 s
