Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (83006)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 152
% Passive clauses: 3762
% Generated clauses: 18957
% Final active clauses: 144
% Final passive clauses: 3478
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 540
% Trivial inequalities: 304
% Fw subsumption resolutions: 885
% Fw demodulations: 1082
% Bw demodulations: 4
% 
% Simple tautologies: 310
% Equational tautologies: 22
% Forward subsumptions: 11635
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 8557
% Backward superposition: 7081
% Self superposition: 208
% Equality factoring: 54
% Equality resolution: 24
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 426
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 219
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 426
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (83024)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 1347
% Generated clauses: 16353
% Final active clauses: 120
% Final passive clauses: 1071
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 585
% Trivial inequalities: 205
% Fw subsumption resolutions: 575
% Fw demodulations: 1031
% Bw demodulations: 121
% Condensations: 893
% 
% Simple tautologies: 177
% Equational tautologies: 25
% Forward subsumptions: 9358
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6267
% Backward superposition: 4830
% Self superposition: 215
% Equality factoring: 74
% Equality resolution: 18
% 
% Split clauses: 111
% Split components: 222
% SAT solver clauses: 2099
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 111
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2099
% TWLsolver calls for satisfiability: 266
% 
% Memory used [KB]: 22899
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f102436,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f102311])).
fof(f102311,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),sK0))) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),sK0)))),
  inference(superposition,[],[f307,f101628])).
fof(f101628,plain,(
  ( ! [X5] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X5))) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X5)))))) )),
  inference(forward_demodulation,[],[f101627,f20437])).
fof(f20437,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(forward_demodulation,[],[f20404,f327])).
fof(f327,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0'),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0'),s('type.canonical.spolynom'('type.num.num'),X0)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X45] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X45))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0'),s('type.canonical.spolynom'('type.num.num'),X45)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_0_1_const.numRing.num_spolynom_simplify')).
fof(f20404,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0'),s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(superposition,[],[f342,f312])).
fof(f312,plain,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f60])).
fof(f60,axiom,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_simplify_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|54','thm.numRing.num_spolynom_simplify_def')).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X7] : ! [X45,X46] : s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.spolynom'(X7),X46))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.spolynom'(X7),'type.canonical.canonical_sum'(X7)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.spolynom'(X7),X46)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_1_2_const.canonical.spolynom_simplify')).
fof(f101627,plain,(
  ( ! [X5] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X5))))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X5)))) )),
  inference(forward_demodulation,[],[f101492,f20234])).
fof(f20234,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(forward_demodulation,[],[f20201,f329])).
fof(f329,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X45] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X45))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X45)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_0_1_const.numRing.num_spolynom_normalize')).
fof(f20201,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0'),s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X0)))) )),
  inference(superposition,[],[f344,f311])).
fof(f311,plain,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f61])).
fof(f61,axiom,(
  s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_spolynom_normalize_0') = s(fun('type.canonical.spolynom'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|54','thm.numRing.num_spolynom_normalize_def')).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X7] : ! [X45,X46] : s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.spolynom'(X7),X46))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.spolynom'(X7),'type.canonical.canonical_sum'(X7)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.spolynom'(X7),X46)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f101492,plain,(
  ( ! [X5] : (s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X5))) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.spolynom'('type.num.num'),X5)))))) )),
  inference(superposition,[],[f19242,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X7] : ! [X11,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X7),X11),s('type.canonical.spolynom'(X7),X6))) = s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X7),X11),s('type.canonical.canonical_sum'(X7),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X7),X11),s('type.canonical.spolynom'(X7),X6)))))),
  file('thm.numRing.num_ring_thms|split|54','thm.canonical.spolynom_simplify_def')).
fof(f19242,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0)))) )),
  inference(forward_demodulation,[],[f19209,f328])).
fof(f328,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0'),s('type.canonical.canonical_sum'('type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0'),s('type.canonical.canonical_sum'('type.num.num'),X0)))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X45] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),X45))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0'),s('type.canonical.canonical_sum'('type.num.num'),X45)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_0_1_const.numRing.num_canonical_sum_simplify')).
fof(f19209,plain,(
  ( ! [X0] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0)))) )),
  inference(superposition,[],[f343,f310])).
fof(f310,plain,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f67])).
fof(f67,axiom,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.numRing.num_canonical_sum_simplify_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|54','thm.numRing.num_canonical_sum_simplify_def')).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X7] : ! [X45,X46] : s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.canonical_sum'(X7),X46))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.canonical_sum'(X7),X46)))),
  file('thm.numRing.num_ring_thms|split|54','ARITY_1_2_const.canonical.canonical_sum_simplify')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),sK0))) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),sK0)))))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),sK0))) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),sK0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X0))) != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0)))))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X0))) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X6))) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_simplify_1'(s('type.canonical.spolynom'('type.num.num'),X6))) = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_simplify_1'(s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_spolynom_normalize_1'(s('type.canonical.spolynom'('type.num.num'),X6)))))),
  file('thm.numRing.num_ring_thms|split|54','thm.numRing.num_ring_thms|split|54')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 492
% Passive clauses: 8217
% Generated clauses: 105512
% Final active clauses: 248
% Final passive clauses: 2014
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 3004
% 
% Duplicate literals: 3127
% Trivial inequalities: 1147
% Fw subsumption resolutions: 2861
% Fw demodulations: 12177
% Bw demodulations: 1803
% Condensations: 4366
% Global subsumptions: 180
% 
% Simple tautologies: 567
% Equational tautologies: 251
% Forward subsumptions: 71859
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2610
% Forward superposition: 40310
% Backward superposition: 31620
% Self superposition: 642
% Equality factoring: 202
% Equality resolution: 48
% 
% Split clauses: 255
% Split components: 515
% SAT solver clauses: 9117
% SAT solver unit clauses: 1002
% SAT solver binary clauses: 2229
% 
% Sat splits: 255
% Sat splitting refutations: 208
% 
% Memory used [KB]: 142257
% Time elapsed: 5.960 s
% ------------------------------
% ------------------------------
% Success in time 7.326 s
