Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (80303)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 151
% Passive clauses: 3131
% Generated clauses: 23269
% Final active clauses: 127
% Final passive clauses: 2635
% Input formulas: 26
% Initial clauses: 51
% 
% Duplicate literals: 892
% Trivial inequalities: 291
% Fw subsumption resolutions: 3566
% Fw demodulations: 3343
% Bw demodulations: 21
% 
% Simple tautologies: 464
% Equational tautologies: 48
% Forward subsumptions: 9536
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 928
% Factoring: 5
% Forward superposition: 7330
% Backward superposition: 6360
% Self superposition: 182
% Equality factoring: 116
% Equality resolution: 32
% 
% Split clauses: 555
% Split components: 1130
% SAT solver clauses: 1374
% SAT solver unit clauses: 34
% SAT solver binary clauses: 251
% 
% Sat splits: 555
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1369
% TWLsolver calls for satisfiability: 218
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (80325)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 1153
% Generated clauses: 10865
% Final active clauses: 101
% Final passive clauses: 1038
% Input formulas: 26
% Initial clauses: 51
% 
% Duplicate literals: 548
% Trivial inequalities: 96
% Fw subsumption resolutions: 1088
% Fw demodulations: 905
% Bw demodulations: 9
% Condensations: 741
% Global subsumptions: 5
% 
% Simple tautologies: 165
% Equational tautologies: 36
% Forward subsumptions: 4128
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 798
% Factoring: 5
% Forward superposition: 3123
% Backward superposition: 2048
% Self superposition: 88
% Equality factoring: 68
% Equality resolution: 12
% 
% Split clauses: 144
% Split components: 297
% SAT solver clauses: 1926
% SAT solver unit clauses: 49
% SAT solver binary clauses: 38
% SAT solver learnt clauses: 10
% 
% Sat splits: 144
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1906
% TWLsolver calls for satisfiability: 244
% 
% Memory used [KB]: 9466
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (80360)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 416
% Passive clauses: 7879
% Generated clauses: 140026
% Final active clauses: 126
% Final passive clauses: 850
% Input formulas: 26
% Initial clauses: 51
% Discarded non-redundant clauses: 384
% 
% Duplicate literals: 8070
% Trivial inequalities: 1204
% Fw subsumption resolutions: 13045
% Fw demodulations: 14951
% Bw demodulations: 5938
% Condensations: 18621
% Global subsumptions: 1672
% 
% Simple tautologies: 1598
% Equational tautologies: 1279
% Forward subsumptions: 69570
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 103
% 
% Unit resulting resolution: 6654
% Factoring: 1
% Forward superposition: 29519
% Backward superposition: 25890
% Self superposition: 758
% Equality factoring: 536
% Equality resolution: 98
% 
% Split clauses: 379
% Split components: 827
% SAT solver clauses: 12193
% SAT solver unit clauses: 56
% SAT solver binary clauses: 557
% 
% Sat splits: 379
% Sat splitting refutations: 1882
% 
% Memory used [KB]: 65627
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2380,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f198,f199,f210,f200,f201,f211,f202,f203,f204,f205,f208,f206,f215,f217,f207,f213,f219,f245,f246,f439,f441,f209,f471,f472,f601,f602,f613,f603,f604,f614,f605,f606,f607,f608,f611,f609,f618,f620,f610,f616,f622,f648,f649,f780,f781,f440,f612,f850,f851,f852,f2376])).
fof(f2376,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f2375,f141])).
fof(f141,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.TRUTH')).
fof(f2375,plain,(
  ~p(s(bool,'T_0')) | $spl0),
  inference(forward_demodulation,[],[f2347,f2209])).
fof(f2209,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1))) | $spl0),
  inference(resolution,[],[f2173,f109])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f2173,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1)))) | $spl0),
  inference(resolution,[],[f209,f112])).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X13,X10,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X13),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X13)))) & p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.pred_set.IN_INTER')).
fof(f2347,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1)))) | $spl0),
  inference(resolution,[],[f2172,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X7,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X13)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X13)))))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.pred_set.IN_COMPL')).
fof(f2172,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl0),
  inference(resolution,[],[f209,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f852,plain,(
  $false | ($spl6 | $spl18 | $spl22)),
  inference(trivial_inequality_removal,[],[f847])).
fof(f847,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl18 | $spl22)),
  inference(superposition,[],[f818,f212])).
fof(f212,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl6),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f818,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f817,f104])).
fof(f104,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f817,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f613,f615])).
fof(f615,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl22),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f851,plain,(
  $false | ($spl18 | $spl22)),
  inference(trivial_inequality_removal,[],[f849])).
fof(f849,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl18 | $spl22)),
  inference(superposition,[],[f818,f615])).
fof(f850,plain,(
  $false | ($spl18 | $spl22)),
  inference(equality_resolution,[],[f818])).
fof(f612,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl16),
  inference(cnf_transformation,[],[f612_D])).
fof(f612_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f440,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f781,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(subsumption_resolution,[],[f768,f102])).
fof(f102,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f768,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl22),
  inference(superposition,[],[f650,f94])).
fof(f94,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f650,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl22),
  inference(backward_demodulation,[],[f615,f358])).
fof(f358,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f148])).
fof(f148,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X1) | s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f93,f99])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X13,X10] : (s(fun(X6,bool),X10) = s(fun(X6,bool),X13) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X10))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X13))))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.pred_set.EXTENSION')).
fof(f93,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f63])).
fof(f63,plain,(
  ? [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.pred_set.COMPL_INTER|split|1')).
fof(f780,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(subsumption_resolution,[],[f767,f102])).
fof(f767,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(superposition,[],[f650,f94])).
fof(f649,plain,(
  $false | $spl23),
  inference(subsumption_resolution,[],[f642,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7] : ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f642,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(trivial_inequality_removal,[],[f636])).
fof(f636,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(superposition,[],[f616,f105])).
fof(f105,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f648,plain,(
  $false | $spl23),
  inference(subsumption_resolution,[],[f643,f142])).
fof(f643,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(trivial_inequality_removal,[],[f635])).
fof(f635,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(superposition,[],[f616,f103])).
fof(f103,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f75])).
fof(f622,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl28),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f616,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl23),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f610,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f358,f110])).
fof(f110,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.COMPL_INTER|split|1',aHL_BOOL_CASES)).
fof(f620,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl26),
  inference(cnf_transformation,[],[f620_D])).
fof(f620_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f618,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl25),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f609,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f358,f110])).
fof(f611,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl17),
  inference(cnf_transformation,[],[f611_D])).
fof(f611_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f608,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f358,f109])).
fof(f607,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f358,f107])).
fof(f107,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.COMPL_INTER|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f606,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f358,f105])).
fof(f605,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f358,f103])).
fof(f614,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl20),
  inference(cnf_transformation,[],[f614_D])).
fof(f614_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f604,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f358,f97])).
fof(f97,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f65])).
fof(f603,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f358,f97])).
fof(f613,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f613_D])).
fof(f613_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f602,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f358,f94])).
fof(f601,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f358,f94])).
fof(f472,plain,(
  $false | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f470])).
fof(f470,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6)),
  inference(superposition,[],[f417,f212])).
fof(f417,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f416,f104])).
fof(f416,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f210,f212])).
fof(f471,plain,(
  $false | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f417])).
fof(f209,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl0),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f441,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f427,f102])).
fof(f427,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f248,f94])).
fof(f248,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl6),
  inference(backward_demodulation,[],[f212,f189])).
fof(f189,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f147])).
fof(f147,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),X0,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f93,f99])).
fof(f439,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f426,f102])).
fof(f426,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(superposition,[],[f248,f94])).
fof(f246,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f239,f142])).
fof(f239,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(trivial_inequality_removal,[],[f233])).
fof(f233,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(superposition,[],[f213,f105])).
fof(f245,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f240,f142])).
fof(f240,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(trivial_inequality_removal,[],[f232])).
fof(f232,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(superposition,[],[f213,f103])).
fof(f219,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl12),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f213,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl7),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f207,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f189,f110])).
fof(f217,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl10),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f215,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl9),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f206,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f189,f110])).
fof(f208,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f205,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f189,f109])).
fof(f204,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f189,f107])).
fof(f203,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f189,f105])).
fof(f202,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f189,f103])).
fof(f211,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl4),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f201,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f189,f97])).
fof(f200,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f189,f97])).
fof(f210,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) | $spl2),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f199,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f189,f94])).
fof(f198,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f189,f94])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 95
% Passive clauses: 698
% Generated clauses: 2226
% Final active clauses: 80
% Final passive clauses: 543
% Input formulas: 26
% Initial clauses: 50
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 22
% Trivial inequalities: 75
% Fw subsumption resolutions: 212
% Fw demodulations: 315
% Bw demodulations: 11
% 
% Simple tautologies: 54
% Forward subsumptions: 862
% 
% Binary resolution: 116
% Forward superposition: 927
% Backward superposition: 521
% Equality resolution: 10
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 35
% SAT solver unit clauses: 5
% SAT solver binary clauses: 24
% 
% Sat splits: 24
% Sat splitting refutations: 10
% 
% Memory used [KB]: 1918
% Time elapsed: 0.059 s
% ------------------------------
% ------------------------------
% Success in time 11.811 s
