Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (60311)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 2853
% Generated clauses: 21000
% Final active clauses: 120
% Final passive clauses: 2440
% Input formulas: 28
% Initial clauses: 66
% 
% Duplicate literals: 826
% Trivial inequalities: 370
% Fw subsumption resolutions: 3295
% Fw demodulations: 3302
% Bw demodulations: 13
% 
% Simple tautologies: 446
% Equational tautologies: 50
% Forward subsumptions: 8570
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 554
% Factoring: 2
% Forward superposition: 6331
% Backward superposition: 5777
% Self superposition: 184
% Equality factoring: 119
% Equality resolution: 31
% 
% Split clauses: 455
% Split components: 927
% SAT solver clauses: 1136
% SAT solver unit clauses: 21
% SAT solver binary clauses: 188
% 
% Sat splits: 455
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1131
% TWLsolver calls for satisfiability: 159
% 
% Memory used [KB]: 11001
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (60331)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 194
% Passive clauses: 1146
% Generated clauses: 19675
% Final active clauses: 103
% Final passive clauses: 595
% Input formulas: 28
% Initial clauses: 66
% 
% Duplicate literals: 494
% Trivial inequalities: 94
% Fw subsumption resolutions: 2099
% Fw demodulations: 2673
% Bw demodulations: 389
% Condensations: 858
% Global subsumptions: 9
% 
% Simple tautologies: 151
% Equational tautologies: 42
% Forward subsumptions: 9631
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 670
% Factoring: 2
% Forward superposition: 5231
% Backward superposition: 4658
% Self superposition: 138
% Equality factoring: 62
% Equality resolution: 18
% 
% Split clauses: 212
% Split components: 427
% SAT solver clauses: 2260
% SAT solver unit clauses: 76
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 1
% 
% Sat splits: 212
% Sat splitting refutations: 36
% 
% TWLsolver clauses: 2223
% TWLsolver calls for satisfiability: 425
% 
% Memory used [KB]: 16630
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1533,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f115,f189,f187,f114,f188,f186,f1027,f1032,f1207,f1208,f1214,f1212,f1209,f1210,f1217,f1211,f1219,f1216,f1351,f1352,f1358,f1356,f1353,f1354,f1361,f1355,f1363,f1360,f1516,f1530,f1531])).
fof(f1531,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1518,f116])).
fof(f116,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SURJ_EMPTY|split|1',aHL_TRUTH)).
fof(f1518,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(superposition,[],[f1494,f1107])).
fof(f1107,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK4(sK3,'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f118,f1060,f249])).
fof(f249,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f243,f118])).
fof(f243,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f168,f119])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SURJ_EMPTY|split|1',aHL_BOOL_CASES)).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1060,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4(sK3,'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(forward_demodulation,[],[f1056,f203])).
fof(f203,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f165,f126])).
fof(f126,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f165,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7] : ! [X13] : ~p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f1056,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK4(sK3,'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK1,sK4(sK3,'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f189,f167])).
fof(f167,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7] : ! [X9,X12] : (s(fun(X7,bool),X9) = s(fun(X7,bool),X12) <=> ! [X13] : s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X12))))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.pred_set.EXTENSION')).
fof(f118,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SURJ_EMPTY|split|1',aHL_FALSITY)).
fof(f1494,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),sK3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1493,f118])).
fof(f1493,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),sK3))))) ) | $spl0),
  inference(forward_demodulation,[],[f1492,f203])).
fof(f1492,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK2),s(sK1,X1))),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1464,f116])).
fof(f1464,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),sK2),s(sK1,X1))),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(superposition,[],[f173,f1089])).
fof(f1089,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f186,f122])).
fof(f122,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f173,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK6(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK6(X8,X4,X3,X2,X1,X0)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,sK8(X4,X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8])],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7,X6] : ! [X8,X9,X12] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),X12)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X7,X13),s(fun(X7,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X7,X6),X8),s(X7,X13))),s(fun(X6,bool),X12))))) & ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X12)))) => ? [X14] : (p(s(bool,'const.bool.IN_2'(s(X7,X14),s(fun(X7,bool),X9)))) & s(X6,X13) = s(X6,happ(s(fun(X7,X6),X8),s(X7,X14)))))))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.pred_set.SURJ_DEF')).
fof(f1530,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1517,f116])).
fof(f1517,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(superposition,[],[f1494,f1067])).
fof(f1067,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK4('const.pred_set.EMPTY_0',sK3,sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f118,f1059,f249])).
fof(f1059,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.pred_set.EMPTY_0',sK3,sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(forward_demodulation,[],[f1058,f203])).
fof(f1058,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK4('const.pred_set.EMPTY_0',sK3,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.pred_set.EMPTY_0',sK3,sK1)),s(fun(sK1,bool),sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f189,f167])).
fof(f1516,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f1067,f1494,f121])).
fof(f121,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f1360,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | $spl18),
  inference(cnf_transformation,[],[f1360_D])).
fof(f1360_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1363,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | $spl23),
  inference(cnf_transformation,[],[f1363_D])).
fof(f1363_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1355,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | $spl3),
  inference(superposition,[],[f1057,f119])).
fof(f1057,plain,(
  s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) != s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | $spl3),
  inference(unit_resulting_resolution,[],[f189,f172])).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SURJ_EMPTY|split|1',aHL_EXT)).
fof(f1361,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(cnf_transformation,[],[f1361_D])).
fof(f1361_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1354,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1))))) ) | $spl3),
  inference(superposition,[],[f1057,f218])).
fof(f218,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f119,f119])).
fof(f1353,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f1057,f218])).
fof(f1356,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1358,plain,(
  s(bool,'T_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | $spl17),
  inference(cnf_transformation,[],[f1358_D])).
fof(f1358_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1352,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,X1) = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f1057,f218])).
fof(f1351,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,X0) = s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0',sK3,bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1057,f218])).
fof(f1216,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | $spl8),
  inference(cnf_transformation,[],[f1216_D])).
fof(f1216_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1219,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | $spl13),
  inference(cnf_transformation,[],[f1219_D])).
fof(f1219_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1211,plain,(
  s(bool,'F_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | $spl3),
  inference(superposition,[],[f1055,f119])).
fof(f1055,plain,(
  s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) != s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | $spl3),
  inference(unit_resulting_resolution,[],[f189,f172])).
fof(f1217,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f1217_D])).
fof(f1217_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1210,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1))))) ) | $spl3),
  inference(superposition,[],[f1055,f218])).
fof(f1209,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f1055,f218])).
fof(f1212,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1214,plain,(
  s(bool,'T_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | $spl7),
  inference(cnf_transformation,[],[f1214_D])).
fof(f1214_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1208,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,X1) = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f1055,f218])).
fof(f1207,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,X0) = s(bool,g(s(sK1,sK5(sK3,'const.pred_set.EMPTY_0',bool,sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1055,f218])).
fof(f1032,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1031,f118])).
fof(f1031,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1030,f203])).
fof(f1030,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1029,f118])).
fof(f1029,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1028,f203])).
fof(f1028,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f981,f118])).
fof(f981,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(superposition,[],[f176,f198])).
fof(f198,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f190,f126])).
fof(f190,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f187,f188])).
fof(f176,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK8(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f1027,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1026,f118])).
fof(f1026,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1025,f203])).
fof(f1025,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1024,f118])).
fof(f1024,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1023,f203])).
fof(f1023,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(condensation,[],[f1022])).
fof(f1022,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f980,f121])).
fof(f980,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK8('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK7('const.pred_set.EMPTY_0','const.pred_set.EMPTY_0',sK2,sK0,sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f176,f310])).
fof(f310,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f307,f118])).
fof(f307,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f249,f198])).
fof(f186,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl0),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f188,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | $spl2),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f114,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3)) & (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <~> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.SURJ_EMPTY|split|1','thm.pred_set.SURJ_EMPTY|split|1')).
fof(f187,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl1),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f189,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | $spl3),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f115,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f76])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 54
% Passive clauses: 142
% Generated clauses: 1385
% Final active clauses: 34
% Final passive clauses: 69
% Input formulas: 28
% Initial clauses: 66
% 
% Duplicate literals: 91
% Trivial inequalities: 27
% Fw subsumption resolutions: 96
% Fw demodulations: 124
% Bw demodulations: 12
% Condensations: 111
% 
% Simple tautologies: 54
% Equational tautologies: 55
% Forward subsumptions: 666
% Fw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 143
% Factoring: 1
% Forward superposition: 334
% Backward superposition: 311
% Self superposition: 25
% Equality factoring: 28
% Equality resolution: 3
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 173
% SAT solver unit clauses: 6
% SAT solver binary clauses: 41
% 
% Sat splits: 12
% Sat splitting refutations: 5
% 
% Memory used [KB]: 1791
% Time elapsed: 0.048 s
% ------------------------------
% ------------------------------
% Success in time 1.356 s
