Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (47810)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1392
% Passive clauses: 9415
% Generated clauses: 41354
% Final active clauses: 1267
% Final passive clauses: 7394
% Input formulas: 261
% Initial clauses: 541
% 
% Duplicate literals: 385
% Trivial inequalities: 82
% Fw subsumption resolutions: 2452
% Fw demodulations: 4817
% Bw demodulations: 119
% 
% Simple tautologies: 398
% Equational tautologies: 35
% Forward subsumptions: 23354
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 11455
% Factoring: 1727
% Forward superposition: 11075
% Backward superposition: 8493
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 592
% Split components: 1186
% SAT solver clauses: 881
% SAT solver unit clauses: 18
% SAT solver binary clauses: 94
% 
% Sat splits: 592
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 878
% TWLsolver calls for satisfiability: 602
% 
% Memory used [KB]: 16886
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (47830)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 569
% Passive clauses: 3232
% Generated clauses: 35397
% Final active clauses: 555
% Final passive clauses: 2498
% Input formulas: 261
% Initial clauses: 541
% 
% Duplicate literals: 94
% Trivial inequalities: 417
% Fw subsumption resolutions: 1072
% Fw demodulations: 11194
% Bw demodulations: 177
% Condensations: 60
% Global subsumptions: 25
% 
% Simple tautologies: 311
% Equational tautologies: 14
% Forward subsumptions: 14873
% Fw demodulations to eq. taut.: 36
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2433
% Factoring: 104
% Forward superposition: 8533
% Backward superposition: 7580
% Self superposition: 75
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 4638
% SAT solver unit clauses: 661
% SAT solver binary clauses: 157
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 14
% 
% Sat splits: 27
% 
% TWLsolver clauses: 4538
% TWLsolver calls for satisfiability: 3705
% 
% Memory used [KB]: 25969
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f31460,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f31459])).
fof(f31459,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15)))))))),
  inference(forward_demodulation,[],[f31457,f1559])).
fof(f1559,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f1138,f1091])).
fof(f1091,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f149])).
fof(f149,axiom,(
  ! [X15,X16] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X15),s('type.integer.int',X16))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16)))))),
  file('thm.integer.INT_ADD_REDUCE|split|5','thm.integer.int_sub')).
fof(f1138,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f468])).
fof(f468,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X7))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))),
  file('thm.integer.INT_ADD_REDUCE|split|5','thm.integer.INT_ADD_CALCULATE|split|3')).
fof(f31457,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))))),
  inference(unit_resulting_resolution,[],[f1521,f1176])).
fof(f1176,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f770])).
fof(f770,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f243])).
fof(f243,axiom,(
  ! [X15,X16] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X15))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))) <=> s('type.integer.int',X15) = s('type.integer.int',X16))),
  file('thm.integer.INT_ADD_REDUCE|split|5','thm.integer.INT_EQ_NEG')).
fof(f1521,plain,(
  s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))),
  inference(backward_demodulation,[],[f1091,f1460])).
fof(f1460,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))),
  inference(forward_demodulation,[],[f1459,f896])).
fof(f896,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X15] : s('type.num.num',X15) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))),
  file('thm.integer.INT_ADD_REDUCE|split|5','thm.arithmetic.NUMERAL_DEF')).
fof(f1459,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK15))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))),
  inference(forward_demodulation,[],[f1458,f896])).
fof(f1458,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK15))))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))),
  inference(forward_demodulation,[],[f1456,f896])).
fof(f1456,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK15))))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))),
  inference(backward_demodulation,[],[f896,f877])).
fof(f877,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK15))))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))))),
  inference(cnf_transformation,[],[f670])).
fof(f670,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK14))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK15))))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK15))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK15),s('type.num.num',sK14)))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f584])).
fof(f584,plain,(
  ? [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))))) != s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))))))) = s('type.integer.int','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))))),
  file('thm.integer.INT_ADD_REDUCE|split|5','thm.integer.INT_ADD_REDUCE|split|5')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 734
% Passive clauses: 4828
% Generated clauses: 31478
% Final active clauses: 645
% Final passive clauses: 2673
% Input formulas: 261
% Initial clauses: 541
% 
% Duplicate literals: 11
% Trivial inequalities: 45
% Fw subsumption resolutions: 755
% Fw demodulations: 6076
% Bw demodulations: 1015
% Condensations: 27
% 
% Simple tautologies: 130
% Equational tautologies: 64
% Forward subsumptions: 18425
% Fw demodulations to eq. taut.: 316
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 8266
% Factoring: 1
% Forward superposition: 8924
% Backward superposition: 4776
% Self superposition: 92
% Equality factoring: 2
% Equality resolution: 18
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 5638
% SAT solver unit clauses: 2625
% SAT solver binary clauses: 2588
% 
% Sat splits: 11
% Sat splitting refutations: 786
% 
% Memory used [KB]: 14072
% Time elapsed: 0.631 s
% ------------------------------
% ------------------------------
% Success in time 1.944 s
