Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (68041)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 122
% Passive clauses: 1828
% Generated clauses: 15279
% Final active clauses: 99
% Final passive clauses: 1599
% Input formulas: 36
% Initial clauses: 39
% 
% Duplicate literals: 666
% Trivial inequalities: 119
% Fw subsumption resolutions: 911
% Fw demodulations: 1331
% Bw demodulations: 21
% 
% Simple tautologies: 224
% Equational tautologies: 23
% Forward subsumptions: 9719
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 78
% Forward superposition: 6921
% Backward superposition: 4747
% Self superposition: 251
% Equality factoring: 72
% Equality resolution: 19
% 
% Split clauses: 287
% Split components: 574
% SAT solver clauses: 461
% SAT solver unit clauses: 3
% SAT solver binary clauses: 71
% 
% Sat splits: 287
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 461
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 20980
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (68074)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 86
% Passive clauses: 1017
% Generated clauses: 11590
% Final active clauses: 84
% Final passive clauses: 899
% Input formulas: 36
% Initial clauses: 39
% 
% Duplicate literals: 504
% Trivial inequalities: 106
% Fw subsumption resolutions: 376
% Fw demodulations: 918
% Bw demodulations: 31
% Condensations: 923
% 
% Simple tautologies: 109
% Equational tautologies: 21
% Forward subsumptions: 6235
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 58
% Forward superposition: 3761
% Backward superposition: 3574
% Self superposition: 213
% Equality factoring: 58
% Equality resolution: 10
% 
% Split clauses: 60
% Split components: 120
% SAT solver clauses: 1415
% SAT solver unit clauses: 66
% SAT solver binary clauses: 50
% 
% Sat splits: 60
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1415
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 23922
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f3887,plain,(
  $false),
  inference(subsumption_resolution,[],[f3886,f113])).
fof(f113,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|3',aHL_TRUTH)).
fof(f3886,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f3884,f639])).
fof(f639,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(equality_resolution,[],[f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(condensation,[],[f431])).
fof(f431,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f428,f154])).
fof(f154,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f114,f115])).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f114,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f428,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f135,f152])).
fof(f152,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f115,f115])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X32,X33,X34] : (p(s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),'=_0'),s(X32,X33))),s(X32,X34)))) <=> s(X32,X33) = s(X32,X34))),
  file('thm.ring.ring_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f3884,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f3842,f1664])).
fof(f1664,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f112,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f111])).
fof(f112,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f102])).
fof(f102,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  file('thm.ring.ring_fupdfupds_comp|split|3','thm.ring.ring_fupdfupds_comp|split|3')).
fof(f3842,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X2))))),s(fun(X3,'type.ring.ring'(X0)),X4))) = s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s(fun(X3,'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun(X3,'type.ring.ring'(X0)),X4)))))) )),
  inference(superposition,[],[f139,f118])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1)))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))) = s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X9)))))),
  file('thm.ring.ring_fupdfupds_comp|split|3','thm.ring.ring_fupdfupds_comp|split|2')).
fof(f139,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7,X13,X14] : ! [X10,X9,X8] : s(fun(X13,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X13,X7),'const.combin.o_2'(s(fun(X14,X7),X9),s(fun(X13,X14),X8))))) = s(fun(X13,X6),'const.combin.o_2'(s(fun(X14,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X14,X7),X9))),s(fun(X13,X14),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 62
% Passive clauses: 469
% Generated clauses: 3827
% Final active clauses: 41
% Final passive clauses: 241
% Input formulas: 36
% Initial clauses: 39
% 
% Duplicate literals: 202
% Trivial inequalities: 18
% Fw subsumption resolutions: 59
% Fw demodulations: 664
% Bw demodulations: 175
% Condensations: 207
% 
% Simple tautologies: 57
% Equational tautologies: 57
% Forward subsumptions: 2032
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 256
% Forward superposition: 1096
% Backward superposition: 874
% Self superposition: 69
% Equality factoring: 32
% Equality resolution: 1
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 500
% SAT solver unit clauses: 116
% SAT solver binary clauses: 150
% 
% Sat splits: 8
% Sat splitting refutations: 13
% 
% Memory used [KB]: 5245
% Time elapsed: 0.165 s
% ------------------------------
% ------------------------------
% Success in time 1.481 s
