Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (113506)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 804
% Passive clauses: 12140
% Generated clauses: 26998
% Final active clauses: 787
% Final passive clauses: 10889
% Input formulas: 69
% Initial clauses: 262
% 
% Duplicate literals: 435
% Trivial inequalities: 14
% Fw subsumption resolutions: 1837
% Fw demodulations: 1761
% Bw demodulations: 9
% 
% Simple tautologies: 293
% Equational tautologies: 61
% Forward subsumptions: 9987
% 
% Binary resolution: 12330
% Factoring: 1668
% Forward superposition: 4079
% Backward superposition: 4320
% Self superposition: 44
% Equality factoring: 26
% Equality resolution: 14
% 
% Split clauses: 388
% Split components: 820
% SAT solver clauses: 661
% SAT solver unit clauses: 24
% SAT solver binary clauses: 116
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 388
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 659
% TWLsolver calls for satisfiability: 404
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (113523)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 488
% Passive clauses: 3146
% Generated clauses: 14654
% Final active clauses: 463
% Final passive clauses: 2576
% Input formulas: 69
% Initial clauses: 262
% 
% Duplicate literals: 223
% Trivial inequalities: 26
% Fw subsumption resolutions: 856
% Fw demodulations: 1580
% Bw demodulations: 95
% Condensations: 319
% Global subsumptions: 16
% 
% Simple tautologies: 294
% Equational tautologies: 22
% Forward subsumptions: 4597
% 
% Binary resolution: 2917
% Factoring: 102
% Forward superposition: 2616
% Backward superposition: 2343
% Self superposition: 42
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 29
% Split components: 60
% SAT solver clauses: 3787
% SAT solver unit clauses: 265
% SAT solver binary clauses: 87
% SAT solver learnt clauses: 32
% SAT solver learnt literals: 12
% 
% Sat splits: 29
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 3723
% TWLsolver calls for satisfiability: 2422
% 
% Memory used [KB]: 19189
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (113551)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1616
% Passive clauses: 14140
% Generated clauses: 413275
% Final active clauses: 1315
% Final passive clauses: 3368
% Input formulas: 69
% Initial clauses: 262
% Discarded non-redundant clauses: 69982
% 
% Duplicate literals: 520
% Trivial inequalities: 77
% Fw subsumption resolutions: 3596
% Fw demodulations: 198418
% Bw demodulations: 2392
% Condensations: 980
% Global subsumptions: 142
% 
% Simple tautologies: 129
% Equational tautologies: 77
% Forward subsumptions: 72146
% Fw demodulations to eq. taut.: 57029
% Bw demodulations to eq. taut.: 82
% 
% Unit resulting resolution: 176003
% Factoring: 1
% Forward superposition: 21577
% Backward superposition: 9882
% Self superposition: 58
% Equality factoring: 28
% Equality resolution: 33
% 
% Split clauses: 181
% Split components: 369
% SAT solver clauses: 16258
% SAT solver unit clauses: 1877
% SAT solver binary clauses: 7610
% 
% Sat splits: 181
% Sat splitting refutations: 1612
% 
% Memory used [KB]: 132663
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f37842,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f584,f586,f587,f588,f591,f688,f716,f733,f735,f736,f737,f740,f837,f865,f994,f989,f1097,f1166,f1165,f1164,f1173,f1175,f1251,f1250,f1256,f1257,f1414,f1413,f1392,f1383,f1382,f1381,f1503,f1504,f1505,f1506,f1567,f1582,f1720,f1583,f1722,f1718,f1584,f1586,f1588,f1590,f1591,f1690,f1723,f1782,f1854,f1899,f2224,f2501,f2349,f2351,f2353,f2355,f2356,f2455,f2457,f2465,f2546,f2845,f2847,f2848,f2849,f2852,f2949,f2953,f2998,f3000,f3001,f3002,f3005,f3102,f3147,f3148,f3150,f3152,f3154,f3155,f3254,f3310,f3319,f3292,f3320,f3322,f3324,f3326,f3327,f3426,f3293,f3481,f3485,f3497,f3501,f3513,f3518,f3530,f3534,f3546,f3550,f3553,f3555,f3557,f3559,f3560,f3659,f3709,f3714,f3725,f3731,f3742,f3748,f3759,f3765,f3766,f3768,f3770,f3772,f3773,f3872,f3880,f3942,f3966,f3982,f3999,f4015,f4030,f4054,f4071,f4088,f4106,f4111,f4123,f4128,f4142,f4147,f4161,f4166,f4174,f4483,f4673,f4701,f4750,f4751,f5453,f5459,f5454,f5480,f5486,f5481,f5639,f5618,f5616,f5608,f5607,f5606,f5730,f5731,f5732,f5733,f5748,f5755,f5771,f5781,f5796,f5807,f5820,f5840,f5862,f5884,f5910,f6422,f6475,f6500,f6542,f6552,f6547,f6558,f6568,f6563,f6583,f6597,f6611,f6622,f6636,f6646,f6662,f6668,f6682,f6692,f6706,f6716,f6755,f6762,f6776,f6783,f6798,f6815,f6830,f6835,f6850,f6855,f6870,f6887,f7376,f7383,f7396,f7404,f7419,f7426,f7439,f7447,f7463,f7494,f7520,f7549,f7586,f7610,f7651,f7679,f7707,f7745,f7766,f7804,f7841,f7861,f7900,f7920,f8693,f8762,f8803,f8806,f8805,f8807,f8808,f9158,f9207,f9210,f9211,f6445,f5460,f6441,f6440,f6438,f6194,f6179,f6175,f4821,f4822,f4823,f9248,f1171,f9249,f9250,f9252,f9209,f9287,f9289,f9291,f9293,f9294,f9393,f9395,f2515,f9402,f9456,f9458,f9459,f9460,f9463,f9560,f9666,f9710,f1858,f1856,f9845,f9859,f9868,f9870,f9871,f9872,f9875,f9972,f10026,f10054,f10070,f10086,f10102,f10118,f10148,f10165,f10182,f10199,f10236,f10253,f10270,f10287,f10316,f10332,f10341,f10348,f10364,f10374,f10382,f10578,f10587,f10611,f10624,f10648,f10669,f10692,f10717,f10733,f4482,f10788,f11272,f11297,f11299,f11301,f11303,f11304,f11403,f11411,f11463,f11490,f11513,f11527,f12031,f12083,f12110,f12145,f6553,f12159,f6569,f12179,f12204,f12229,f12249,f12278,f12305,f12331,f12353,f12383,f12405,f12453,f12478,f12518,f12534,f12586,f12628,f12661,f12698,f7558,f12721,f12756,f12787,f12807,f12828,f12885,f7750,f12912,f7771,f12935,f12969,f12997,f13046,f13074,f13297,f8754,f13412,f13416,f13427,f13431,f13470,f13474,f13485,f13490,f13500,f13505,f13515,f13520,f13531,f13537,f13567,f13573,f13578,f13591,f13584,f13592,f13604,f13607,f13608,f13667,f13745,f13746,f13747,f13757,f2505,f2503,f13896,f13898,f13762,f13764,f13765,f13766,f13769,f13866,f13899,f13906,f14049,f13907,f13909,f13911,f13913,f13914,f14013,f14051,f14076,f14078,f14079,f14080,f14083,f14180,f14230,f14247,f14290,f14074,f14050,f2508,f14294,f14296,f14297,f14298,f14301,f14398,f14502,f14506,f14508,f14510,f14512,f14513,f14612,f3690,f14651,f14653,f14655,f14657,f14658,f14757,f14922,f14935,f15043,f15044,f15045,f15065,f15067,f15068,f15069,f15072,f15169,f14217,f15217,f14242,f15249,f14257,f15264,f15273,f15279,f15288,f15323,f15332,f15338,f15348,f15353,f15363,f15368,f15378,f15384,f15395,f15423,f15434,f15439,f15058,f15440,f15451,f15458,f15466,f15522,f13752,f15523,f13755,f15524,f13754,f15582,f8756,f15643,f15781,f15843,f15883,f15173,f15891,f15893,f15894,f15895,f15898,f15995,f16033,f16196,f5487,f16875,f17363,f15842,f17421,f17457,f15847,f15845,f17465,f17467,f17469,f17471,f17472,f17571,f14792,f17621,f17630,f17646,f17635,f10241,f17631,f17573,f14241,f17724,f17725,f17726,f17743,f17780,f17825,f17826,f17827,f17828,f17829,f17831,f17833,f17835,f17836,f17935,f18784,f15059,f18785,f15062,f18786,f15061,f19656,f15841,f15840,f19662,f15844,f19667,f19703,f19836,f19897,f19898,f19938,f19902,f19900,f19944,f19896,f19895,f19950,f19899,f19955,f19960,f20235,f20226,f20225,f20224,f20349,f20350,f20351,f20352,f20625,f20686,f20684,f20687,f20731,f20691,f20689,f20951,f20967,f20985,f21001,f21012,f21014,f21016,f21018,f21019,f21118,f21164,f21120,f21169,f21175,f21178,f2533,f21185,f21189,f2517,f21209,f21226,f21243,f21259,f21275,f21404,f21414,f21428,f21421,f21441,f21444,f21445,f21462,f21463,f21449,f21456,f21464,f21465,f21479,f21480,f21475,f21594,f21595,f21599,f21606,f21612,f21613,f21617,f21624,f21630,f21649,f21634,f21635,f21636,f21637,f21639,f21646,f21728,f21729,f21725,f21736,f21737,f21739,f21801,f21803,f21804,f21805,f21808,f21905,f21945,f21965,f21966,f21962,f22013,f22014,f22010,f21730,f22027,f22028,f22024,f22029,f22030,f22031,f22045,f22046,f21461,f22035,f22042,f22065,f22066,f22079,f22080,f21478,f22076,f22081,f22082,f21609,f22086,f22093,f22098,f22099,f21627,f22103,f22110,f22115,f22134,f21650,f22119,f21653,f22120,f21656,f22121,f21655,f22122,f22124,f22131,f22136,f22139,f22157,f22140,f9247,f22142,f1174,f22143,f22145,f22152,f22179,f22194,f22195,f22183,f22190,f22198,f22200,f22213,f21459,f22202,f22209,f22260,f22264,f22277,f22265,f21460,f22272,f22278,f22280,f22282,f22284,f22285,f22384,f22420,f22392,f22450,f22451,f22466,f22456,f21652,f22463,f22471,f22485,f22482,f22491,f22506,f22503,f22513,f22529,f22526,f22537,f22553,f22550,f22561,f22578,f22575,f22586,f22602,f22599,f22607,f22638,f22635,f22643,f22661,f22658,f22668,f22714,f22716,f22717,f22718,f22721,f22818,f22853,f23082,f23094,f23121,f10655,f23161,f10676,f23183,f10701,f23207,f10726,f23222,f4700,f23667,f23672,f23701,f23706,f23729,f6450,f23773,f5827,f23796,f5847,f23821,f5869,f23844,f5891,f23882,f6218,f24418,f6485,f24470,f6511,f24646,f12189,f24677,f12215,f24724,f12234,f24752,f12261,f24776,f12288,f24806,f12315,f24837,f12336,f24882,f12367,f24907,f12388,f24932,f12419,f24956,f12464,f24985,f12489,f25030,f12523,f25048,f12539,f25053,f25055,f25057,f25059,f25060,f25159,f25198,f25170,f25205,f25225,f25228,f25233,f25250,f25243,f25251,f25288,f25293,f25308,f25313,f25331,f25337,f25355,f25361,f25380,f25385,f25427,f25431,f25448,f25452,f25471,f25475,f25679,f25701,f25680,f25704,f25681,f25703,f27821,f27974,f2532,f2530,f27838,f27840,f27841,f27842,f27845,f27942,f27980,f28133,f28148,f28150,f28151,f28152,f28155,f28252,f28289,f28302,f28304,f28305,f28306,f28309,f28406,f28443,f30901,f31044,f31046,f31047,f31048,f31051,f31148,f31183,f31697,f31715,f31717,f31718,f31719,f31722,f31819,f31854,f31888,f31890,f31891,f31892,f31895,f31992,f32032,f32088,f32090,f32092,f32094,f32095,f32194,f32227,f32351,f32364,f32562,f32565,f32564,f32405,f32407,f32408,f32409,f32412,f32509,f32566,f32631,f32649,f32722,f32724,f32725,f32726,f32729,f32826,f32883,f33706,f33708,f33709,f33710,f33713,f33810,f33841,f33848,f33850,f33851,f33852,f33855,f33952,f33992,f34005,f34021,f34044,f34046,f34051,f34065,f34066,f34068,f34069,f34070,f34073,f34170,f34210,f34285,f34298,f34305,f34307,f34308,f34309,f34312,f34409,f34437,f34455,f34468,f34480,f34493,f34507,f34521,f34532,f34546,f34555,f34569,f34576,f34578,f34579,f34580,f34583,f34680,f34708,f34722,f34737,f34744,f34759,f34765,f34783,f34767,f34784,f34793,f34805,f34020,f34811,f34813,f34814,f34815,f34818,f34915,f34943,f35080,f35082,f35083,f35084,f35087,f35184,f35212,f35289,f35307,f35504,f35509,f33705,f35537,f35535,f28296,f35544,f35542,f3291,f27975,f35548,f35540,f35546,f26895,f7928,f26864,f7907,f26834,f7869,f26803,f7848,f26739,f7826,f26678,f7729,f26637,f7690,f26594,f7662,f26522,f7620,f26477,f7596,f26395,f7531,f26352,f7505,f26302,f7480,f35566,f35577,f3317,f35592,f9864,f35614,f35636,f35638,f35640,f35642,f35643,f35742,f2510,f35799,f2509,f2507,f13897,f35824,f35853,f35885,f35916,f35949,f36097,f36127,f36155,f20957,f36285,f36287,f36288,f36289,f36292,f36389,f14434,f36451,f10030,f36476,f10058,f36500,f10074,f36510,f36512,f36513,f36514,f36517,f36614,f2985,f36618,f2527,f2989,f36676,f10091,f36698,f10107,f36718,f10123,f36737,f10154,f36757,f10171,f36772,f36774,f36775,f36776,f36779,f36876,f36925,f36889,f36931,f36939,f36954,f36943,f10632,f36951,f10188,f36971,f10202,f36991,f3946,f37014,f3970,f37034,f3987,f37057,f4003,f37081,f4019,f37092,f5488,f37093,f10590,f37094,f5461,f37107,f4035,f37115,f10353,f37133,f4060,f37152,f4077,f37167,f37169,f37170,f37171,f37174,f37271,f37323,f37284,f37329,f2528,f37339,f10386,f3290,f37351,f4094,f37364,f37366,f37367,f37368,f37371,f37468,f3139,f37513,f37515,f37517,f717,f37519,f719,f37520,f718,f37619,f720,f3908,f37627,f3914,f3316,f37822,f37826])).
fof(f37826,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f37825,f555])).
fof(f555,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.LIST_REL_trans',aHL_FALSITY)).
fof(f37825,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(forward_demodulation,[],[f37824,f1677])).
fof(f1677,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f324,f510])).
fof(f510,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X21))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_F_EQ_F')).
fof(f324,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))))),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2,X3,X4] : (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X3),s('type.list.list'(X0),X4)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3,X4] : ((! [X5] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4] : ((! [X5] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3)))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((! [X5] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3)))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X4))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X3),s('type.list.list'(X0),X4))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((! [X11] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X9)))))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X9))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X10))))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X10))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X6),X9)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X9),s('type.list.list'(X6),X10))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X6),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((! [X11] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X9)))))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X9))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X10))))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X10))))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X6),X9)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X9),s('type.list.list'(X6),X10))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X6),X10)))))),
  file('thm.list.LIST_REL_trans','thm.list.LIST_REL_trans')).
fof(f37824,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | $spl16),
  inference(subsumption_resolution,[],[f37823,f1717])).
fof(f1717,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | $spl16),
  inference(cnf_transformation,[],[f1717_D])).
fof(f1717_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f37823,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(subsumption_resolution,[],[f37666,f327])).
fof(f327,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X12] : ! [X7,X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X12,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X12),X9)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X12),X9))) & ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) => p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X6,fun(X12,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X12,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X12),X9)))))))))),
  file('thm.list.LIST_REL_trans','thm.list.LIST_REL_EL_EQN')).
fof(f37666,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f37512,f328])).
fof(f328,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK18(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f217])).
fof(f37512,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) )),
  inference(subsumption_resolution,[],[f2499,f583])).
fof(f583,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) )),
  inference(resolution,[],[f322,f326])).
fof(f326,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f322,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))))),
  inference(cnf_transformation,[],[f213])).
fof(f2499,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))))))) )),
  inference(duplicate_literal_removal,[],[f2498])).
fof(f2498,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) )),
  inference(forward_demodulation,[],[f2497,f582])).
fof(f582,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))),
  inference(resolution,[],[f322,f325])).
fof(f325,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f2497,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) )),
  inference(subsumption_resolution,[],[f2342,f323])).
fof(f323,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(cnf_transformation,[],[f213])).
fof(f2342,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))) )),
  inference(resolution,[],[f321,f326])).
fof(f321,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) )),
  inference(cnf_transformation,[],[f213])).
fof(f37822,plain,(
  $false | ($spl18 | $spl21)),
  inference(subsumption_resolution,[],[f37665,f1720])).
fof(f37665,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl21),
  inference(resolution,[],[f37512,f1722])).
fof(f3316,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl62),
  inference(cnf_transformation,[],[f3316_D])).
fof(f3316_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3914,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl89),
  inference(cnf_transformation,[],[f3914_D])).
fof(f3914_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f37627,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ($spl42 | $spl57)),
  inference(superposition,[],[f20463,f568])).
fof(f568,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30])],[f157])).
fof(f157,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X25] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | ? [X11] : s('type.num.num',X25) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))),
  file('thm.list.LIST_REL_trans','thm.arithmetic.num_CASES')).
fof(f20463,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ($spl42 | $spl57)),
  inference(resolution,[],[f20454,f3823])).
fof(f3823,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X123)) ) | $spl57),
  inference(resolution,[],[f3291,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f20454,plain,(
  ( ! [X15] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X15,'F_0') | ~p(s(bool,X15))) ) | $spl42),
  inference(subsumption_resolution,[],[f20427,f540])).
fof(f540,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.EQ_CLAUSES|split|3')).
fof(f20427,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | ~p(s(bool,X15)) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X15,'F_0')) ) | $spl42),
  inference(resolution,[],[f9522,f10450])).
fof(f10450,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f10449,f545])).
fof(f545,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.EQ_CLAUSES|split|0')).
fof(f10449,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f10446,f10426])).
fof(f10426,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f10421,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f10421,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f9259,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f9259,plain,(
  ( ! [X8,X7] : (sP9(X8,'F_0',X7) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) )),
  inference(superposition,[],[f4977,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f103,f205,f204,f203,f202])).
fof(f203,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X36,X37,X38] : ((p(s(bool,X38)) <=> s(bool,X36) = s(bool,X37)) <=> ((p(s(bool,X38)) | p(s(bool,X37)) | p(s(bool,X36))) & (p(s(bool,X38)) | ~p(s(bool,X36)) | ~p(s(bool,X37))) & (p(s(bool,X37)) | ~p(s(bool,X36)) | ~p(s(bool,X38))) & (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)))))),
  file('thm.list.LIST_REL_trans','thm.sat.sat_2|49|_prove_ltb')).
fof(f4977,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f686,f684])).
fof(f684,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))),
  inference(resolution,[],[f322,f540])).
fof(f686,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))),
  inference(resolution,[],[f322,f545])).
fof(f10446,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f10427,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f10427,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f10422,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f264])).
fof(f264,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f203])).
fof(f10422,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f9259,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f259])).
fof(f9522,plain,(
  ( ! [X147,X148] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X148,X147) | s(bool,X147) = s(bool,X148)) ) | $spl42),
  inference(resolution,[],[f2527,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f3908,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl86),
  inference(cnf_transformation,[],[f3908_D])).
fof(f3908_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f720,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) | $spl8),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f37619,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X17,X26,X16,X27] : ((s(bool,X17) = s(bool,X26) & (p(s(bool,X26)) => s(bool,X16) = s(bool,X27))) => ((p(s(bool,X17)) => p(s(bool,X16))) <=> (p(s(bool,X26)) => p(s(bool,X27)))))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CONG')).
fof(f718,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X12))) ) | $spl4),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X12))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f37520,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f719,plain,(
  ( ! [X18] : (~p(s(bool,X18))) ) | $spl6),
  inference(cnf_transformation,[],[f719_D])).
fof(f719_D,plain,(
  ( ! [X18] : (~p(s(bool,X18))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f37519,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f717,plain,(
  ( ! [X7,X5] : (s(bool,X5) != s(bool,X7) | p(s(bool,X5))) ) | $spl2),
  inference(cnf_transformation,[],[f717_D])).
fof(f717_D,plain,(
  ( ! [X7,X5] : (s(bool,X5) != s(bool,X7) | p(s(bool,X5))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f37517,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f37515,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f37513,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl42 | $spl57)),
  inference(resolution,[],[f20463,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f3139,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl54),
  inference(cnf_transformation,[],[f3139_D])).
fof(f3139_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f37468,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f35569,f570])).
fof(f35569,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl42),
  inference(resolution,[],[f3082,f1724])).
fof(f1724,plain,(
  ( ! [X166] : (~sP12('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)),X166) | p(s(bool,X166))) )),
  inference(subsumption_resolution,[],[f1667,f527])).
fof(f527,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X21] : ((p(s(bool,'F_0')) & p(s(bool,X21))) <=> p(s(bool,'F_0')))),
  file('thm.list.LIST_REL_trans','thm.bool.AND_CLAUSES|split|2')).
fof(f1667,plain,(
  ( ! [X166] : (p(s(bool,'F_0')) | p(s(bool,X166)) | ~sP12('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)),X166)) )),
  inference(resolution,[],[f324,f482])).
fof(f482,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f3082,plain,(
  ( ! [X175] : (sP12(X175,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f483])).
fof(f483,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f283])).
fof(f37371,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl42),
  inference(resolution,[],[f35569,f375])).
fof(f37368,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f35569,f374])).
fof(f37367,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f35569,f373])).
fof(f37366,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl42),
  inference(resolution,[],[f35569,f371])).
fof(f37364,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f35569,f369])).
fof(f4094,plain,(
  ( ! [X10,X11] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl106),
  inference(cnf_transformation,[],[f4094_D])).
fof(f4094_D,plain,(
  ( ! [X10,X11] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f37351,plain,(
  ( ! [X10,X11] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3827,f568])).
fof(f3827,plain,(
  ( ! [X130,X131] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X130,X131)) ) | $spl57),
  inference(resolution,[],[f3291,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f265])).
fof(f3290,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl56),
  inference(cnf_transformation,[],[f3290_D])).
fof(f3290_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f10386,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl272),
  inference(cnf_transformation,[],[f10386_D])).
fof(f10386_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f37339,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X4)) ) | $spl57),
  inference(resolution,[],[f3827,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f2528,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl43),
  inference(cnf_transformation,[],[f2528_D])).
fof(f2528_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f37329,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) | $spl654),
  inference(cnf_transformation,[],[f37329_D])).
fof(f37329_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) <=> ~$spl654),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl654])])).
fof(f37284,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) )),
  inference(superposition,[],[f37039,f568])).
fof(f37039,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) )),
  inference(forward_demodulation,[],[f732,f582])).
fof(f732,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) )),
  inference(resolution,[],[f323,f326])).
fof(f37323,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17))))) | p(s(bool,X12))) ) | $spl652),
  inference(cnf_transformation,[],[f37323_D])).
fof(f37323_D,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17))))) | p(s(bool,X12))) ) <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f37271,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X292),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X293)) | p(s(bool,X294)) | s(bool,X294) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X292),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X292),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f37039,f570])).
fof(f37174,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X38),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X38),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f37039,f375])).
fof(f37171,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X26),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X26),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f37039,f374])).
fof(f37170,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f37039,f373])).
fof(f37169,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X21)) )),
  inference(resolution,[],[f37039,f371])).
fof(f37167,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f37039,f369])).
fof(f4077,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl104),
  inference(cnf_transformation,[],[f4077_D])).
fof(f4077_D,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f37152,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3826,f568])).
fof(f3826,plain,(
  ( ! [X128,X129] : (sP7(X128,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X129)) ) | $spl57),
  inference(resolution,[],[f3291,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f265])).
fof(f4060,plain,(
  ( ! [X10,X11] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl102),
  inference(cnf_transformation,[],[f4060_D])).
fof(f4060_D,plain,(
  ( ! [X10,X11] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f37133,plain,(
  ( ! [X10,X11] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3824,f568])).
fof(f3824,plain,(
  ( ! [X125,X124] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X124,X125)) ) | $spl57),
  inference(resolution,[],[f3291,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f10353,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl268),
  inference(cnf_transformation,[],[f10353_D])).
fof(f10353_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f37115,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3824,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f4035,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl100),
  inference(cnf_transformation,[],[f4035_D])).
fof(f4035_D,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f37107,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3823,f568])).
fof(f5461,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl125),
  inference(cnf_transformation,[],[f5461_D])).
fof(f5461_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f37094,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | $spl57),
  inference(resolution,[],[f3823,f10572])).
fof(f10572,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f10571,f443])).
fof(f10571,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f10570,f508])).
fof(f508,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CLAUSES|split|4')).
fof(f10570,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f10563,f441])).
fof(f10563,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f9261,f432])).
fof(f9261,plain,(
  ( ! [X12,X11] : (~sP9(X12,'F_0',X11) | ~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X11)) )),
  inference(superposition,[],[f4977,f448])).
fof(f10590,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl274),
  inference(cnf_transformation,[],[f10590_D])).
fof(f10590_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f37093,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3823,f10444])).
fof(f10444,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f10443,f555])).
fof(f10443,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f10440,f441])).
fof(f10440,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f10426,f432])).
fof(f5488,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl128),
  inference(cnf_transformation,[],[f5488_D])).
fof(f5488_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f37092,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X8))) ) | $spl57),
  inference(resolution,[],[f3823,f10519])).
fof(f10519,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f10513,f443])).
fof(f10513,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) )),
  inference(resolution,[],[f10474,f448])).
fof(f10474,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f10472,f555])).
fof(f10472,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f10469,f432])).
fof(f10469,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f10463,f439])).
fof(f10463,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f10460,f437])).
fof(f10460,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f10456,f440])).
fof(f10456,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f10450,f428])).
fof(f4019,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl98),
  inference(cnf_transformation,[],[f4019_D])).
fof(f4019_D,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f37081,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3811,f568])).
fof(f3811,plain,(
  ( ! [X99,X98] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X98,X99)) ) | $spl57),
  inference(resolution,[],[f3291,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f4003,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl96),
  inference(cnf_transformation,[],[f4003_D])).
fof(f4003_D,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f37057,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3810,f568])).
fof(f3810,plain,(
  ( ! [X97,X96] : (sP4(X96,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X97)) ) | $spl57),
  inference(resolution,[],[f3291,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f254])).
fof(f3987,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl94),
  inference(cnf_transformation,[],[f3987_D])).
fof(f3987_D,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f37034,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3797,f568])).
fof(f3797,plain,(
  ( ! [X70,X71] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X70,X71)) ) | $spl57),
  inference(resolution,[],[f3291,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f3970,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl92),
  inference(cnf_transformation,[],[f3970_D])).
fof(f3970_D,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f37014,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3781,f568])).
fof(f3781,plain,(
  ( ! [X39,X38] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X38,X39)) ) | $spl57),
  inference(resolution,[],[f3291,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f3946,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl90),
  inference(cnf_transformation,[],[f3946_D])).
fof(f3946_D,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f36991,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3780,f568])).
fof(f3780,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f10202,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl254),
  inference(cnf_transformation,[],[f10202_D])).
fof(f10202_D,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f36971,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3071,f568])).
fof(f3071,plain,(
  ( ! [X161,X162] : (sP10(X161,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X162)) ) | $spl42),
  inference(resolution,[],[f2527,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f105,f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6,X12,X42] : ((p(s(bool,X6)) | (p(s(bool,X12)) & p(s(bool,X42)))) <=> ((p(s(bool,X6)) | p(s(bool,X12))) & (p(s(bool,X6)) | p(s(bool,X42)))))),
  file('thm.list.LIST_REL_trans','thm.bool.LEFT_OR_OVER_AND')).
fof(f10188,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl252),
  inference(cnf_transformation,[],[f10188_D])).
fof(f10188_D,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f36951,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3062,f568])).
fof(f3062,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X143,X144)) ) | $spl42),
  inference(resolution,[],[f2527,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f10632,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl277),
  inference(cnf_transformation,[],[f10632_D])).
fof(f10632_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f36943,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f3062,f10520])).
fof(f10520,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(subsumption_resolution,[],[f10514,f437])).
fof(f10514,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f10474,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f36954,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X3)) ) | $spl640),
  inference(cnf_transformation,[],[f36954_D])).
fof(f36954_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X3)) ) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f36939,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X2)) ) | $spl42),
  inference(resolution,[],[f3062,f431])).
fof(f36931,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) | $spl638),
  inference(cnf_transformation,[],[f36931_D])).
fof(f36931_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f36889,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) )),
  inference(superposition,[],[f583,f568])).
fof(f36925,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16))))) | p(s(bool,X11))) ) | $spl636),
  inference(cnf_transformation,[],[f36925_D])).
fof(f36925_D,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16))))) | p(s(bool,X11))) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f36876,plain,(
  ( ! [X292,X293,X291] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X291),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X292)) | p(s(bool,X293)) | s(bool,X293) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X291),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X291),s('type.list.list'(sK13),sK16)))))) )),
  inference(resolution,[],[f583,f570])).
fof(f36779,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X37),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X37),s('type.list.list'(sK13),sK16))))) | s(bool,X39) != s(bool,X40)) )),
  inference(resolution,[],[f583,f375])).
fof(f36776,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(sK13),sK16)))))) )),
  inference(resolution,[],[f583,f374])).
fof(f36775,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(sK13),sK16)))))) )),
  inference(resolution,[],[f583,f373])).
fof(f36774,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK13),sK16))))) | s(bool,X19) != s(bool,X20)) )),
  inference(resolution,[],[f583,f371])).
fof(f36772,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16)))))) )),
  inference(resolution,[],[f583,f369])).
fof(f10171,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl250),
  inference(cnf_transformation,[],[f10171_D])).
fof(f10171_D,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f36757,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3061,f568])).
fof(f3061,plain,(
  ( ! [X142,X141] : (sP6(X141,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X142)) ) | $spl42),
  inference(resolution,[],[f2527,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f10154,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl248),
  inference(cnf_transformation,[],[f10154_D])).
fof(f10154_D,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f36737,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3060,f568])).
fof(f3060,plain,(
  ( ! [X140,X139] : (sP6(X139,X140,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f443])).
fof(f10123,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl246),
  inference(cnf_transformation,[],[f10123_D])).
fof(f10123_D,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f36718,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3059,f568])).
fof(f3059,plain,(
  ( ! [X138,X137] : (sP7(X137,X138,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f439])).
fof(f10107,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl244),
  inference(cnf_transformation,[],[f10107_D])).
fof(f10107_D,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f36698,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3056,f568])).
fof(f3056,plain,(
  ( ! [X132,X131] : (sP8(X131,X132,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f262])).
fof(f10091,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl242),
  inference(cnf_transformation,[],[f10091_D])).
fof(f10091_D,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f36676,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3043,f568])).
fof(f3043,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f2989,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl52),
  inference(cnf_transformation,[],[f2989_D])).
fof(f2989_D,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2527,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl42),
  inference(cnf_transformation,[],[f2527_D])).
fof(f2527_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f36618,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl18),
  inference(superposition,[],[f1720,f568])).
fof(f2985,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl50),
  inference(cnf_transformation,[],[f2985_D])).
fof(f2985_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f36614,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f570])).
fof(f36517,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl18),
  inference(resolution,[],[f1720,f375])).
fof(f36514,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f374])).
fof(f36513,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f373])).
fof(f36512,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl18),
  inference(resolution,[],[f1720,f371])).
fof(f36510,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f369])).
fof(f10074,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl240),
  inference(cnf_transformation,[],[f10074_D])).
fof(f10074_D,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f36500,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3027,f568])).
fof(f3027,plain,(
  ( ! [X74,X73] : (sP2(X73,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X74)) ) | $spl42),
  inference(resolution,[],[f2527,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f247])).
fof(f10058,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl238),
  inference(cnf_transformation,[],[f10058_D])).
fof(f10058_D,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f36476,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3026,f568])).
fof(f3026,plain,(
  ( ! [X72,X71] : (sP2(X71,X72,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f10030,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl236),
  inference(cnf_transformation,[],[f10030_D])).
fof(f10030_D,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f36451,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3013,f568])).
fof(f3013,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X46)) ) | $spl42),
  inference(resolution,[],[f2527,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f14434,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl362),
  inference(cnf_transformation,[],[f14434_D])).
fof(f14434_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f36389,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f570])).
fof(f36292,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl34),
  inference(resolution,[],[f2508,f375])).
fof(f36289,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f374])).
fof(f36288,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f373])).
fof(f36287,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl34),
  inference(resolution,[],[f2508,f371])).
fof(f36285,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f369])).
fof(f20957,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl428),
  inference(cnf_transformation,[],[f20957_D])).
fof(f20957_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f36155,plain,(
  ( ! [X6,X7] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22339,f568])).
fof(f22339,plain,(
  ( ! [X130,X131] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X130,X131)) ) | $spl437),
  inference(resolution,[],[f21460,f441])).
fof(f36127,plain,(
  ( ! [X6,X7] : (sP7(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22338,f568])).
fof(f22338,plain,(
  ( ! [X128,X129] : (sP7(X128,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X129)) ) | $spl437),
  inference(resolution,[],[f21460,f440])).
fof(f36097,plain,(
  ( ! [X6,X7] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22336,f568])).
fof(f22336,plain,(
  ( ! [X125,X124] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X124,X125)) ) | $spl437),
  inference(resolution,[],[f21460,f437])).
fof(f35949,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22335,f568])).
fof(f22335,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X123)) ) | $spl437),
  inference(resolution,[],[f21460,f436])).
fof(f35916,plain,(
  ( ! [X6,X7] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22323,f568])).
fof(f22323,plain,(
  ( ! [X99,X98] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X98,X99)) ) | $spl437),
  inference(resolution,[],[f21460,f420])).
fof(f35885,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22322,f568])).
fof(f22322,plain,(
  ( ! [X97,X96] : (sP4(X96,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X97)) ) | $spl437),
  inference(resolution,[],[f21460,f419])).
fof(f35853,plain,(
  ( ! [X6,X7] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22309,f568])).
fof(f22309,plain,(
  ( ! [X70,X71] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X70,X71)) ) | $spl437),
  inference(resolution,[],[f21460,f403])).
fof(f35824,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22293,f568])).
fof(f22293,plain,(
  ( ! [X39,X38] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X38,X39)) ) | $spl437),
  inference(resolution,[],[f21460,f386])).
fof(f13897,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | $spl345),
  inference(cnf_transformation,[],[f13897_D])).
fof(f13897_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f2507,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | $spl32),
  inference(cnf_transformation,[],[f2507_D])).
fof(f2507_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2509,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl35),
  inference(cnf_transformation,[],[f2509_D])).
fof(f2509_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f35799,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ($spl30 | $spl356)),
  inference(forward_demodulation,[],[f35798,f14241])).
fof(f35798,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ($spl30 | $spl356)),
  inference(subsumption_resolution,[],[f35797,f2504])).
fof(f2504,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | $spl30),
  inference(cnf_transformation,[],[f2504_D])).
fof(f2504_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f35797,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl356),
  inference(forward_demodulation,[],[f35796,f14241])).
fof(f35796,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl356),
  inference(forward_demodulation,[],[f35755,f14241])).
fof(f35755,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl356),
  inference(superposition,[],[f321,f14241])).
fof(f2510,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))))) ) | $spl36),
  inference(cnf_transformation,[],[f2510_D])).
fof(f2510_D,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))) | ~p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f35742,plain,(
  ( ! [X296,X295,X297] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X295),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X296)) | ~p(s(bool,X297)) | s(bool,X297) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f570])).
fof(f35643,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f321,f375])).
fof(f35642,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f374])).
fof(f35640,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f373])).
fof(f35638,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f321,f371])).
fof(f35636,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f369])).
fof(f35614,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22292,f568])).
fof(f22292,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f384])).
fof(f9864,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8)) ) | $spl234),
  inference(cnf_transformation,[],[f9864_D])).
fof(f9864_D,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f35592,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3083,f568])).
fof(f3083,plain,(
  ( ! [X176] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X176)) ) | $spl42),
  inference(resolution,[],[f2527,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f3317,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f3317_D])).
fof(f3317_D,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f35577,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3082,f568])).
fof(f35566,plain,(
  $false | ($spl224 | $spl424)),
  inference(subsumption_resolution,[],[f20689,f20641])).
fof(f20641,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | $spl224),
  inference(superposition,[],[f567,f8754])).
fof(f567,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X11] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))),
  file('thm.list.LIST_REL_trans','thm.num.NOT_SUC')).
fof(f7480,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl192),
  inference(cnf_transformation,[],[f7480_D])).
fof(f7480_D,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f26302,plain,(
  ( ! [X6,X7] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22286,f568])).
fof(f22286,plain,(
  ( ! [X24,X25] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X25,X24) | p(s(bool,X24))) ) | $spl437),
  inference(resolution,[],[f21460,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f7505,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) | $spl194),
  inference(cnf_transformation,[],[f7505_D])).
fof(f7505_D,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f26352,plain,(
  ( ! [X6,X7] : (~sP1(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22287,f568])).
fof(f22287,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X26))) ) | $spl437),
  inference(resolution,[],[f21460,f376])).
fof(f7531,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl196),
  inference(cnf_transformation,[],[f7531_D])).
fof(f7531_D,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f26395,plain,(
  ( ! [X6,X7] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22288,f568])).
fof(f22288,plain,(
  ( ! [X28,X29] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X28,X29) | ~p(s(bool,X28))) ) | $spl437),
  inference(resolution,[],[f21460,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f7596,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) | $spl200),
  inference(cnf_transformation,[],[f7596_D])).
fof(f7596_D,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f26477,plain,(
  ( ! [X6,X7] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22301,f568])).
fof(f22301,plain,(
  ( ! [X54,X55] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X55,X54) | ~p(s(bool,X54))) ) | $spl437),
  inference(resolution,[],[f21460,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f7620,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl202),
  inference(cnf_transformation,[],[f7620_D])).
fof(f7620_D,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f26522,plain,(
  ( ! [X6,X7] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22302,f568])).
fof(f22302,plain,(
  ( ! [X57,X56] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X56,X57) | ~p(s(bool,X56))) ) | $spl437),
  inference(resolution,[],[f21460,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f244])).
fof(f7662,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl204),
  inference(cnf_transformation,[],[f7662_D])).
fof(f7662_D,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f26594,plain,(
  ( ! [X6,X7] : (~sP5(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22317,f568])).
fof(f22317,plain,(
  ( ! [X87,X86] : (~sP5(X86,X87,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X86))) ) | $spl437),
  inference(resolution,[],[f21460,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f7690,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) | $spl206),
  inference(cnf_transformation,[],[f7690_D])).
fof(f7690_D,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f26637,plain,(
  ( ! [X6,X7] : (~sP5(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22318,f568])).
fof(f22318,plain,(
  ( ! [X88,X89] : (~sP5(X88,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X89) | ~p(s(bool,X88))) ) | $spl437),
  inference(resolution,[],[f21460,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f7729,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl208),
  inference(cnf_transformation,[],[f7729_D])).
fof(f7729_D,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f26678,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22324,f568])).
fof(f22324,plain,(
  ( ! [X101,X100] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X100,X101) | p(s(bool,X101))) ) | $spl437),
  inference(resolution,[],[f21460,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f101,f200,f199])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X36,X37,X38] : ((p(s(bool,X38)) <=> (p(s(bool,X37)) & p(s(bool,X36)))) <=> ((p(s(bool,X38)) | ~p(s(bool,X37)) | ~p(s(bool,X36))) & (p(s(bool,X37)) | ~p(s(bool,X38))) & (p(s(bool,X36)) | ~p(s(bool,X38)))))),
  file('thm.list.LIST_REL_trans','thm.sat.sat_23_prove_ltb')).
fof(f7826,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) | $spl214),
  inference(cnf_transformation,[],[f7826_D])).
fof(f7826_D,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f26739,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22327,f568])).
fof(f22327,plain,(
  ( ! [X107,X106] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X106,X107) | p(s(bool,X106))) ) | $spl437),
  inference(resolution,[],[f21460,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f7848,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11))) ) | $spl216),
  inference(cnf_transformation,[],[f7848_D])).
fof(f7848_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f26803,plain,(
  ( ! [X6,X7] : (~sP10(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22345,f568])).
fof(f22345,plain,(
  ( ! [X142,X143] : (~sP10(X143,X142,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X142))) ) | $spl437),
  inference(resolution,[],[f21460,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f272])).
fof(f272,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f207])).
fof(f7869,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11))) ) | $spl218),
  inference(cnf_transformation,[],[f7869_D])).
fof(f7869_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f26834,plain,(
  ( ! [X6,X7] : (~sP10(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22346,f568])).
fof(f22346,plain,(
  ( ! [X144,X145] : (~sP10(X145,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X144) | p(s(bool,X144))) ) | $spl437),
  inference(resolution,[],[f21460,f456])).
fof(f7907,plain,(
  ( ! [X10,X11] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) | $spl220),
  inference(cnf_transformation,[],[f7907_D])).
fof(f7907_D,plain,(
  ( ! [X10,X11] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f26864,plain,(
  ( ! [X6,X7] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22347,f568])).
fof(f22347,plain,(
  ( ! [X146,X147] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X146,X147) | p(s(bool,X146))) ) | $spl437),
  inference(resolution,[],[f21460,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f273])).
fof(f7928,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl222),
  inference(cnf_transformation,[],[f7928_D])).
fof(f7928_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f26895,plain,(
  ( ! [X6,X7] : (~sP10(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22348,f568])).
fof(f22348,plain,(
  ( ! [X149,X148] : (~sP10(X148,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X149) | p(s(bool,X148))) ) | $spl437),
  inference(resolution,[],[f21460,f457])).
fof(f35546,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl609),
  inference(cnf_transformation,[],[f35546_D])).
fof(f35546_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) <=> ~$spl609),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl609])])).
fof(f35540,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | $spl603),
  inference(cnf_transformation,[],[f35540_D])).
fof(f35540_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) <=> ~$spl603),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl603])])).
fof(f35548,plain,(
  s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) | $spl611),
  inference(cnf_transformation,[],[f35548_D])).
fof(f35548_D,plain,(
  s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) <=> ~$spl611),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl611])])).
fof(f27975,plain,(
  s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ($spl29 | $spl40)),
  inference(subsumption_resolution,[],[f27832,f2503])).
fof(f27832,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl40),
  inference(resolution,[],[f2516,f349])).
fof(f349,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))) | s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X3))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X3))))))))) & (! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))) | p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29])],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | ? [X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))))))) & (! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))) | p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | ? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))))))))) & (! [X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))) | p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) <=> ? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) <=> ? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) <=> ? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X3)))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.list.LLEX_3'(s(fun(X6,fun(X6,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(X6),X9)))) <=> ? [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9)))))) & s('type.list.list'(X6),'const.list.TAKE_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.TAKE_2'(s('type.num.num',X11),s('type.list.list'(X6),X9))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X8))))),s(X6,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(X6),X9)))))))))),
  file('thm.list.LIST_REL_trans','thm.list.LLEX_EL_THM')).
fof(f2516,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | $spl40),
  inference(cnf_transformation,[],[f2516_D])).
fof(f2516_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3291,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(cnf_transformation,[],[f3291_D])).
fof(f3291_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f35542,plain,(
  s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) | $spl605),
  inference(cnf_transformation,[],[f35542_D])).
fof(f35542_D,plain,(
  s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) <=> ~$spl605),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl605])])).
fof(f35544,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | $spl606),
  inference(cnf_transformation,[],[f35544_D])).
fof(f35544_D,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f28296,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl46),
  inference(resolution,[],[f2532,f349])).
fof(f35535,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl599),
  inference(cnf_transformation,[],[f35535_D])).
fof(f35535_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl599),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl599])])).
fof(f35537,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | $spl600),
  inference(cnf_transformation,[],[f35537_D])).
fof(f35537_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f33705,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl544),
  inference(resolution,[],[f32564,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,'const.list.EL_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X2))))),s(X0,'const.list.EL_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f35509,plain,(
  $false | ($spl16 | $spl18 | $spl226)),
  inference(subsumption_resolution,[],[f35508,f8756])).
fof(f35508,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | ($spl16 | $spl18 | $spl226)),
  inference(forward_demodulation,[],[f35507,f20759])).
fof(f20759,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | ($spl16 | $spl226)),
  inference(backward_demodulation,[],[f8756,f1717])).
fof(f35507,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | ($spl18 | $spl226)),
  inference(subsumption_resolution,[],[f35506,f555])).
fof(f35506,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | ($spl18 | $spl226)),
  inference(forward_demodulation,[],[f35505,f1677])).
fof(f35505,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | ($spl18 | $spl226)),
  inference(subsumption_resolution,[],[f35343,f31710])).
fof(f31710,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl18 | $spl226)),
  inference(forward_demodulation,[],[f1720,f8756])).
fof(f35343,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | $spl226),
  inference(resolution,[],[f35329,f328])).
fof(f35329,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl226),
  inference(subsumption_resolution,[],[f35328,f32329])).
fof(f32329,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl226),
  inference(forward_demodulation,[],[f583,f8756])).
fof(f35328,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(forward_demodulation,[],[f2499,f8756])).
fof(f35504,plain,(
  $false | ($spl18 | $spl21 | $spl226)),
  inference(subsumption_resolution,[],[f35342,f31710])).
fof(f35342,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl21 | $spl226)),
  inference(resolution,[],[f35329,f1722])).
fof(f35307,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl597),
  inference(cnf_transformation,[],[f35307_D])).
fof(f35307_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl597),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl597])])).
fof(f35289,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f34130,f10520])).
fof(f34130,plain,(
  ( ! [X136,X137] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X136,X137)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f445])).
fof(f33839,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl544)),
  inference(forward_demodulation,[],[f33703,f20758])).
fof(f20758,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) | $spl226),
  inference(backward_demodulation,[],[f8756,f582])).
fof(f33703,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl544),
  inference(resolution,[],[f32564,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f35212,plain,(
  ( ! [X10,X8] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13)))) | p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl594),
  inference(cnf_transformation,[],[f35212_D])).
fof(f35212_D,plain,(
  ( ! [X10,X8] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13)))) | p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f35184,plain,(
  ( ! [X292,X290,X291] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X290),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X290),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X291)) | p(s(bool,X292)) | s(bool,X292) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X290,sK14,sK13,sK13)))),s(sK13,sK20(X290,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f727,f570])).
fof(f727,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X0,sK14,sK13,sK13)))),s(sK13,sK20(X0,sK14,sK13,sK13))))) | p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X0),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X0),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) )),
  inference(resolution,[],[f323,f329])).
fof(f329,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK19(X3,X2,X1,X0)))),s(X1,sK20(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK19(X3,X2,X1,X0)))),s(X1,sK20(X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK21(X3,X2,X1,X0)),s('type.list.list'(X1),sK22(X3,X2,X1,X0))))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),sK21(X3,X2,X1,X0)))),s('type.list.list'(X1),sK22(X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK19(X3,X2,X1,X0)),s('type.list.list'(X0),sK21(X3,X2,X1,X0)))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK20(X3,X2,X1,X0)),s('type.list.list'(X1),sK22(X3,X2,X1,X0)))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20,sK21,sK22])],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X6,X7,X8,X9] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X8))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X7),s('type.list.list'(X1),X9))))))) | ! [X4,X5] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7)))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) | ? [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) & ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) | ! [X8,X9] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) | p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3] : ((p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0')))) & ! [X4,X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X6),s('type.list.list'(X1),X7)))) & p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X6))),s('type.list.list'(X1),X7))))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X6))))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X7)))))))) => ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X8),s('type.list.list'(X1),X9)))) => p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X8))),s('type.list.list'(X1),X9))))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6,X12] : ! [X7,X31] : ((p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X12),bool)),X31),s('type.list.list'(X6),'const.list.NIL_0'))),s('type.list.list'(X12),'const.list.NIL_0')))) & ! [X32,X33,X13,X14] : ((p(s(bool,happ(s(fun(X12,bool),happ(s(fun(X6,fun(X12,bool)),X7),s(X6,X32))),s(X12,X33)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X12,bool)),X7),s('type.list.list'(X6),X13),s('type.list.list'(X12),X14)))) & p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X12),bool)),X31),s('type.list.list'(X6),X13))),s('type.list.list'(X12),X14))))) => p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X12),bool)),X31),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X32),s('type.list.list'(X6),X13))))),s('type.list.list'(X12),'const.list.CONS_2'(s(X12,X33),s('type.list.list'(X12),X14)))))))) => ! [X34,X35] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X12,bool)),X7),s('type.list.list'(X6),X34),s('type.list.list'(X12),X35)))) => p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.list.list'(X6),fun('type.list.list'(X12),bool)),X31),s('type.list.list'(X6),X34))),s('type.list.list'(X12),X35))))))),
  file('thm.list.LIST_REL_trans','thm.list.LIST_REL_strongind')).
fof(f35087,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X36),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X36),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X36,sK14,sK13,sK13)))),s(sK13,sK20(X36,sK14,sK13,sK13)))) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f727,f375])).
fof(f35084,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X24),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X24),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X24,sK14,sK13,sK13)))),s(sK13,sK20(X24,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f727,f374])).
fof(f35083,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X20),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X20),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X20,sK14,sK13,sK13)))),s(sK13,sK20(X20,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f727,f373])).
fof(f35082,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X16),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X16),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X16,sK14,sK13,sK13)))),s(sK13,sK20(X16,sK14,sK13,sK13)))) | s(bool,X18) != s(bool,X19)) )),
  inference(resolution,[],[f727,f371])).
fof(f35080,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f727,f369])).
fof(f34943,plain,(
  ( ! [X10,X8] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13)))) | p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl592),
  inference(cnf_transformation,[],[f34943_D])).
fof(f34943_D,plain,(
  ( ! [X10,X8] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13)))) | p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f34915,plain,(
  ( ! [X292,X290,X291] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X290),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X290),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X291)) | p(s(bool,X292)) | s(bool,X292) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X290,sK14,sK13,sK13)))),s(sK13,sK20(X290,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f578,f570])).
fof(f578,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X0,sK14,sK13,sK13)))),s(sK13,sK20(X0,sK14,sK13,sK13))))) | p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X0),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X0),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) )),
  inference(resolution,[],[f322,f329])).
fof(f34818,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X36),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X36),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X36,sK14,sK13,sK13)))),s(sK13,sK20(X36,sK14,sK13,sK13)))) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f578,f375])).
fof(f34815,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X24),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X24),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X24,sK14,sK13,sK13)))),s(sK13,sK20(X24,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f578,f374])).
fof(f34814,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X20),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X20),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X20,sK14,sK13,sK13)))),s(sK13,sK20(X20,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f578,f373])).
fof(f34813,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X16),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X16),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X16,sK14,sK13,sK13)))),s(sK13,sK20(X16,sK14,sK13,sK13)))) | s(bool,X18) != s(bool,X19)) )),
  inference(resolution,[],[f578,f371])).
fof(f34811,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X8),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(X8,sK14,sK13,sK13)))),s(sK13,sK20(X8,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f578,f369])).
fof(f34020,plain,(
  s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13))))) | $spl558),
  inference(cnf_transformation,[],[f34020_D])).
fof(f34020_D,plain,(
  s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13))))) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f34805,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl590),
  inference(cnf_transformation,[],[f34805_D])).
fof(f34805_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f34793,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33921,f568])).
fof(f33921,plain,(
  ( ! [X154,X155] : (sP10(X154,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X155)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f462])).
fof(f33838,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl544)),
  inference(forward_demodulation,[],[f33702,f8756])).
fof(f33702,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl544),
  inference(resolution,[],[f32564,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f34784,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl588),
  inference(cnf_transformation,[],[f34784_D])).
fof(f34784_D,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f34767,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0,X1) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33912,f568])).
fof(f33912,plain,(
  ( ! [X136,X137] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X136,X137)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f445])).
fof(f34783,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl587),
  inference(cnf_transformation,[],[f34783_D])).
fof(f34783_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl587),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl587])])).
fof(f34765,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33912,f10520])).
fof(f34759,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl584),
  inference(cnf_transformation,[],[f34759_D])).
fof(f34759_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f34744,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33911,f568])).
fof(f33911,plain,(
  ( ! [X134,X135] : (sP6(X134,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X135)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f444])).
fof(f34737,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl582),
  inference(cnf_transformation,[],[f34737_D])).
fof(f34737_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f34722,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33910,f568])).
fof(f33910,plain,(
  ( ! [X132,X133] : (sP6(X132,X133,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f443])).
fof(f34708,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13)))) | p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl580),
  inference(cnf_transformation,[],[f34708_D])).
fof(f34708_D,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13)))) | p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f34680,plain,(
  ( ! [X294,X293,X295] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X293),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X293),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X293,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X293,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f728,f570])).
fof(f728,plain,(
  ( ! [X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X1,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X1,sK14,sK13,sK13))))) | p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X1),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X1),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) )),
  inference(resolution,[],[f323,f330])).
fof(f330,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),X4))),s('type.list.list'(X1),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X4),s('type.list.list'(X1),X5)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK21(X3,X2,X1,X0)),s('type.list.list'(X1),sK22(X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X1),bool),happ(s(fun('type.list.list'(X0),fun('type.list.list'(X1),bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X1),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f34583,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X39),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X39),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X39,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X39,sK14,sK13,sK13)))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f728,f375])).
fof(f34580,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X27),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X27),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X27,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X27,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f728,f374])).
fof(f34579,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X23),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X23),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X23,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X23,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f728,f373])).
fof(f34578,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X19),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X19),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X19,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X19,sK14,sK13,sK13)))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f728,f371])).
fof(f34576,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK16))),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f728,f369])).
fof(f34569,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl578),
  inference(cnf_transformation,[],[f34569_D])).
fof(f34569_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f34555,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33909,f568])).
fof(f33909,plain,(
  ( ! [X130,X131] : (sP7(X130,X131,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f439])).
fof(f34546,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl576),
  inference(cnf_transformation,[],[f34546_D])).
fof(f34546_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f34532,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33906,f568])).
fof(f33906,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f435])).
fof(f34521,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl574),
  inference(cnf_transformation,[],[f34521_D])).
fof(f34521_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f34507,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33893,f568])).
fof(f33893,plain,(
  ( ! [X99,X98] : (sP4(X98,X99,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f418])).
fof(f34493,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl572),
  inference(cnf_transformation,[],[f34493_D])).
fof(f34493_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f34480,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33877,f568])).
fof(f33877,plain,(
  ( ! [X66,X67] : (sP2(X66,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X67)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f402])).
fof(f34468,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl570),
  inference(cnf_transformation,[],[f34468_D])).
fof(f34468_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f34455,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33876,f568])).
fof(f33876,plain,(
  ( ! [X64,X65] : (sP2(X64,X65,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f401])).
fof(f34437,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13)))) | p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl568),
  inference(cnf_transformation,[],[f34437_D])).
fof(f34437_D,plain,(
  ( ! [X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13)))) | p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f34409,plain,(
  ( ! [X294,X293,X295] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X293),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X293),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X294)) | p(s(bool,X295)) | s(bool,X295) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X293,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X293,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f579,f570])).
fof(f579,plain,(
  ( ! [X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X1,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X1,sK14,sK13,sK13))))) | p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X1),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X1),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0'))))) )),
  inference(resolution,[],[f322,f330])).
fof(f34312,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X39),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X39),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X39,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X39,sK14,sK13,sK13)))) | s(bool,X41) != s(bool,X42)) )),
  inference(resolution,[],[f579,f375])).
fof(f34309,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X27),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X27),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X28)) | ~p(s(bool,X29)) | p(s(bool,X30)) | s(bool,X30) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X27,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X27,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f579,f374])).
fof(f34308,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X23),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X23),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X23,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X23,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f579,f373])).
fof(f34307,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X19),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X19),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X19,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X19,sK14,sK13,sK13)))) | s(bool,X21) != s(bool,X22)) )),
  inference(resolution,[],[f579,f371])).
fof(f34305,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK16)))) | ~p(s(bool,happ(s(fun('type.list.list'(sK13),bool),happ(s(fun('type.list.list'(sK13),fun('type.list.list'(sK13),bool)),X11),s('type.list.list'(sK13),'const.list.NIL_0'))),s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK21(X11,sK14,sK13,sK13)),s('type.list.list'(sK13),sK22(X11,sK14,sK13,sK13))))) )),
  inference(resolution,[],[f579,f369])).
fof(f34298,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl566),
  inference(cnf_transformation,[],[f34298_D])).
fof(f34298_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f34285,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33863,f568])).
fof(f33863,plain,(
  ( ! [X39,X38] : (sP0(X38,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X39)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f385])).
fof(f34210,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl564),
  inference(cnf_transformation,[],[f34210_D])).
fof(f34210_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f34170,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f570])).
fof(f34073,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X22) != s(bool,X23)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f375])).
fof(f34070,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f374])).
fof(f34069,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f373])).
fof(f34068,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) != s(bool,X8)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f371])).
fof(f34066,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33839,f369])).
fof(f34065,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0)) ) | $spl562),
  inference(cnf_transformation,[],[f34065_D])).
fof(f34065_D,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0)) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f34051,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33933,f568])).
fof(f33933,plain,(
  ( ! [X169] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X169)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f484])).
fof(f34046,plain,(
  $false | ($spl437 | $spl560)),
  inference(subsumption_resolution,[],[f34045,f555])).
fof(f34045,plain,(
  p(s(bool,'F_0')) | ($spl437 | $spl560)),
  inference(forward_demodulation,[],[f34027,f22371])).
fof(f22371,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl437),
  inference(resolution,[],[f21460,f510])).
fof(f34027,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl437 | $spl560)),
  inference(resolution,[],[f34021,f22423])).
fof(f22423,plain,(
  ( ! [X166] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X166) | p(s(bool,X166))) ) | $spl437),
  inference(subsumption_resolution,[],[f22361,f527])).
fof(f22361,plain,(
  ( ! [X166] : (p(s(bool,'F_0')) | p(s(bool,X166)) | ~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X166)) ) | $spl437),
  inference(resolution,[],[f21460,f482])).
fof(f34044,plain,(
  $false | ($spl437 | $spl560)),
  inference(subsumption_resolution,[],[f34043,f555])).
fof(f34043,plain,(
  p(s(bool,'F_0')) | ($spl437 | $spl560)),
  inference(forward_demodulation,[],[f34026,f22371])).
fof(f34026,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl560),
  inference(resolution,[],[f34021,f1724])).
fof(f34021,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl560),
  inference(cnf_transformation,[],[f34021_D])).
fof(f34021_D,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f34005,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',sK29(sK16,sK15,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK16,sK15,sK14,sK13)))))) ) | ($spl226 | $spl544)),
  inference(superposition,[],[f33932,f568])).
fof(f33932,plain,(
  ( ! [X168] : (sP12(X168,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f483])).
fof(f33992,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl556),
  inference(cnf_transformation,[],[f33992_D])).
fof(f33992_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl556])])).
fof(f33952,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f570])).
fof(f33855,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X22) != s(bool,X23)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f375])).
fof(f33852,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f374])).
fof(f33851,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f373])).
fof(f33850,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) != s(bool,X8)) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f371])).
fof(f33848,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK16,sK15,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl544)),
  inference(resolution,[],[f33838,f369])).
fof(f33841,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl554),
  inference(cnf_transformation,[],[f33841_D])).
fof(f33841_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f33810,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl544),
  inference(resolution,[],[f32564,f570])).
fof(f33713,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))) | s(bool,X22) != s(bool,X23)) ) | $spl544),
  inference(resolution,[],[f32564,f375])).
fof(f33710,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl544),
  inference(resolution,[],[f32564,f374])).
fof(f33709,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl544),
  inference(resolution,[],[f32564,f373])).
fof(f33708,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))) | s(bool,X7) != s(bool,X8)) ) | $spl544),
  inference(resolution,[],[f32564,f371])).
fof(f33706,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl544),
  inference(resolution,[],[f32564,f369])).
fof(f32883,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17))))) | p(s(bool,X12))) ) | $spl552),
  inference(cnf_transformation,[],[f32883_D])).
fof(f32883_D,plain,(
  ( ! [X12,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17))))) | p(s(bool,X12))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f32826,plain,(
  ( ! [X292,X294,X293] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X292),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X293)) | p(s(bool,X294)) | s(bool,X294) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X292),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X292),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f32605,f570])).
fof(f32605,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl226),
  inference(forward_demodulation,[],[f732,f20758])).
fof(f32729,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X38),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X38),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X41)) ) | $spl226),
  inference(resolution,[],[f32605,f375])).
fof(f32726,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X26),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X26),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f32605,f374])).
fof(f32725,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X23) != s(bool,X25) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X22),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f32605,f373])).
fof(f32724,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X21)) ) | $spl226),
  inference(resolution,[],[f32605,f371])).
fof(f32722,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X10),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f32605,f369])).
fof(f32649,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl551),
  inference(cnf_transformation,[],[f32649_D])).
fof(f32649_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl551),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl551])])).
fof(f32631,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31701,f10520])).
fof(f31701,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X143,X144)) ) | ($spl18 | $spl226)),
  inference(forward_demodulation,[],[f2909,f8756])).
fof(f2909,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X143,X144)) ) | $spl18),
  inference(resolution,[],[f1720,f445])).
fof(f32566,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16))))) | p(s(bool,X11))) ) | $spl548),
  inference(cnf_transformation,[],[f32566_D])).
fof(f32566_D,plain,(
  ( ! [X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16))))) | p(s(bool,X11))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f32509,plain,(
  ( ! [X292,X293,X291] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X291),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X292)) | p(s(bool,X293)) | s(bool,X293) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X291),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X291),s('type.list.list'(sK13),sK16)))))) ) | $spl226),
  inference(resolution,[],[f32329,f570])).
fof(f32412,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X37),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X37),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X37),s('type.list.list'(sK13),sK16))))) | s(bool,X39) != s(bool,X40)) ) | $spl226),
  inference(resolution,[],[f32329,f375])).
fof(f32409,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(sK13),sK16)))))) ) | $spl226),
  inference(resolution,[],[f32329,f374])).
fof(f32408,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X22)) | p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,X23) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(sK13),sK16)))))) ) | $spl226),
  inference(resolution,[],[f32329,f373])).
fof(f32407,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK13),sK16))))) | s(bool,X19) != s(bool,X20)) ) | $spl226),
  inference(resolution,[],[f32329,f371])).
fof(f32405,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK13),sK16)))))) ) | $spl226),
  inference(resolution,[],[f32329,f369])).
fof(f32564,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | $spl544),
  inference(cnf_transformation,[],[f32564_D])).
fof(f32564_D,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f32565,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl546),
  inference(cnf_transformation,[],[f32565_D])).
fof(f32565_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f32562,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))) ) | $spl226),
  inference(forward_demodulation,[],[f32561,f8756])).
fof(f32561,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl226),
  inference(duplicate_literal_removal,[],[f32560])).
fof(f32560,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl226),
  inference(forward_demodulation,[],[f32399,f20758])).
fof(f32399,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl226),
  inference(resolution,[],[f32329,f349])).
fof(f32364,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)))) | $spl543),
  inference(cnf_transformation,[],[f32364_D])).
fof(f32364_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)))) <=> ~$spl543),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl543])])).
fof(f32351,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl52),
  inference(resolution,[],[f31779,f10520])).
fof(f31779,plain,(
  ( ! [X136,X137] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))),X136,X137)) ) | $spl52),
  inference(resolution,[],[f19847,f445])).
fof(f19847,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl52),
  inference(superposition,[],[f560,f2989])).
fof(f560,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))))),
  file('thm.list.LIST_REL_trans','thm.prim_rec.LESS_|48|')).
fof(f32227,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) | $spl540),
  inference(cnf_transformation,[],[f32227_D])).
fof(f32227_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f32194,plain,(
  ( ! [X184,X183] : (p(s(bool,X183)) | ~p(s(bool,X184)) | s(bool,X184) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) | $spl21),
  inference(resolution,[],[f1722,f570])).
fof(f32095,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17))))) | s(bool,X26) != s(bool,X28)) ) | $spl21),
  inference(resolution,[],[f1722,f375])).
fof(f32094,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) | $spl21),
  inference(resolution,[],[f1722,f374])).
fof(f32092,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) | $spl21),
  inference(resolution,[],[f1722,f373])).
fof(f32090,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17))))) | s(bool,X11) != s(bool,X13)) ) | $spl21),
  inference(resolution,[],[f1722,f371])).
fof(f32088,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) ) | $spl21),
  inference(resolution,[],[f1722,f369])).
fof(f32032,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl538),
  inference(cnf_transformation,[],[f32032_D])).
fof(f32032_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f31992,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f570])).
fof(f31895,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X22) != s(bool,X23)) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f375])).
fof(f31892,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f374])).
fof(f31891,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f373])).
fof(f31890,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) != s(bool,X8)) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f371])).
fof(f31888,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl18 | $spl226)),
  inference(resolution,[],[f31710,f369])).
fof(f31854,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl536),
  inference(cnf_transformation,[],[f31854_D])).
fof(f31854_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f31819,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl52),
  inference(resolution,[],[f19847,f570])).
fof(f31722,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl52),
  inference(resolution,[],[f19847,f375])).
fof(f31719,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl52),
  inference(resolution,[],[f19847,f374])).
fof(f31718,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl52),
  inference(resolution,[],[f19847,f373])).
fof(f31717,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl52),
  inference(resolution,[],[f19847,f371])).
fof(f31715,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13))))) ) | $spl52),
  inference(resolution,[],[f19847,f369])).
fof(f31697,plain,(
  $false | ($spl21 | $spl46 | $spl414)),
  inference(subsumption_resolution,[],[f31696,f2532])).
fof(f31696,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ($spl21 | $spl414)),
  inference(forward_demodulation,[],[f1722,f19900])).
fof(f31183,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) | $spl534),
  inference(cnf_transformation,[],[f31183_D])).
fof(f31183_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f31148,plain,(
  ( ! [X182,X181] : (p(s(bool,X181)) | p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) | $spl30),
  inference(resolution,[],[f2504,f570])).
fof(f31051,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))) | s(bool,X25) != s(bool,X26)) ) | $spl30),
  inference(resolution,[],[f2504,f375])).
fof(f31048,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) | $spl30),
  inference(resolution,[],[f2504,f374])).
fof(f31047,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) | $spl30),
  inference(resolution,[],[f2504,f373])).
fof(f31046,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))) | s(bool,X10) != s(bool,X11)) ) | $spl30),
  inference(resolution,[],[f2504,f371])).
fof(f31044,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) ) | $spl30),
  inference(resolution,[],[f2504,f369])).
fof(f30901,plain,(
  $false | ($spl31 | $spl34 | $spl226)),
  inference(subsumption_resolution,[],[f30900,f557])).
fof(f557,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f39])).
fof(f39,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_trans','thm.bool.TRUTH')).
fof(f30900,plain,(
  ~p(s(bool,'T_0')) | ($spl31 | $spl34 | $spl226)),
  inference(forward_demodulation,[],[f30899,f686])).
fof(f30899,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | ($spl31 | $spl34 | $spl226)),
  inference(subsumption_resolution,[],[f30898,f575])).
fof(f575,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f503])).
fof(f503,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X17] : (s(X6,X17) = s(X6,X17) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_trans','thm.bool.REFL_CLAUSE')).
fof(f30898,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | ($spl31 | $spl34 | $spl226)),
  inference(forward_demodulation,[],[f30897,f21770])).
fof(f21770,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14441,f8756])).
fof(f14441,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl34),
  inference(subsumption_resolution,[],[f14429,f555])).
fof(f14429,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl34),
  inference(superposition,[],[f2508,f569])).
fof(f569,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.LIST_REL_trans',aHL_BOOL_CASES)).
fof(f30897,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | ($spl31 | $spl226)),
  inference(forward_demodulation,[],[f30755,f8756])).
fof(f30755,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | $spl31),
  inference(resolution,[],[f2505,f326])).
fof(f28443,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl532),
  inference(cnf_transformation,[],[f28443_D])).
fof(f28443_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f28406,plain,(
  ( ! [X182,X181] : (p(s(bool,X181)) | p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl46),
  inference(resolution,[],[f2532,f570])).
fof(f28309,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))))) | s(bool,X25) != s(bool,X26)) ) | $spl46),
  inference(resolution,[],[f2532,f375])).
fof(f28306,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl46),
  inference(resolution,[],[f2532,f374])).
fof(f28305,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl46),
  inference(resolution,[],[f2532,f373])).
fof(f28304,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))))) | s(bool,X10) != s(bool,X11)) ) | $spl46),
  inference(resolution,[],[f2532,f371])).
fof(f28302,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl46),
  inference(resolution,[],[f2532,f369])).
fof(f28289,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) | $spl530),
  inference(cnf_transformation,[],[f28289_D])).
fof(f28289_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f28252,plain,(
  ( ! [X182,X181] : (p(s(bool,X181)) | p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) | $spl44),
  inference(resolution,[],[f2529,f570])).
fof(f2529,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | $spl44),
  inference(cnf_transformation,[],[f2529_D])).
fof(f2529_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f28155,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))) | s(bool,X25) != s(bool,X26)) ) | $spl44),
  inference(resolution,[],[f2529,f375])).
fof(f28152,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) | $spl44),
  inference(resolution,[],[f2529,f374])).
fof(f28151,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) | $spl44),
  inference(resolution,[],[f2529,f373])).
fof(f28150,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))) | s(bool,X10) != s(bool,X11)) ) | $spl44),
  inference(resolution,[],[f2529,f371])).
fof(f28148,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) ) | $spl44),
  inference(resolution,[],[f2529,f369])).
fof(f28133,plain,(
  $false | ($spl45 | $spl226)),
  inference(subsumption_resolution,[],[f28132,f557])).
fof(f28132,plain,(
  ~p(s(bool,'T_0')) | ($spl45 | $spl226)),
  inference(forward_demodulation,[],[f28131,f686])).
fof(f28131,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | ($spl45 | $spl226)),
  inference(subsumption_resolution,[],[f28130,f560])).
fof(f28130,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | ($spl45 | $spl226)),
  inference(forward_demodulation,[],[f27987,f8756])).
fof(f27987,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) | $spl45),
  inference(resolution,[],[f2530,f326])).
fof(f27980,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl528),
  inference(cnf_transformation,[],[f27980_D])).
fof(f27980_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f27942,plain,(
  ( ! [X182,X181] : (p(s(bool,X181)) | p(s(bool,X182)) | s(bool,X182) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl40),
  inference(resolution,[],[f2516,f570])).
fof(f27845,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))))) | s(bool,X25) != s(bool,X26)) ) | $spl40),
  inference(resolution,[],[f2516,f375])).
fof(f27842,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl40),
  inference(resolution,[],[f2516,f374])).
fof(f27841,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl40),
  inference(resolution,[],[f2516,f373])).
fof(f27840,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17))))) | s(bool,X10) != s(bool,X11)) ) | $spl40),
  inference(resolution,[],[f2516,f371])).
fof(f27838,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) ) | $spl40),
  inference(resolution,[],[f2516,f369])).
fof(f2530,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | $spl45),
  inference(cnf_transformation,[],[f2530_D])).
fof(f2530_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f2532,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | $spl46),
  inference(cnf_transformation,[],[f2532_D])).
fof(f2532_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f27974,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | ($spl40 | $spl226)),
  inference(subsumption_resolution,[],[f27831,f560])).
fof(f27831,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | ($spl40 | $spl226)),
  inference(resolution,[],[f2516,f20757])).
fof(f20757,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(backward_demodulation,[],[f8756,f321])).
fof(f27821,plain,(
  $false | ($spl41 | $spl226)),
  inference(subsumption_resolution,[],[f27820,f557])).
fof(f27820,plain,(
  ~p(s(bool,'T_0')) | ($spl41 | $spl226)),
  inference(forward_demodulation,[],[f27819,f835])).
fof(f835,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f323,f545])).
fof(f27819,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | ($spl41 | $spl226)),
  inference(subsumption_resolution,[],[f27818,f560])).
fof(f27818,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | ($spl41 | $spl226)),
  inference(forward_demodulation,[],[f27675,f20758])).
fof(f27675,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | $spl41),
  inference(resolution,[],[f2517,f326])).
fof(f25703,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl525),
  inference(cnf_transformation,[],[f25703_D])).
fof(f25703_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl525),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl525])])).
fof(f25681,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X10))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21752,f10572])).
fof(f21752,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X123)) ) | ($spl226 | $spl349)),
  inference(forward_demodulation,[],[f14708,f8756])).
fof(f14708,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X123)) ) | $spl349),
  inference(resolution,[],[f14049,f436])).
fof(f25704,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl526),
  inference(cnf_transformation,[],[f25704_D])).
fof(f25704_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f25680,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21752,f10444])).
fof(f25701,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl522),
  inference(cnf_transformation,[],[f25701_D])).
fof(f25701_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f25679,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21752,f10519])).
fof(f25475,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl520),
  inference(cnf_transformation,[],[f25475_D])).
fof(f25475_D,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f25471,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21795,f568])).
fof(f21795,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X46)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14309,f8756])).
fof(f14309,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X46)) ) | $spl34),
  inference(resolution,[],[f2508,f385])).
fof(f25452,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl518),
  inference(cnf_transformation,[],[f25452_D])).
fof(f25452_D,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f25448,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21794,f568])).
fof(f21794,plain,(
  ( ! [X72,X71] : (sP2(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14322,f8756])).
fof(f14322,plain,(
  ( ! [X72,X71] : (sP2(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f401])).
fof(f25431,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl516),
  inference(cnf_transformation,[],[f25431_D])).
fof(f25431_D,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f25427,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21793,f568])).
fof(f21793,plain,(
  ( ! [X74,X73] : (sP2(X73,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X74)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14323,f8756])).
fof(f14323,plain,(
  ( ! [X74,X73] : (sP2(X73,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X74)) ) | $spl34),
  inference(resolution,[],[f2508,f402])).
fof(f25385,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl514),
  inference(cnf_transformation,[],[f25385_D])).
fof(f25385_D,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f25380,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21792,f568])).
fof(f21792,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14339,f8756])).
fof(f14339,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f418])).
fof(f25361,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl512),
  inference(cnf_transformation,[],[f25361_D])).
fof(f25361_D,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f25355,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21791,f568])).
fof(f21791,plain,(
  ( ! [X132,X131] : (sP8(X131,X132,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14352,f8756])).
fof(f14352,plain,(
  ( ! [X132,X131] : (sP8(X131,X132,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f435])).
fof(f25337,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl510),
  inference(cnf_transformation,[],[f25337_D])).
fof(f25337_D,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f25331,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21790,f568])).
fof(f21790,plain,(
  ( ! [X138,X137] : (sP7(X137,X138,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14355,f8756])).
fof(f14355,plain,(
  ( ! [X138,X137] : (sP7(X137,X138,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f439])).
fof(f25313,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl508),
  inference(cnf_transformation,[],[f25313_D])).
fof(f25313_D,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f25308,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21789,f568])).
fof(f21789,plain,(
  ( ! [X140,X139] : (sP6(X139,X140,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14356,f8756])).
fof(f14356,plain,(
  ( ! [X140,X139] : (sP6(X139,X140,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f443])).
fof(f25293,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl506),
  inference(cnf_transformation,[],[f25293_D])).
fof(f25293_D,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f25288,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21788,f568])).
fof(f21788,plain,(
  ( ! [X142,X141] : (sP6(X141,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X142)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14357,f8756])).
fof(f14357,plain,(
  ( ! [X142,X141] : (sP6(X141,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X142)) ) | $spl34),
  inference(resolution,[],[f2508,f444])).
fof(f25251,plain,(
  ( ! [X6,X7] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X6,X7)) ) | $spl504),
  inference(cnf_transformation,[],[f25251_D])).
fof(f25251_D,plain,(
  ( ! [X6,X7] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X6,X7)) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f25243,plain,(
  ( ! [X6,X7] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21787,f568])).
fof(f21787,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X143,X144)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14358,f8756])).
fof(f14358,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X143,X144)) ) | $spl34),
  inference(resolution,[],[f2508,f445])).
fof(f25250,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl503),
  inference(cnf_transformation,[],[f25250_D])).
fof(f25250_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl503),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl503])])).
fof(f25233,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21787,f10520])).
fof(f25228,plain,(
  ( ! [X6,X7] : (sP10(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl500),
  inference(cnf_transformation,[],[f25228_D])).
fof(f25228_D,plain,(
  ( ! [X6,X7] : (sP10(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f25225,plain,(
  ( ! [X6,X7] : (sP10(X6,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21786,f568])).
fof(f21786,plain,(
  ( ! [X161,X162] : (sP10(X161,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X162)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14367,f8756])).
fof(f14367,plain,(
  ( ! [X161,X162] : (sP10(X161,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X162)) ) | $spl34),
  inference(resolution,[],[f2508,f462])).
fof(f25205,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) | $spl499),
  inference(cnf_transformation,[],[f25205_D])).
fof(f25205_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl499),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl499])])).
fof(f25170,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl349)),
  inference(superposition,[],[f21768,f568])).
fof(f21768,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl349)),
  inference(forward_demodulation,[],[f14049,f8756])).
fof(f25198,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl496),
  inference(cnf_transformation,[],[f25198_D])).
fof(f25198_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f25159,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f570])).
fof(f25060,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X21) != s(bool,X23)) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f375])).
fof(f25059,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f374])).
fof(f25057,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f373])).
fof(f25055,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X6) != s(bool,X8)) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f371])).
fof(f25053,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl349)),
  inference(resolution,[],[f21768,f369])).
fof(f12539,plain,(
  ( ! [X10,X11] : (sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl312),
  inference(cnf_transformation,[],[f12539_D])).
fof(f12539_D,plain,(
  ( ! [X10,X11] : (sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f25048,plain,(
  ( ! [X6,X7] : (sP10(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20799,f568])).
fof(f20799,plain,(
  ( ! [X159,X160] : (sP10(X159,X160,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X159))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9528])).
fof(f9528,plain,(
  ( ! [X159,X160] : (sP10(X159,X160,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X159))) ) | $spl42),
  inference(resolution,[],[f2527,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f273])).
fof(f12523,plain,(
  ( ! [X10,X11] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) | $spl310),
  inference(cnf_transformation,[],[f12523_D])).
fof(f12523_D,plain,(
  ( ! [X10,X11] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f25030,plain,(
  ( ! [X6,X7] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20798,f568])).
fof(f20798,plain,(
  ( ! [X158,X157] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X157,X158) | ~p(s(bool,X158))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9527])).
fof(f9527,plain,(
  ( ! [X158,X157] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X157,X158) | ~p(s(bool,X158))) ) | $spl42),
  inference(resolution,[],[f2527,f461])).
fof(f12489,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) | $spl308),
  inference(cnf_transformation,[],[f12489_D])).
fof(f12489_D,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f24985,plain,(
  ( ! [X6,X7] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20795,f568])).
fof(f20795,plain,(
  ( ! [X92,X91] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X91,X92) | p(s(bool,X91))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9494])).
fof(f9494,plain,(
  ( ! [X92,X91] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X91,X92) | p(s(bool,X91))) ) | $spl42),
  inference(resolution,[],[f2527,f412])).
fof(f12464,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl306),
  inference(cnf_transformation,[],[f12464_D])).
fof(f12464_D,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f24956,plain,(
  ( ! [X6,X7] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20794,f568])).
fof(f20794,plain,(
  ( ! [X90,X89] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X90,X89) | p(s(bool,X89))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9493])).
fof(f9493,plain,(
  ( ! [X90,X89] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X90,X89) | p(s(bool,X89))) ) | $spl42),
  inference(resolution,[],[f2527,f411])).
fof(f12419,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl304),
  inference(cnf_transformation,[],[f12419_D])).
fof(f12419_D,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f24932,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20793,f568])).
fof(f20793,plain,(
  ( ! [X81,X82] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X81,X82) | ~p(s(bool,X81))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9489])).
fof(f9489,plain,(
  ( ! [X81,X82] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X81,X82) | ~p(s(bool,X81))) ) | $spl42),
  inference(resolution,[],[f2527,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f99,f197,f196])).
fof(f99,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X36,X37,X38] : ((p(s(bool,X38)) <=> (p(s(bool,X37)) | p(s(bool,X36)))) <=> ((p(s(bool,X38)) | ~p(s(bool,X37))) & (p(s(bool,X38)) | ~p(s(bool,X36))) & (p(s(bool,X37)) | p(s(bool,X36)) | ~p(s(bool,X38)))))),
  file('thm.list.LIST_REL_trans','thm.sat.sat_25_prove_ltb')).
fof(f12388,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) | $spl302),
  inference(cnf_transformation,[],[f12388_D])).
fof(f12388_D,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f24907,plain,(
  ( ! [X6,X7] : (sP3(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20792,f568])).
fof(f20792,plain,(
  ( ! [X80,X79] : (sP3(X79,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X80) | ~p(s(bool,X79))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9488])).
fof(f9488,plain,(
  ( ! [X80,X79] : (sP3(X79,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X80) | ~p(s(bool,X79))) ) | $spl42),
  inference(resolution,[],[f2527,f406])).
fof(f12367,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) | $spl300),
  inference(cnf_transformation,[],[f12367_D])).
fof(f12367_D,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f24882,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20791,f568])).
fof(f20791,plain,(
  ( ! [X78,X77] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X77,X78) | ~p(s(bool,X78))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9487])).
fof(f9487,plain,(
  ( ! [X78,X77] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X77,X78) | ~p(s(bool,X78))) ) | $spl42),
  inference(resolution,[],[f2527,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f12336,plain,(
  ( ! [X10,X11] : (sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl298),
  inference(cnf_transformation,[],[f12336_D])).
fof(f12336_D,plain,(
  ( ! [X10,X11] : (sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f24837,plain,(
  ( ! [X6,X7] : (sP3(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20790,f568])).
fof(f20790,plain,(
  ( ! [X76,X75] : (sP3(X75,X76,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X75))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9486])).
fof(f9486,plain,(
  ( ! [X76,X75] : (sP3(X75,X76,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X75))) ) | $spl42),
  inference(resolution,[],[f2527,f405])).
fof(f12315,plain,(
  ( ! [X10,X11] : (~sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl296),
  inference(cnf_transformation,[],[f12315_D])).
fof(f12315_D,plain,(
  ( ! [X10,X11] : (~sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f24806,plain,(
  ( ! [X6,X7] : (~sP3(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20789,f568])).
fof(f20789,plain,(
  ( ! [X64,X63] : (~sP3(X63,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X64) | p(s(bool,X63))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9480])).
fof(f9480,plain,(
  ( ! [X64,X63] : (~sP3(X63,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X64) | p(s(bool,X63))) ) | $spl42),
  inference(resolution,[],[f2527,f394])).
fof(f12288,plain,(
  ( ! [X10,X11] : (~sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) | $spl294),
  inference(cnf_transformation,[],[f12288_D])).
fof(f12288_D,plain,(
  ( ! [X10,X11] : (~sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f24776,plain,(
  ( ! [X6,X7] : (~sP3(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20788,f568])).
fof(f20788,plain,(
  ( ! [X61,X62] : (~sP3(X61,X62,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X61))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9479])).
fof(f9479,plain,(
  ( ! [X61,X62] : (~sP3(X61,X62,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X61))) ) | $spl42),
  inference(resolution,[],[f2527,f393])).
fof(f12261,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl292),
  inference(cnf_transformation,[],[f12261_D])).
fof(f12261_D,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f24752,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20787,f568])).
fof(f20787,plain,(
  ( ! [X54,X53] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X53,X54) | ~p(s(bool,X53))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9475])).
fof(f9475,plain,(
  ( ! [X54,X53] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X53,X54) | ~p(s(bool,X53))) ) | $spl42),
  inference(resolution,[],[f2527,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f172,f194,f193])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X36,X37,X38] : ((p(s(bool,X38)) <=> (p(s(bool,X37)) => p(s(bool,X36)))) <=> ((p(s(bool,X38)) | p(s(bool,X37))) & (p(s(bool,X38)) | ~p(s(bool,X36))) & (~p(s(bool,X37)) | p(s(bool,X36)) | ~p(s(bool,X38)))))),
  file('thm.list.LIST_REL_trans','thm.sat.sat_27_prove_ltb')).
fof(f12234,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) | $spl290),
  inference(cnf_transformation,[],[f12234_D])).
fof(f12234_D,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f24724,plain,(
  ( ! [X6,X7] : (sP1(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | ~p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20786,f568])).
fof(f20786,plain,(
  ( ! [X52,X51] : (sP1(X51,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X52) | ~p(s(bool,X51))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9474])).
fof(f9474,plain,(
  ( ! [X52,X51] : (sP1(X51,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X52) | ~p(s(bool,X51))) ) | $spl42),
  inference(resolution,[],[f2527,f389])).
fof(f12215,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl288),
  inference(cnf_transformation,[],[f12215_D])).
fof(f12215_D,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f24677,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X7)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20785,f568])).
fof(f20785,plain,(
  ( ! [X50,X49] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X49,X50) | p(s(bool,X50))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9473])).
fof(f9473,plain,(
  ( ! [X50,X49] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X49,X50) | p(s(bool,X50))) ) | $spl42),
  inference(resolution,[],[f2527,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f12189,plain,(
  ( ! [X10,X11] : (~sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl286),
  inference(cnf_transformation,[],[f12189_D])).
fof(f12189_D,plain,(
  ( ! [X10,X11] : (~sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f24646,plain,(
  ( ! [X6,X7] : (~sP1(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | p(s(bool,X6)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20784,f568])).
fof(f20784,plain,(
  ( ! [X31,X32] : (~sP1(X31,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X32) | p(s(bool,X31))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9464])).
fof(f9464,plain,(
  ( ! [X31,X32] : (~sP1(X31,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X32) | p(s(bool,X31))) ) | $spl42),
  inference(resolution,[],[f2527,f377])).
fof(f6511,plain,(
  ( ! [X8] : (~sP12(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) | $spl150),
  inference(cnf_transformation,[],[f6511_D])).
fof(f6511_D,plain,(
  ( ! [X8] : (~sP12(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f24470,plain,(
  ( ! [X4] : (~sP12(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22424,f568])).
fof(f22424,plain,(
  ( ! [X167] : (~sP12(X167,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X167))) ) | $spl437),
  inference(subsumption_resolution,[],[f22362,f527])).
fof(f22362,plain,(
  ( ! [X167] : (p(s(bool,'F_0')) | p(s(bool,X167)) | ~sP12(X167,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f482])).
fof(f6485,plain,(
  ( ! [X8] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) | $spl148),
  inference(cnf_transformation,[],[f6485_D])).
fof(f6485_D,plain,(
  ( ! [X8] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f24418,plain,(
  ( ! [X4] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22423,f568])).
fof(f6218,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl144),
  inference(cnf_transformation,[],[f6218_D])).
fof(f6218_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f23882,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl437),
  inference(superposition,[],[f22371,f568])).
fof(f5891,plain,(
  ( ! [X8] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) | $spl142),
  inference(cnf_transformation,[],[f5891_D])).
fof(f5891_D,plain,(
  ( ! [X8] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f23844,plain,(
  ( ! [X4] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | ~p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22356,f568])).
fof(f22356,plain,(
  ( ! [X161] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X161) | ~p(s(bool,X161))) ) | $spl437),
  inference(resolution,[],[f21460,f471])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP11(X1,X0)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f5869,plain,(
  ( ! [X8] : (sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) | $spl140),
  inference(cnf_transformation,[],[f5869_D])).
fof(f5869_D,plain,(
  ( ! [X8] : (sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f23821,plain,(
  ( ! [X4] : (sP11(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22355,f568])).
fof(f22355,plain,(
  ( ! [X160] : (sP11(X160,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X160))) ) | $spl437),
  inference(resolution,[],[f21460,f468])).
fof(f468,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f5847,plain,(
  ( ! [X8] : (~sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) | $spl138),
  inference(cnf_transformation,[],[f5847_D])).
fof(f5847_D,plain,(
  ( ! [X8] : (~sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f23796,plain,(
  ( ! [X4] : (~sP11(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22354,f568])).
fof(f22354,plain,(
  ( ! [X159] : (~sP11(X159,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X159))) ) | $spl437),
  inference(resolution,[],[f21460,f466])).
fof(f466,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f278])).
fof(f5827,plain,(
  ( ! [X8] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) | $spl136),
  inference(cnf_transformation,[],[f5827_D])).
fof(f5827_D,plain,(
  ( ! [X8] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f23773,plain,(
  ( ! [X4] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl437),
  inference(superposition,[],[f22353,f568])).
fof(f22353,plain,(
  ( ! [X158] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X158) | p(s(bool,X158))) ) | $spl437),
  inference(resolution,[],[f21460,f466])).
fof(f6450,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl147),
  inference(cnf_transformation,[],[f6450_D])).
fof(f6450_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f23729,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl451),
  inference(superposition,[],[f21655,f568])).
fof(f23706,plain,(
  ( ! [X4] : (sP12(X4,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl494),
  inference(cnf_transformation,[],[f23706_D])).
fof(f23706_D,plain,(
  ( ! [X4] : (sP12(X4,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f23701,plain,(
  ( ! [X4] : (sP12(X4,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21785,f568])).
fof(f21785,plain,(
  ( ! [X175] : (sP12(X175,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14378,f8756])).
fof(f14378,plain,(
  ( ! [X175] : (sP12(X175,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f483])).
fof(f23672,plain,(
  ( ! [X4] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl492),
  inference(cnf_transformation,[],[f23672_D])).
fof(f23672_D,plain,(
  ( ! [X4] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f23667,plain,(
  ( ! [X4] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.|48|_0')),X4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl34 | $spl226)),
  inference(superposition,[],[f21784,f568])).
fof(f21784,plain,(
  ( ! [X176] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X176)) ) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f14379,f8756])).
fof(f14379,plain,(
  ( ! [X176] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X176)) ) | $spl34),
  inference(resolution,[],[f2508,f484])).
fof(f4700,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl118),
  inference(cnf_transformation,[],[f4700_D])).
fof(f4700_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f23222,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl42 | $spl226)),
  inference(superposition,[],[f20804,f568])).
fof(f20804,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9558])).
fof(f9558,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl42),
  inference(resolution,[],[f2527,f545])).
fof(f10726,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) | $spl284),
  inference(cnf_transformation,[],[f10726_D])).
fof(f10726_D,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f23207,plain,(
  ( ! [X4] : (sP11(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20803,f568])).
fof(f20803,plain,(
  ( ! [X170] : (sP11(X170,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X170))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9535])).
fof(f9535,plain,(
  ( ! [X170] : (sP11(X170,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X170))) ) | $spl42),
  inference(resolution,[],[f2527,f471])).
fof(f10701,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) | $spl282),
  inference(cnf_transformation,[],[f10701_D])).
fof(f10701_D,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f23183,plain,(
  ( ! [X4] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20802,f568])).
fof(f20802,plain,(
  ( ! [X169] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X169) | p(s(bool,X169))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9534])).
fof(f9534,plain,(
  ( ! [X169] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X169) | p(s(bool,X169))) ) | $spl42),
  inference(resolution,[],[f2527,f468])).
fof(f10676,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) | $spl280),
  inference(cnf_transformation,[],[f10676_D])).
fof(f10676_D,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f23161,plain,(
  ( ! [X4] : (~sP11(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20801,f568])).
fof(f20801,plain,(
  ( ! [X168] : (~sP11(X168,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X168))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9533])).
fof(f9533,plain,(
  ( ! [X168] : (~sP11(X168,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X168))) ) | $spl42),
  inference(resolution,[],[f2527,f467])).
fof(f467,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f278])).
fof(f10655,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) | $spl278),
  inference(cnf_transformation,[],[f10655_D])).
fof(f10655_D,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f23121,plain,(
  ( ! [X4] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | ~p(s(bool,X4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20800,f568])).
fof(f20800,plain,(
  ( ! [X167] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X167) | ~p(s(bool,X167))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f9532])).
fof(f9532,plain,(
  ( ! [X167] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X167) | ~p(s(bool,X167))) ) | $spl42),
  inference(resolution,[],[f2527,f467])).
fof(f23094,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13)))) | $spl491),
  inference(cnf_transformation,[],[f23094_D])).
fof(f23094_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13)))) <=> ~$spl491),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl491])])).
fof(f23082,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl356),
  inference(resolution,[],[f22778,f10520])).
fof(f22778,plain,(
  ( ! [X136,X137] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))),X136,X137)) ) | $spl356),
  inference(resolution,[],[f17374,f445])).
fof(f17374,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) | $spl356),
  inference(superposition,[],[f560,f14241])).
fof(f22853,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl488),
  inference(cnf_transformation,[],[f22853_D])).
fof(f22853_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f22818,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl356),
  inference(resolution,[],[f17374,f570])).
fof(f22721,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13)))) | s(bool,X22) != s(bool,X23)) ) | $spl356),
  inference(resolution,[],[f17374,f375])).
fof(f22718,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl356),
  inference(resolution,[],[f17374,f374])).
fof(f22717,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl356),
  inference(resolution,[],[f17374,f373])).
fof(f22716,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13)))) | s(bool,X7) != s(bool,X8)) ) | $spl356),
  inference(resolution,[],[f17374,f371])).
fof(f22714,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK29(sK17,sK16,sK14,sK13))))) ) | $spl356),
  inference(resolution,[],[f17374,f369])).
fof(f22668,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl486),
  inference(cnf_transformation,[],[f22668_D])).
fof(f22668_D,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f22658,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl374)),
  inference(superposition,[],[f21671,f568])).
fof(f21671,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | ($spl226 | $spl374)),
  inference(forward_demodulation,[],[f15273,f8756])).
fof(f22661,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl226 | $spl374 | $spl437)),
  inference(subsumption_resolution,[],[f22644,f21460])).
fof(f22644,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X1))) ) | ($spl226 | $spl374)),
  inference(resolution,[],[f21671,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f22643,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl484),
  inference(cnf_transformation,[],[f22643_D])).
fof(f22643_D,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f22635,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl376)),
  inference(superposition,[],[f21670,f568])).
fof(f21670,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl376)),
  inference(forward_demodulation,[],[f15288,f8756])).
fof(f22638,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl376 | $spl437)),
  inference(subsumption_resolution,[],[f22623,f21460])).
fof(f22623,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl376)),
  inference(resolution,[],[f21670,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f22607,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl482),
  inference(cnf_transformation,[],[f22607_D])).
fof(f22607_D,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f22599,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl378)),
  inference(superposition,[],[f21669,f568])).
fof(f21669,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | ($spl226 | $spl378)),
  inference(forward_demodulation,[],[f15332,f8756])).
fof(f22602,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl378 | $spl437)),
  inference(subsumption_resolution,[],[f22587,f21460])).
fof(f22587,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl378)),
  inference(resolution,[],[f21669,f400])).
fof(f22586,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl480),
  inference(cnf_transformation,[],[f22586_D])).
fof(f22586_D,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f22575,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl380)),
  inference(superposition,[],[f21668,f568])).
fof(f21668,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl380)),
  inference(forward_demodulation,[],[f15348,f8756])).
fof(f22578,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl226 | $spl380 | $spl437)),
  inference(subsumption_resolution,[],[f22562,f21460])).
fof(f22562,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl380)),
  inference(resolution,[],[f21668,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f254])).
fof(f22561,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl478),
  inference(cnf_transformation,[],[f22561_D])).
fof(f22561_D,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f22550,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl382)),
  inference(superposition,[],[f21667,f568])).
fof(f21667,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl382)),
  inference(forward_demodulation,[],[f15363,f8756])).
fof(f22553,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl226 | $spl382 | $spl437)),
  inference(subsumption_resolution,[],[f22538,f21460])).
fof(f22538,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl382)),
  inference(resolution,[],[f21667,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f262])).
fof(f22537,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl476),
  inference(cnf_transformation,[],[f22537_D])).
fof(f22537_D,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f22526,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl384)),
  inference(superposition,[],[f21666,f568])).
fof(f21666,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl384)),
  inference(forward_demodulation,[],[f15378,f8756])).
fof(f22529,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl226 | $spl384 | $spl437)),
  inference(subsumption_resolution,[],[f22514,f21460])).
fof(f22514,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl384)),
  inference(resolution,[],[f21666,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f265])).
fof(f22513,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl474),
  inference(cnf_transformation,[],[f22513_D])).
fof(f22513_D,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f22503,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl386)),
  inference(superposition,[],[f21665,f568])).
fof(f21665,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl386)),
  inference(forward_demodulation,[],[f15395,f8756])).
fof(f22506,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl386 | $spl437)),
  inference(subsumption_resolution,[],[f22492,f21460])).
fof(f22492,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl386)),
  inference(resolution,[],[f21665,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f22491,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) | $spl472),
  inference(cnf_transformation,[],[f22491_D])).
fof(f22491_D,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7)) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f22482,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl388)),
  inference(superposition,[],[f21664,f568])).
fof(f21664,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | ($spl226 | $spl388)),
  inference(forward_demodulation,[],[f15434,f8756])).
fof(f22485,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl388 | $spl437)),
  inference(subsumption_resolution,[],[f22472,f21460])).
fof(f22472,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl388)),
  inference(resolution,[],[f21664,f442])).
fof(f22471,plain,(
  ( ! [X6,X7] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7)) ) | $spl470),
  inference(cnf_transformation,[],[f22471_D])).
fof(f22471_D,plain,(
  ( ! [X6,X7] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7)) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f22463,plain,(
  ( ! [X6,X7] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl390)),
  inference(superposition,[],[f21663,f568])).
fof(f21663,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | ($spl226 | $spl390)),
  inference(forward_demodulation,[],[f15451,f8756])).
fof(f21652,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl449),
  inference(cnf_transformation,[],[f21652_D])).
fof(f21652_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl449),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl449])])).
fof(f22456,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl390)),
  inference(resolution,[],[f21663,f10520])).
fof(f22466,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl390 | $spl437)),
  inference(subsumption_resolution,[],[f22452,f21460])).
fof(f22452,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl226 | $spl390)),
  inference(resolution,[],[f21663,f442])).
fof(f22451,plain,(
  $false | ($spl226 | $spl392 | $spl437)),
  inference(subsumption_resolution,[],[f22436,f21460])).
fof(f22436,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl392)),
  inference(resolution,[],[f21662,f1653])).
fof(f1653,plain,(
  ( ! [X146,X147] : (~sP10('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)),X146,X147) | p(s(bool,X146))) )),
  inference(resolution,[],[f324,f457])).
fof(f21662,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | ($spl226 | $spl392)),
  inference(forward_demodulation,[],[f15466,f8756])).
fof(f22450,plain,(
  $false | ($spl226 | $spl392 | $spl437)),
  inference(subsumption_resolution,[],[f22435,f21460])).
fof(f22435,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl226 | $spl392)),
  inference(resolution,[],[f21662,f1651])).
fof(f1651,plain,(
  ( ! [X142,X143] : (~sP10(X143,X142,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | p(s(bool,X142))) )),
  inference(resolution,[],[f324,f456])).
fof(f22392,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl437),
  inference(superposition,[],[f21460,f568])).
fof(f22420,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl468),
  inference(cnf_transformation,[],[f22420_D])).
fof(f22420_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f22384,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f570])).
fof(f22285,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X21) != s(bool,X23)) ) | $spl437),
  inference(resolution,[],[f21460,f375])).
fof(f22284,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f374])).
fof(f22282,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f373])).
fof(f22280,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X6) != s(bool,X8)) ) | $spl437),
  inference(resolution,[],[f21460,f371])).
fof(f22278,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl437),
  inference(resolution,[],[f21460,f369])).
fof(f22272,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f22000,f568])).
fof(f22000,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl57 | $spl226)),
  inference(forward_demodulation,[],[f3780,f8756])).
fof(f21460,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl437),
  inference(cnf_transformation,[],[f21460_D])).
fof(f21460_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl437),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl437])])).
fof(f22265,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X10))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f22000,f597])).
fof(f597,plain,(
  ( ! [X39,X40] : (~sP0('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X39,X40) | ~p(s(bool,X40)) | p(s(bool,X39))) )),
  inference(resolution,[],[f322,f383])).
fof(f22277,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X9))) ) | $spl466),
  inference(cnf_transformation,[],[f22277_D])).
fof(f22277_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X9))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f22264,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f22000,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f22260,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f22000,f383])).
fof(f22209,plain,(
  ( ! [X6,X7] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20781,f568])).
fof(f20781,plain,(
  ( ! [X130,X131] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X130,X131)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3827])).
fof(f22202,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20781,f644])).
fof(f644,plain,(
  ( ! [X134,X133] : (~sP7(X133,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X134) | p(s(bool,X134)) | ~p(s(bool,X133))) )),
  inference(resolution,[],[f322,f438])).
fof(f21459,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl436),
  inference(cnf_transformation,[],[f21459_D])).
fof(f21459_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f22213,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl464),
  inference(cnf_transformation,[],[f22213_D])).
fof(f22213_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f22200,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4)) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20781,f432])).
fof(f22198,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20781,f438])).
fof(f22190,plain,(
  ( ! [X6,X7] : (sP7(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20780,f568])).
fof(f20780,plain,(
  ( ! [X128,X129] : (sP7(X128,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X129)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3826])).
fof(f22183,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20780,f643])).
fof(f643,plain,(
  ( ! [X132,X131] : (~sP7('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X131,X132) | p(s(bool,X132)) | ~p(s(bool,X131))) )),
  inference(resolution,[],[f322,f438])).
fof(f22195,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl462),
  inference(cnf_transformation,[],[f22195_D])).
fof(f22195_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f22194,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22193,f20778])).
fof(f20778,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X123)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3823])).
fof(f22193,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22180,f445])).
fof(f22180,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3) | ~sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20780,f431])).
fof(f22179,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20780,f438])).
fof(f22152,plain,(
  ( ! [X6,X7] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20779,f568])).
fof(f20779,plain,(
  ( ! [X125,X124] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X124,X125)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3824])).
fof(f22145,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f641])).
fof(f641,plain,(
  ( ! [X127,X128] : (~sP8(X127,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X128) | p(s(bool,X128)) | ~p(s(bool,X127))) )),
  inference(resolution,[],[f322,f434])).
fof(f22143,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f10459])).
fof(f10459,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f10453,f540])).
fof(f10453,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP8(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) )),
  inference(resolution,[],[f10450,f448])).
fof(f1174,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,X53)) ) | $spl14),
  inference(cnf_transformation,[],[f1174_D])).
fof(f1174_D,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,X53)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f22142,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f10572])).
fof(f9247,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) | $spl232),
  inference(cnf_transformation,[],[f9247_D])).
fof(f9247_D,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f22140,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,X8) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f10519])).
fof(f22157,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl460),
  inference(cnf_transformation,[],[f22157_D])).
fof(f22157_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f22139,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f433])).
fof(f22136,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20779,f434])).
fof(f22131,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20778,f568])).
fof(f22124,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f640])).
fof(f640,plain,(
  ( ! [X125,X126] : (~sP8('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X125,X126) | p(s(bool,X126)) | ~p(s(bool,X125))) )),
  inference(resolution,[],[f322,f434])).
fof(f22122,plain,(
  ( ! [X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X11))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10459])).
fof(f21655,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl451),
  inference(cnf_transformation,[],[f21655_D])).
fof(f21655_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl451),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl451])])).
fof(f22121,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X10))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10572])).
fof(f21656,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl452),
  inference(cnf_transformation,[],[f21656_D])).
fof(f21656_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f22120,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10444])).
fof(f21653,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl448),
  inference(cnf_transformation,[],[f21653_D])).
fof(f21653_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f22119,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10519])).
fof(f21650,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2)) ) | $spl446),
  inference(cnf_transformation,[],[f21650_D])).
fof(f21650_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2)) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f22134,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2)) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22116,f445])).
fof(f22116,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2) | ~sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f431])).
fof(f22115,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f434])).
fof(f22110,plain,(
  ( ! [X6,X7] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20777,f568])).
fof(f20777,plain,(
  ( ! [X99,X98] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X98,X99)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3811])).
fof(f22103,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f628])).
fof(f628,plain,(
  ( ! [X101,X102] : (~sP4(X101,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X102) | p(s(bool,X102)) | ~p(s(bool,X101))) )),
  inference(resolution,[],[f322,f417])).
fof(f21627,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3) | ~p(s(bool,X3))) ) | $spl444),
  inference(cnf_transformation,[],[f21627_D])).
fof(f21627_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3) | ~p(s(bool,X3))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f22099,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X3))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f252])).
fof(f22098,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f417])).
fof(f22093,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20776,f568])).
fof(f20776,plain,(
  ( ! [X97,X96] : (sP4(X96,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X97)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3810])).
fof(f22086,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f627])).
fof(f627,plain,(
  ( ! [X99,X100] : (~sP4('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)),X99,X100) | p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f322,f417])).
fof(f21609,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X3))) ) | $spl442),
  inference(cnf_transformation,[],[f21609_D])).
fof(f21609_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X3))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f22082,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f413])).
fof(f22081,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f417])).
fof(f22076,plain,(
  ( ! [X6,X7] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20775,f568])).
fof(f20775,plain,(
  ( ! [X70,X71] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X70,X71)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3797])).
fof(f21478,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) | $spl440),
  inference(cnf_transformation,[],[f21478_D])).
fof(f21478_D,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f22080,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22068,f406])).
fof(f22068,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f22079,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22067,f405])).
fof(f22067,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X5))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f22066,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f22065,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f400])).
fof(f22042,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20774,f568])).
fof(f20774,plain,(
  ( ! [X39,X38] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X38,X39)) ) | ($spl57 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3781])).
fof(f22035,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f598])).
fof(f598,plain,(
  ( ! [X41,X42] : (~sP0(X41,X42,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))) | p(s(bool,X42)) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f322,f383])).
fof(f21461,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) | $spl438),
  inference(cnf_transformation,[],[f21461_D])).
fof(f21461_D,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f22046,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22033,f389])).
fof(f22033,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f22045,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f22032,f388])).
fof(f22032,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X5))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f22031,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f22030,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f383])).
fof(f22029,plain,(
  ( ! [X4] : (sP12(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl458),
  inference(cnf_transformation,[],[f22029_D])).
fof(f22029_D,plain,(
  ( ! [X4] : (sP12(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f22024,plain,(
  ( ! [X4] : (sP12(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl358)),
  inference(superposition,[],[f21675,f568])).
fof(f21675,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl226 | $spl358)),
  inference(forward_demodulation,[],[f14242,f8756])).
fof(f22028,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl358)),
  inference(subsumption_resolution,[],[f22016,f527])).
fof(f22016,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl358)),
  inference(resolution,[],[f21675,f482])).
fof(f22027,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl358)),
  inference(subsumption_resolution,[],[f22015,f527])).
fof(f22015,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl358)),
  inference(resolution,[],[f21675,f486])).
fof(f486,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f178,f211])).
fof(f178,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X12,X6] : ((~(p(s(bool,X6)) | p(s(bool,X12))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X12)) => p(s(bool,'F_0')))))),
  file('thm.list.LIST_REL_trans','thm.sat.sat_9_prove_ltb')).
fof(f21730,plain,(
  ( ! [X4] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl454),
  inference(cnf_transformation,[],[f21730_D])).
fof(f21730_D,plain,(
  ( ! [X4] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f22010,plain,(
  ( ! [X4] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl360)),
  inference(superposition,[],[f21674,f568])).
fof(f21674,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X0)) ) | ($spl226 | $spl360)),
  inference(forward_demodulation,[],[f14257,f8756])).
fof(f22014,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl360)),
  inference(subsumption_resolution,[],[f22002,f527])).
fof(f22002,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl360)),
  inference(resolution,[],[f21674,f482])).
fof(f22013,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl226 | $spl360)),
  inference(subsumption_resolution,[],[f22001,f527])).
fof(f22001,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl226 | $spl360)),
  inference(resolution,[],[f21674,f486])).
fof(f21962,plain,(
  ( ! [X4] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl56 | $spl226)),
  inference(superposition,[],[f21707,f568])).
fof(f21707,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X173)) ) | ($spl56 | $spl226)),
  inference(forward_demodulation,[],[f15976,f8756])).
fof(f15976,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X173)) ) | $spl56),
  inference(resolution,[],[f3290,f484])).
fof(f21966,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl56 | $spl226)),
  inference(subsumption_resolution,[],[f21954,f527])).
fof(f21954,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl56 | $spl226)),
  inference(resolution,[],[f21707,f482])).
fof(f21965,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl56 | $spl226)),
  inference(subsumption_resolution,[],[f21953,f527])).
fof(f21953,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl56 | $spl226)),
  inference(resolution,[],[f21707,f486])).
fof(f21945,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl456),
  inference(cnf_transformation,[],[f21945_D])).
fof(f21945_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f21905,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f570])).
fof(f21796,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl34 | $spl226)),
  inference(forward_demodulation,[],[f2508,f8756])).
fof(f21808,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X22) != s(bool,X23)) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f375])).
fof(f21805,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f374])).
fof(f21804,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f373])).
fof(f21803,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) != s(bool,X8)) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f371])).
fof(f21801,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl34 | $spl226)),
  inference(resolution,[],[f21796,f369])).
fof(f21739,plain,(
  $false | ($spl356 | $spl400)),
  inference(subsumption_resolution,[],[f21738,f567])).
fof(f21738,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | ($spl356 | $spl400)),
  inference(forward_demodulation,[],[f14241,f15845])).
fof(f21737,plain,(
  $false | ($spl356 | $spl400)),
  inference(subsumption_resolution,[],[f17379,f15845])).
fof(f17379,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl356),
  inference(superposition,[],[f567,f14241])).
fof(f21736,plain,(
  $false | ($spl356 | $spl400)),
  inference(subsumption_resolution,[],[f21735,f559])).
fof(f559,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))))),
  file('thm.list.LIST_REL_trans','thm.prim_rec.NOT_LESS_|48|')).
fof(f21735,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl356 | $spl400)),
  inference(forward_demodulation,[],[f17374,f15845])).
fof(f21725,plain,(
  ( ! [X4] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl226 | $spl348 | $spl400)),
  inference(superposition,[],[f21685,f568])).
fof(f21685,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X173)) ) | ($spl226 | $spl348 | $spl400)),
  inference(forward_demodulation,[],[f21684,f15845])).
fof(f21684,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X173)) ) | ($spl226 | $spl348)),
  inference(forward_demodulation,[],[f15150,f8756])).
fof(f15150,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X173)) ) | $spl348),
  inference(resolution,[],[f14048,f484])).
fof(f14048,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl348),
  inference(cnf_transformation,[],[f14048_D])).
fof(f14048_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f21729,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl348 | $spl400)),
  inference(subsumption_resolution,[],[f21717,f527])).
fof(f21717,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl226 | $spl348 | $spl400)),
  inference(resolution,[],[f21685,f482])).
fof(f21728,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl226 | $spl348 | $spl400)),
  inference(subsumption_resolution,[],[f21716,f527])).
fof(f21716,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl226 | $spl348 | $spl400)),
  inference(resolution,[],[f21685,f486])).
fof(f21646,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20778,f568])).
fof(f21639,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f640])).
fof(f21637,plain,(
  ( ! [X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X11))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10459])).
fof(f21636,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X10))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10572])).
fof(f21635,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10444])).
fof(f21634,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X8))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f10519])).
fof(f21649,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2)) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f21631,f445])).
fof(f21631,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2) | ~sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f431])).
fof(f21630,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20778,f434])).
fof(f21624,plain,(
  ( ! [X6,X7] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20777,f568])).
fof(f21617,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f628])).
fof(f21613,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X3))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f413])).
fof(f21612,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20777,f417])).
fof(f21606,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20776,f568])).
fof(f21599,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f627])).
fof(f21595,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f413])).
fof(f21594,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20776,f417])).
fof(f21475,plain,(
  ( ! [X6,X7] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20775,f568])).
fof(f21480,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f21467,f406])).
fof(f21467,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f398])).
fof(f21479,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f21466,f405])).
fof(f21466,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X5))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f397])).
fof(f21465,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f572])).
fof(f21464,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20775,f400])).
fof(f21456,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(superposition,[],[f20774,f568])).
fof(f21449,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f598])).
fof(f21463,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f21447,f389])).
fof(f21447,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f381])).
fof(f21462,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(subsumption_resolution,[],[f21446,f388])).
fof(f21446,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X5))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f380])).
fof(f21445,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f571])).
fof(f21444,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl57 | $spl226)),
  inference(resolution,[],[f20774,f383])).
fof(f21441,plain,(
  ( ! [X6,X7] : (sP10(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20771,f568])).
fof(f20771,plain,(
  ( ! [X161,X162] : (sP10(X161,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X162)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3071])).
fof(f21421,plain,(
  ( ! [X6,X7] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X6,X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20770,f568])).
fof(f20770,plain,(
  ( ! [X144,X143] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X143,X144)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3062])).
fof(f21428,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl435),
  inference(cnf_transformation,[],[f21428_D])).
fof(f21428_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl435),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl435])])).
fof(f21414,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(resolution,[],[f20770,f10520])).
fof(f21404,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20769,f568])).
fof(f20769,plain,(
  ( ! [X142,X141] : (sP6(X141,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X142)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3061])).
fof(f21275,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20768,f568])).
fof(f20768,plain,(
  ( ! [X140,X139] : (sP6(X139,X140,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3060])).
fof(f21259,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20767,f568])).
fof(f20767,plain,(
  ( ! [X138,X137] : (sP7(X137,X138,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3059])).
fof(f21243,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20766,f568])).
fof(f20766,plain,(
  ( ! [X132,X131] : (sP8(X131,X132,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3056])).
fof(f21226,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20765,f568])).
fof(f20765,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3043])).
fof(f21209,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20764,f568])).
fof(f20764,plain,(
  ( ! [X74,X73] : (sP2(X73,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X74)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3027])).
fof(f2517,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | $spl41),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f21189,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ($spl226 | $spl400)),
  inference(forward_demodulation,[],[f21188,f15845])).
fof(f21188,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl400)),
  inference(forward_demodulation,[],[f21187,f15845])).
fof(f21187,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl400)),
  inference(subsumption_resolution,[],[f21186,f560])).
fof(f21186,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl400)),
  inference(forward_demodulation,[],[f21133,f15845])).
fof(f21133,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl400)),
  inference(superposition,[],[f20757,f15845])).
fof(f21185,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ($spl226 | $spl414)),
  inference(forward_demodulation,[],[f21184,f19900])).
fof(f21184,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl414)),
  inference(forward_demodulation,[],[f21183,f19900])).
fof(f21183,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl414)),
  inference(subsumption_resolution,[],[f21182,f560])).
fof(f21182,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl414)),
  inference(forward_demodulation,[],[f21132,f19900])).
fof(f21132,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK16)))))) | ($spl226 | $spl414)),
  inference(superposition,[],[f20757,f19900])).
fof(f2533,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) | $spl48),
  inference(cnf_transformation,[],[f2533_D])).
fof(f2533_D,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f21178,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) | $spl226),
  inference(subsumption_resolution,[],[f21128,f560])).
fof(f21128,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) ) | $spl226),
  inference(superposition,[],[f20757,f568])).
fof(f21175,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | $spl226),
  inference(forward_demodulation,[],[f21174,f562])).
fof(f562,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.LIST_REL_trans','thm.list.LENGTH|split|0')).
fof(f21174,plain,(
  ( ! [X18] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(forward_demodulation,[],[f21173,f562])).
fof(f21173,plain,(
  ( ! [X18] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(subsumption_resolution,[],[f21172,f560])).
fof(f21172,plain,(
  ( ! [X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(forward_demodulation,[],[f21126,f562])).
fof(f21126,plain,(
  ( ! [X18] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),'const.list.NIL_0'))),s('type.list.list'(sK13),sK16))))))) ) | $spl226),
  inference(superposition,[],[f20757,f562])).
fof(f21169,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17))))))) ) | $spl432),
  inference(cnf_transformation,[],[f21169_D])).
fof(f21169_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17))))))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f21120,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0))))) ) | $spl226),
  inference(superposition,[],[f20757,f568])).
fof(f21164,plain,(
  ( ! [X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))) | ~p(s(bool,X13)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))))) ) | $spl430),
  inference(cnf_transformation,[],[f21164_D])).
fof(f21164_D,plain,(
  ( ! [X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))) | ~p(s(bool,X13)) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17))))))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f21118,plain,(
  ( ! [X296,X295,X297] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X295),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16)))))) | p(s(bool,X296)) | ~p(s(bool,X297)) | s(bool,X297) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f20757,f570])).
fof(f21019,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X42)) ) | $spl226),
  inference(resolution,[],[f20757,f375])).
fof(f21018,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16)))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f20757,f374])).
fof(f21016,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16)))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f20757,f373])).
fof(f21014,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X22)) ) | $spl226),
  inference(resolution,[],[f20757,f371])).
fof(f21012,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) ) | $spl226),
  inference(resolution,[],[f20757,f369])).
fof(f21001,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20763,f568])).
fof(f20763,plain,(
  ( ! [X72,X71] : (sP2(X71,X72,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3026])).
fof(f20985,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X7) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20762,f568])).
fof(f20762,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X46)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3013])).
fof(f20967,plain,(
  ( ! [X4] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20773,f568])).
fof(f20773,plain,(
  ( ! [X176] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X176)) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3083])).
fof(f20951,plain,(
  ( ! [X4] : (sP12(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl42 | $spl226)),
  inference(superposition,[],[f20772,f568])).
fof(f20772,plain,(
  ( ! [X175] : (sP12(X175,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl42 | $spl226)),
  inference(backward_demodulation,[],[f8756,f3082])).
fof(f20689,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | $spl424),
  inference(cnf_transformation,[],[f20689_D])).
fof(f20689_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f20691,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) | $spl426),
  inference(cnf_transformation,[],[f20691_D])).
fof(f20691_D,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f20731,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | $spl224),
  inference(forward_demodulation,[],[f20668,f8754])).
fof(f20668,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) | $spl224),
  inference(superposition,[],[f568,f8754])).
fof(f20687,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | $spl224),
  inference(forward_demodulation,[],[f20630,f8754])).
fof(f20630,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))))) | $spl224),
  inference(superposition,[],[f8754,f568])).
fof(f20684,plain,(
  s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))) | $spl420),
  inference(cnf_transformation,[],[f20684_D])).
fof(f20684_D,plain,(
  s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f20686,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | $spl422),
  inference(cnf_transformation,[],[f20686_D])).
fof(f20686_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f20625,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) | s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))))) | $spl224),
  inference(superposition,[],[f8754,f568])).
fof(f20352,plain,(
  $false | ($spl42 | $spl60)),
  inference(subsumption_resolution,[],[f20221,f20107])).
fof(f20107,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | ($spl42 | $spl60)),
  inference(subsumption_resolution,[],[f20106,f547])).
fof(f547,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.EQ_CLAUSES|split|1')).
fof(f20106,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3))) ) | ($spl42 | $spl60)),
  inference(forward_demodulation,[],[f3293,f9558])).
fof(f20221,plain,(
  ( ! [X285] : (p(s(bool,X285))) ) | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f522])).
fof(f522,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CLAUSES|split|1')).
fof(f20351,plain,(
  $false | ($spl42 | $spl60)),
  inference(subsumption_resolution,[],[f20219,f20107])).
fof(f20219,plain,(
  ( ! [X283] : (p(s(bool,X283))) ) | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f519])).
fof(f519,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X21] : ((p(s(bool,X21)) => p(s(bool,X21))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CLAUSES|split|3')).
fof(f20350,plain,(
  $false | ($spl42 | $spl60)),
  inference(subsumption_resolution,[],[f20217,f20107])).
fof(f20217,plain,(
  p(s(bool,'F_0')) | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f516])).
fof(f516,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X21] : ((p(s(bool,'F_0')) => p(s(bool,X21))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CLAUSES|split|2')).
fof(f20349,plain,(
  $false | ($spl42 | $spl60)),
  inference(subsumption_resolution,[],[f20216,f20107])).
fof(f20216,plain,(
  ( ! [X281] : (p(s(bool,X281))) ) | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f513])).
fof(f513,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X21] : ((p(s(bool,'T_0')) => p(s(bool,X21))) <=> p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_CLAUSES|split|0')).
fof(f20224,plain,(
  $false | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f556])).
fof(f556,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_trans',aHL_TRUTH)).
fof(f20225,plain,(
  $false | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f557])).
fof(f20226,plain,(
  $false | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f575])).
fof(f20235,plain,(
  $false | ($spl42 | $spl60)),
  inference(resolution,[],[f20107,f560])).
fof(f19960,plain,(
  $false | ($spl53 | $spl415)),
  inference(subsumption_resolution,[],[f19959,f2988])).
fof(f2988,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl53),
  inference(cnf_transformation,[],[f2988_D])).
fof(f2988_D,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f19959,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl415),
  inference(trivial_inequality_removal,[],[f19956])).
fof(f19956,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl415),
  inference(superposition,[],[f19899,f568])).
fof(f19955,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | $spl418),
  inference(cnf_transformation,[],[f19955_D])).
fof(f19955_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f19899,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl415),
  inference(cnf_transformation,[],[f19899_D])).
fof(f19899_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) <=> ~$spl415),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl415])])).
fof(f19950,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | $spl53),
  inference(superposition,[],[f2988,f568])).
fof(f19895,plain,(
  s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) | $spl410),
  inference(cnf_transformation,[],[f19895_D])).
fof(f19895_D,plain,(
  s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f19896,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl413),
  inference(cnf_transformation,[],[f19896_D])).
fof(f19896_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f19944,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) | $spl53),
  inference(superposition,[],[f2988,f568])).
fof(f19900,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl414),
  inference(cnf_transformation,[],[f19900_D])).
fof(f19900_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f19902,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | $spl416),
  inference(cnf_transformation,[],[f19902_D])).
fof(f19902_D,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f19938,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl52),
  inference(forward_demodulation,[],[f19879,f2989])).
fof(f19879,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | $spl52),
  inference(superposition,[],[f568,f2989])).
fof(f19898,plain,(
  s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl52),
  inference(forward_demodulation,[],[f19841,f2989])).
fof(f19841,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))))) | $spl52),
  inference(superposition,[],[f2989,f568])).
fof(f19897,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | $spl412),
  inference(cnf_transformation,[],[f19897_D])).
fof(f19897_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f19836,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) | s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) | $spl52),
  inference(superposition,[],[f2989,f568])).
fof(f19703,plain,(
  $false | ($spl357 | $spl401)),
  inference(subsumption_resolution,[],[f19702,f14240])).
fof(f14240,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl357),
  inference(cnf_transformation,[],[f14240_D])).
fof(f14240_D,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) <=> ~$spl357),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl357])])).
fof(f19702,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl401),
  inference(trivial_inequality_removal,[],[f19699])).
fof(f19699,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl401),
  inference(superposition,[],[f15844,f568])).
fof(f19667,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | $spl408),
  inference(cnf_transformation,[],[f19667_D])).
fof(f19667_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f15844,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl401),
  inference(cnf_transformation,[],[f15844_D])).
fof(f15844_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) <=> ~$spl401),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl401])])).
fof(f19662,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | $spl357),
  inference(superposition,[],[f14240,f568])).
fof(f15840,plain,(
  s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK29(sK17,sK16,sK14,sK13)))))) | $spl396),
  inference(cnf_transformation,[],[f15840_D])).
fof(f15840_D,plain,(
  s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK29(sK17,sK16,sK14,sK13)))))) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f15841,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl399),
  inference(cnf_transformation,[],[f15841_D])).
fof(f15841_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f19656,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK29(sK17,sK16,sK14,sK13)))))) | $spl357),
  inference(superposition,[],[f14240,f568])).
fof(f15061,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl371),
  inference(cnf_transformation,[],[f15061_D])).
fof(f15061_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f18786,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | $spl349),
  inference(resolution,[],[f14708,f10572])).
fof(f15062,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl372),
  inference(cnf_transformation,[],[f15062_D])).
fof(f15062_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f18785,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14708,f10444])).
fof(f15059,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl368),
  inference(cnf_transformation,[],[f15059_D])).
fof(f15059_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f18784,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X8))) ) | $spl349),
  inference(resolution,[],[f14708,f10519])).
fof(f17935,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl125),
  inference(resolution,[],[f17823,f570])).
fof(f17823,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl125),
  inference(trivial_inequality_removal,[],[f17819])).
fof(f17819,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl125),
  inference(superposition,[],[f5461,f547])).
fof(f17836,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl125),
  inference(resolution,[],[f17823,f375])).
fof(f17835,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl125),
  inference(resolution,[],[f17823,f374])).
fof(f17833,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl125),
  inference(resolution,[],[f17823,f373])).
fof(f17831,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl125),
  inference(resolution,[],[f17823,f371])).
fof(f17829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl125),
  inference(resolution,[],[f17823,f369])).
fof(f17828,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl125),
  inference(subsumption_resolution,[],[f17813,f545])).
fof(f17813,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl125),
  inference(superposition,[],[f5461,f500])).
fof(f500,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.list.LIST_REL_trans','thm.bool.IMP_ANTISYM_AX')).
fof(f17827,plain,(
  ( ! [X14] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X14))) ) | $spl125),
  inference(subsumption_resolution,[],[f17812,f545])).
fof(f17812,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X14))) ) | $spl125),
  inference(superposition,[],[f5461,f500])).
fof(f17826,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl125),
  inference(subsumption_resolution,[],[f17809,f545])).
fof(f17809,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl125),
  inference(superposition,[],[f5461,f496])).
fof(f496,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13,X14] : (s(bool,X13) = s(bool,X14) <=> ((p(s(bool,X13)) => p(s(bool,X14))) & (p(s(bool,X14)) => p(s(bool,X13)))))),
  file('thm.list.LIST_REL_trans','thm.bool.EQ_IMP_THM')).
fof(f17825,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X10))) ) | $spl125),
  inference(subsumption_resolution,[],[f17808,f545])).
fof(f17808,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X10))) ) | $spl125),
  inference(superposition,[],[f5461,f496])).
fof(f17780,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X4)) ) | $spl57),
  inference(resolution,[],[f3827,f432])).
fof(f17743,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3824,f433])).
fof(f17726,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | $spl57),
  inference(resolution,[],[f3823,f10572])).
fof(f17725,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3823,f10444])).
fof(f17724,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X8))) ) | $spl57),
  inference(resolution,[],[f3823,f10519])).
fof(f14241,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl356),
  inference(cnf_transformation,[],[f14241_D])).
fof(f14241_D,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f17573,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl349),
  inference(superposition,[],[f14049,f568])).
fof(f17631,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl406),
  inference(resolution,[],[f17630,f383])).
fof(f10241,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X9))) ) | $spl256),
  inference(cnf_transformation,[],[f10241_D])).
fof(f10241_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X9))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f17635,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl406),
  inference(resolution,[],[f17630,f382])).
fof(f17646,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl56 | $spl406)),
  inference(subsumption_resolution,[],[f17645,f575])).
fof(f17645,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl56 | $spl406)),
  inference(forward_demodulation,[],[f17631,f15993])).
fof(f15993,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl56),
  inference(resolution,[],[f3290,f545])).
fof(f17630,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl406),
  inference(cnf_transformation,[],[f17630_D])).
fof(f17630_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f17621,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl349),
  inference(superposition,[],[f14665,f568])).
fof(f14665,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f384])).
fof(f14792,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl364),
  inference(cnf_transformation,[],[f14792_D])).
fof(f14792_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f17571,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f570])).
fof(f17472,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl349),
  inference(resolution,[],[f14049,f375])).
fof(f17471,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f374])).
fof(f17469,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f373])).
fof(f17467,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl349),
  inference(resolution,[],[f14049,f371])).
fof(f17465,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f369])).
fof(f15845,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl400),
  inference(cnf_transformation,[],[f15845_D])).
fof(f15845_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f15847,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | $spl402),
  inference(cnf_transformation,[],[f15847_D])).
fof(f15847_D,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f17457,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl356),
  inference(forward_demodulation,[],[f17406,f14241])).
fof(f17406,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | $spl356),
  inference(superposition,[],[f568,f14241])).
fof(f17421,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl356),
  inference(forward_demodulation,[],[f17368,f14241])).
fof(f17368,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | $spl356),
  inference(superposition,[],[f14241,f568])).
fof(f15842,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | $spl398),
  inference(cnf_transformation,[],[f15842_D])).
fof(f15842_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f17363,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK29(sK17,sK16,sK14,sK13)))))) | $spl356),
  inference(superposition,[],[f14241,f568])).
fof(f16875,plain,(
  $false | ($spl56 | $spl125)),
  inference(subsumption_resolution,[],[f5461,f15993])).
fof(f5487,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl129),
  inference(cnf_transformation,[],[f5487_D])).
fof(f5487_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f16196,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl56),
  inference(resolution,[],[f15955,f10520])).
fof(f15955,plain,(
  ( ! [X140,X141] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X140,X141)) ) | $spl56),
  inference(resolution,[],[f3290,f445])).
fof(f16033,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl404),
  inference(cnf_transformation,[],[f16033_D])).
fof(f16033_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f15995,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl56),
  inference(resolution,[],[f3290,f570])).
fof(f15898,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X26) != s(bool,X27)) ) | $spl56),
  inference(resolution,[],[f3290,f375])).
fof(f15895,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl56),
  inference(resolution,[],[f3290,f374])).
fof(f15894,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl56),
  inference(resolution,[],[f3290,f373])).
fof(f15893,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X11) != s(bool,X12)) ) | $spl56),
  inference(resolution,[],[f3290,f371])).
fof(f15891,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl56),
  inference(resolution,[],[f3290,f369])).
fof(f15173,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | $spl348),
  inference(superposition,[],[f14048,f568])).
fof(f15883,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | ($spl57 | $spl348)),
  inference(forward_demodulation,[],[f15824,f15208])).
fof(f15208,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | ($spl57 | $spl348)),
  inference(subsumption_resolution,[],[f15207,f555])).
fof(f15207,plain,(
  p(s(bool,'F_0')) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) | ($spl57 | $spl348)),
  inference(forward_demodulation,[],[f15173,f3859])).
fof(f3859,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl57),
  inference(resolution,[],[f3291,f510])).
fof(f15824,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | ($spl57 | $spl348)),
  inference(superposition,[],[f568,f15208])).
fof(f15843,plain,(
  s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | ($spl57 | $spl348)),
  inference(forward_demodulation,[],[f15786,f15208])).
fof(f15786,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))))) | ($spl57 | $spl348)),
  inference(superposition,[],[f15208,f568])).
fof(f15781,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK29(sK17,sK16,sK14,sK13)) | s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30(sK29(sK17,sK16,sK14,sK13)))))) | ($spl57 | $spl348)),
  inference(superposition,[],[f15208,f568])).
fof(f15643,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))))))) | $spl394),
  inference(cnf_transformation,[],[f15643_D])).
fof(f15643_D,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))))))) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f8756,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | $spl226),
  inference(cnf_transformation,[],[f8756_D])).
fof(f8756_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f15582,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))))))) | $spl228),
  inference(superposition,[],[f8805,f568])).
fof(f13754,plain,(
  s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | $spl341),
  inference(cnf_transformation,[],[f13754_D])).
fof(f13754_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) <=> ~$spl341),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl341])])).
fof(f15524,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X10))) ) | $spl29),
  inference(resolution,[],[f3610,f10572])).
fof(f3610,plain,(
  ( ! [X125,X124] : (sP8(X124,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)),X125)) ) | $spl29),
  inference(resolution,[],[f2503,f436])).
fof(f13755,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl342),
  inference(cnf_transformation,[],[f13755_D])).
fof(f13755_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f15523,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f3610,f10444])).
fof(f13752,plain,(
  s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | $spl338),
  inference(cnf_transformation,[],[f13752_D])).
fof(f13752_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f15522,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X8))) ) | $spl29),
  inference(resolution,[],[f3610,f10519])).
fof(f15466,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) | $spl392),
  inference(cnf_transformation,[],[f15466_D])).
fof(f15466_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f15458,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15138,f568])).
fof(f15138,plain,(
  ( ! [X158,X159] : (sP10(X158,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X159)) ) | $spl348),
  inference(resolution,[],[f14048,f462])).
fof(f15451,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0,X1)) ) | $spl390),
  inference(cnf_transformation,[],[f15451_D])).
fof(f15451_D,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0,X1)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f15440,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0,X1) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15129,f568])).
fof(f15129,plain,(
  ( ! [X140,X141] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X140,X141)) ) | $spl348),
  inference(resolution,[],[f14048,f445])).
fof(f15058,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl369),
  inference(cnf_transformation,[],[f15058_D])).
fof(f15058_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f15439,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f15129,f10520])).
fof(f15434,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) | $spl388),
  inference(cnf_transformation,[],[f15434_D])).
fof(f15434_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f15423,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15128,f568])).
fof(f15128,plain,(
  ( ! [X138,X139] : (sP6(X138,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X139)) ) | $spl348),
  inference(resolution,[],[f14048,f444])).
fof(f15395,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl386),
  inference(cnf_transformation,[],[f15395_D])).
fof(f15395_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f15384,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15127,f568])).
fof(f15127,plain,(
  ( ! [X136,X137] : (sP6(X136,X137,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f443])).
fof(f15378,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl384),
  inference(cnf_transformation,[],[f15378_D])).
fof(f15378_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f15368,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15126,f568])).
fof(f15126,plain,(
  ( ! [X134,X135] : (sP7(X134,X135,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f439])).
fof(f15363,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl382),
  inference(cnf_transformation,[],[f15363_D])).
fof(f15363_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f15353,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15123,f568])).
fof(f15123,plain,(
  ( ! [X128,X129] : (sP8(X128,X129,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f435])).
fof(f15348,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl380),
  inference(cnf_transformation,[],[f15348_D])).
fof(f15348_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f15338,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15110,f568])).
fof(f15110,plain,(
  ( ! [X103,X102] : (sP4(X102,X103,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f418])).
fof(f15332,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) | $spl378),
  inference(cnf_transformation,[],[f15332_D])).
fof(f15332_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f15323,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15094,f568])).
fof(f15094,plain,(
  ( ! [X70,X71] : (sP2(X70,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X71)) ) | $spl348),
  inference(resolution,[],[f14048,f402])).
fof(f15288,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl376),
  inference(cnf_transformation,[],[f15288_D])).
fof(f15288_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f15279,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15093,f568])).
fof(f15093,plain,(
  ( ! [X68,X69] : (sP2(X68,X69,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f401])).
fof(f15273,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) | $spl374),
  inference(cnf_transformation,[],[f15273_D])).
fof(f15273_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f15264,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X1) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15080,f568])).
fof(f15080,plain,(
  ( ! [X43,X42] : (sP0(X42,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X43)) ) | $spl348),
  inference(resolution,[],[f14048,f385])).
fof(f14257,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0)) ) | $spl360),
  inference(cnf_transformation,[],[f14257_D])).
fof(f14257_D,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0)) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f15249,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15150,f568])).
fof(f14242,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl358),
  inference(cnf_transformation,[],[f14242_D])).
fof(f14242_D,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f15217,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl348),
  inference(superposition,[],[f15149,f568])).
fof(f15149,plain,(
  ( ! [X172] : (sP12(X172,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f483])).
fof(f14217,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl354),
  inference(cnf_transformation,[],[f14217_D])).
fof(f14217_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f15169,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f570])).
fof(f15072,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X26) != s(bool,X27)) ) | $spl348),
  inference(resolution,[],[f14048,f375])).
fof(f15069,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f374])).
fof(f15068,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f373])).
fof(f15067,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X11) != s(bool,X12)) ) | $spl348),
  inference(resolution,[],[f14048,f371])).
fof(f15065,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl348),
  inference(resolution,[],[f14048,f369])).
fof(f15045,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | $spl349),
  inference(resolution,[],[f14708,f10572])).
fof(f15044,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14708,f10444])).
fof(f15043,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X8))) ) | $spl349),
  inference(resolution,[],[f14708,f10519])).
fof(f14935,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl367),
  inference(cnf_transformation,[],[f14935_D])).
fof(f14935_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f14922,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f14358,f10520])).
fof(f14757,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f570])).
fof(f14658,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl349),
  inference(resolution,[],[f14049,f375])).
fof(f14657,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f374])).
fof(f14655,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f373])).
fof(f14653,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl349),
  inference(resolution,[],[f14049,f371])).
fof(f14651,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl349),
  inference(resolution,[],[f14049,f369])).
fof(f3690,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl76),
  inference(cnf_transformation,[],[f3690_D])).
fof(f3690_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f14612,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f570])).
fof(f14513,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X23) != s(bool,X25)) ) | $spl29),
  inference(resolution,[],[f2503,f375])).
fof(f14512,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f374])).
fof(f14510,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f373])).
fof(f14508,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X8) != s(bool,X10)) ) | $spl29),
  inference(resolution,[],[f2503,f371])).
fof(f14506,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f369])).
fof(f14502,plain,(
  $false | ($spl28 | $spl349)),
  inference(subsumption_resolution,[],[f14049,f13894])).
fof(f13894,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl28),
  inference(forward_demodulation,[],[f13758,f582])).
fof(f13758,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl28),
  inference(resolution,[],[f2502,f344])).
fof(f2502,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | $spl28),
  inference(cnf_transformation,[],[f2502_D])).
fof(f2502_D,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f14398,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f570])).
fof(f14301,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl34),
  inference(resolution,[],[f2508,f375])).
fof(f14298,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f374])).
fof(f14297,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f373])).
fof(f14296,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl34),
  inference(resolution,[],[f2508,f371])).
fof(f14294,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl34),
  inference(resolution,[],[f2508,f369])).
fof(f2508,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl34),
  inference(cnf_transformation,[],[f2508_D])).
fof(f2508_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f14050,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),X1)))) ) | $spl350),
  inference(cnf_transformation,[],[f14050_D])).
fof(f14050_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),X1)))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f14074,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1))))))) ) | $spl28),
  inference(resolution,[],[f13894,f348])).
fof(f348,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2)))))) | s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.TAKE_2'(s('type.num.num',X4),s('type.list.list'(X0),X3))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f14290,plain,(
  $false | ($spl16 | $spl28 | $spl35)),
  inference(subsumption_resolution,[],[f13895,f2509])).
fof(f13895,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ($spl16 | $spl28)),
  inference(forward_demodulation,[],[f13759,f1717])).
fof(f13759,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | $spl28),
  inference(resolution,[],[f2502,f345])).
fof(f14247,plain,(
  ( ! [X0] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl28),
  inference(superposition,[],[f14161,f568])).
fof(f14161,plain,(
  ( ! [X173] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X173)) ) | $spl28),
  inference(resolution,[],[f13894,f484])).
fof(f14230,plain,(
  ( ! [X0] : (sP12(X0,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK29(sK17,sK16,sK14,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK29(sK17,sK16,sK14,sK13)))))) ) | $spl28),
  inference(superposition,[],[f14160,f568])).
fof(f14160,plain,(
  ( ! [X172] : (sP12(X172,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl28),
  inference(resolution,[],[f13894,f483])).
fof(f14180,plain,(
  ( ! [X182,X183] : (p(s(bool,X182)) | p(s(bool,X183)) | s(bool,X183) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl28),
  inference(resolution,[],[f13894,f570])).
fof(f14083,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X26) != s(bool,X27)) ) | $spl28),
  inference(resolution,[],[f13894,f375])).
fof(f14080,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl28),
  inference(resolution,[],[f13894,f374])).
fof(f14079,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl28),
  inference(resolution,[],[f13894,f373])).
fof(f14078,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X11) != s(bool,X12)) ) | $spl28),
  inference(resolution,[],[f13894,f371])).
fof(f14076,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl28),
  inference(resolution,[],[f13894,f369])).
fof(f14051,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl352),
  inference(cnf_transformation,[],[f14051_D])).
fof(f14051_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f14013,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl35),
  inference(resolution,[],[f2509,f570])).
fof(f13914,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X23) != s(bool,X25)) ) | $spl35),
  inference(resolution,[],[f2509,f375])).
fof(f13913,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl35),
  inference(resolution,[],[f2509,f374])).
fof(f13911,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl35),
  inference(resolution,[],[f2509,f373])).
fof(f13909,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X8) != s(bool,X10)) ) | $spl35),
  inference(resolution,[],[f2509,f371])).
fof(f13907,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl35),
  inference(resolution,[],[f2509,f369])).
fof(f14049,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl349),
  inference(cnf_transformation,[],[f14049_D])).
fof(f14049_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f13906,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl35),
  inference(resolution,[],[f2509,f348])).
fof(f13899,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl346),
  inference(cnf_transformation,[],[f13899_D])).
fof(f13899_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f13866,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X179) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl28),
  inference(resolution,[],[f2502,f570])).
fof(f13769,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X22) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f2502,f375])).
fof(f13766,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl28),
  inference(resolution,[],[f2502,f374])).
fof(f13765,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl28),
  inference(resolution,[],[f2502,f373])).
fof(f13764,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X7) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f2502,f371])).
fof(f13762,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl28),
  inference(resolution,[],[f2502,f369])).
fof(f13898,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | $spl344),
  inference(cnf_transformation,[],[f13898_D])).
fof(f13898_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f13896,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | $spl28),
  inference(forward_demodulation,[],[f13761,f582])).
fof(f13761,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl28),
  inference(resolution,[],[f2502,f347])).
fof(f2503,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) | $spl29),
  inference(cnf_transformation,[],[f2503_D])).
fof(f2503_D,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f2505,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | $spl31),
  inference(cnf_transformation,[],[f2505_D])).
fof(f2505_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f13757,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(duplicate_literal_removal,[],[f13756])).
fof(f13756,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(forward_demodulation,[],[f9281,f582])).
fof(f9281,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(resolution,[],[f321,f347])).
fof(f13747,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X10))) ) | $spl29),
  inference(resolution,[],[f3610,f10572])).
fof(f13746,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f3610,f10444])).
fof(f13745,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X8))) ) | $spl29),
  inference(resolution,[],[f3610,f10519])).
fof(f13667,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) | $spl336),
  inference(cnf_transformation,[],[f13667_D])).
fof(f13667_D,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f13608,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) | $spl18),
  inference(superposition,[],[f2990,f568])).
fof(f2990,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl18),
  inference(subsumption_resolution,[],[f2958,f559])).
fof(f2958,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl18),
  inference(superposition,[],[f1720,f568])).
fof(f13607,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl334),
  inference(cnf_transformation,[],[f13607_D])).
fof(f13607_D,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f13604,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2918,f568])).
fof(f2918,plain,(
  ( ! [X161,X162] : (sP10(X161,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X162)) ) | $spl18),
  inference(resolution,[],[f1720,f462])).
fof(f13592,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl332),
  inference(cnf_transformation,[],[f13592_D])).
fof(f13592_D,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f13584,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2909,f568])).
fof(f13591,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl331),
  inference(cnf_transformation,[],[f13591_D])).
fof(f13591_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f13578,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f2909,f10520])).
fof(f13573,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl328),
  inference(cnf_transformation,[],[f13573_D])).
fof(f13573_D,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f13567,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2908,f568])).
fof(f2908,plain,(
  ( ! [X142,X141] : (sP6(X141,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X142)) ) | $spl18),
  inference(resolution,[],[f1720,f444])).
fof(f13537,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl326),
  inference(cnf_transformation,[],[f13537_D])).
fof(f13537_D,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f13531,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2907,f568])).
fof(f2907,plain,(
  ( ! [X140,X139] : (sP6(X139,X140,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f443])).
fof(f13520,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl324),
  inference(cnf_transformation,[],[f13520_D])).
fof(f13520_D,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f13515,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2906,f568])).
fof(f2906,plain,(
  ( ! [X138,X137] : (sP7(X137,X138,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f439])).
fof(f13505,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl322),
  inference(cnf_transformation,[],[f13505_D])).
fof(f13505_D,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f13500,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2903,f568])).
fof(f2903,plain,(
  ( ! [X132,X131] : (sP8(X131,X132,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f435])).
fof(f13490,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl320),
  inference(cnf_transformation,[],[f13490_D])).
fof(f13490_D,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f13485,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2890,f568])).
fof(f2890,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f418])).
fof(f13474,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl318),
  inference(cnf_transformation,[],[f13474_D])).
fof(f13474_D,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f13470,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2874,f568])).
fof(f2874,plain,(
  ( ! [X74,X73] : (sP2(X73,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X74)) ) | $spl18),
  inference(resolution,[],[f1720,f402])).
fof(f13431,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl316),
  inference(cnf_transformation,[],[f13431_D])).
fof(f13431_D,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f13427,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2873,f568])).
fof(f2873,plain,(
  ( ! [X72,X71] : (sP2(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f401])).
fof(f13416,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl314),
  inference(cnf_transformation,[],[f13416_D])).
fof(f13416_D,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f13412,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2860,f568])).
fof(f2860,plain,(
  ( ! [X45,X46] : (sP0(X45,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X46)) ) | $spl18),
  inference(resolution,[],[f1720,f385])).
fof(f8754,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) | $spl224),
  inference(cnf_transformation,[],[f8754_D])).
fof(f8754_D,plain,(
  s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f13297,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) | $spl26),
  inference(superposition,[],[f1858,f568])).
fof(f13074,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3836,f568])).
fof(f3836,plain,(
  ( ! [X149,X148] : (~sP10(X148,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X149) | p(s(bool,X148))) ) | $spl57),
  inference(resolution,[],[f3291,f457])).
fof(f13046,plain,(
  ( ! [X10,X11] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3835,f568])).
fof(f3835,plain,(
  ( ! [X146,X147] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X146,X147) | p(s(bool,X146))) ) | $spl57),
  inference(resolution,[],[f3291,f457])).
fof(f12997,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3834,f568])).
fof(f3834,plain,(
  ( ! [X144,X145] : (~sP10(X145,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X144) | p(s(bool,X144))) ) | $spl57),
  inference(resolution,[],[f3291,f456])).
fof(f12969,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3833,f568])).
fof(f3833,plain,(
  ( ! [X142,X143] : (~sP10(X143,X142,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X142))) ) | $spl57),
  inference(resolution,[],[f3291,f456])).
fof(f12935,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3815,f568])).
fof(f3815,plain,(
  ( ! [X107,X106] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X106,X107) | p(s(bool,X106))) ) | $spl57),
  inference(resolution,[],[f3291,f422])).
fof(f7771,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl212),
  inference(cnf_transformation,[],[f7771_D])).
fof(f7771_D,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f12912,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3814,f568])).
fof(f3814,plain,(
  ( ! [X105,X104] : (sP5(X104,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X105) | p(s(bool,X104))) ) | $spl57),
  inference(resolution,[],[f3291,f422])).
fof(f7750,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) | $spl210),
  inference(cnf_transformation,[],[f7750_D])).
fof(f7750_D,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f12885,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3813,f568])).
fof(f3813,plain,(
  ( ! [X103,X102] : (sP5(X102,X103,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X102))) ) | $spl57),
  inference(resolution,[],[f3291,f421])).
fof(f12828,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3812,f568])).
fof(f3812,plain,(
  ( ! [X101,X100] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X100,X101) | p(s(bool,X101))) ) | $spl57),
  inference(resolution,[],[f3291,f421])).
fof(f12807,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3806,f568])).
fof(f3806,plain,(
  ( ! [X88,X89] : (~sP5(X88,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X89) | ~p(s(bool,X88))) ) | $spl57),
  inference(resolution,[],[f3291,f412])).
fof(f12787,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3805,f568])).
fof(f3805,plain,(
  ( ! [X87,X86] : (~sP5(X86,X87,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X86))) ) | $spl57),
  inference(resolution,[],[f3291,f411])).
fof(f12756,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3790,f568])).
fof(f3790,plain,(
  ( ! [X57,X56] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X56,X57) | ~p(s(bool,X56))) ) | $spl57),
  inference(resolution,[],[f3291,f394])).
fof(f12721,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3789,f568])).
fof(f3789,plain,(
  ( ! [X54,X55] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X55,X54) | ~p(s(bool,X54))) ) | $spl57),
  inference(resolution,[],[f3291,f393])).
fof(f7558,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl198),
  inference(cnf_transformation,[],[f7558_D])).
fof(f7558_D,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f12698,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3784,f568])).
fof(f3784,plain,(
  ( ! [X45,X44] : (sP1(X44,X45,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X44))) ) | $spl57),
  inference(resolution,[],[f3291,f388])).
fof(f12661,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3776,f568])).
fof(f3776,plain,(
  ( ! [X28,X29] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X28,X29) | ~p(s(bool,X28))) ) | $spl57),
  inference(resolution,[],[f3291,f377])).
fof(f12628,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3775,f568])).
fof(f3775,plain,(
  ( ! [X26,X27] : (~sP1(X26,X27,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X26))) ) | $spl57),
  inference(resolution,[],[f3291,f376])).
fof(f12586,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3774,f568])).
fof(f3774,plain,(
  ( ! [X24,X25] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X25,X24) | p(s(bool,X24))) ) | $spl57),
  inference(resolution,[],[f3291,f376])).
fof(f12534,plain,(
  ( ! [X10,X11] : (sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9528,f568])).
fof(f12518,plain,(
  ( ! [X10,X11] : (sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9527,f568])).
fof(f12478,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9494,f568])).
fof(f12453,plain,(
  ( ! [X10,X11] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9493,f568])).
fof(f12405,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9489,f568])).
fof(f12383,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9488,f568])).
fof(f12353,plain,(
  ( ! [X10,X11] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9487,f568])).
fof(f12331,plain,(
  ( ! [X10,X11] : (sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9486,f568])).
fof(f12305,plain,(
  ( ! [X10,X11] : (~sP3(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9480,f568])).
fof(f12278,plain,(
  ( ! [X10,X11] : (~sP3(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9479,f568])).
fof(f12249,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9475,f568])).
fof(f12229,plain,(
  ( ! [X10,X11] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9474,f568])).
fof(f12204,plain,(
  ( ! [X10,X11] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9473,f568])).
fof(f12179,plain,(
  ( ! [X10,X11] : (~sP1(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9464,f568])).
fof(f6569,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X8)) ) | $spl158),
  inference(cnf_transformation,[],[f6569_D])).
fof(f6569_D,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X8)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f12159,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X8) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2930,f568])).
fof(f2930,plain,(
  ( ! [X176] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X176)) ) | $spl18),
  inference(resolution,[],[f1720,f484])).
fof(f6553,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) | $spl154),
  inference(cnf_transformation,[],[f6553_D])).
fof(f6553_D,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f12145,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2929,f568])).
fof(f2929,plain,(
  ( ! [X175] : (sP12(X175,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f483])).
fof(f12110,plain,(
  ( ! [X8] : (~sP12(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3910,f568])).
fof(f3910,plain,(
  ( ! [X167] : (~sP12(X167,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X167))) ) | $spl57),
  inference(subsumption_resolution,[],[f3850,f527])).
fof(f3850,plain,(
  ( ! [X167] : (p(s(bool,'F_0')) | p(s(bool,X167)) | ~sP12(X167,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f482])).
fof(f12083,plain,(
  ( ! [X8] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3909,f568])).
fof(f3909,plain,(
  ( ! [X166] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X166) | p(s(bool,X166))) ) | $spl57),
  inference(subsumption_resolution,[],[f3849,f527])).
fof(f3849,plain,(
  ( ! [X166] : (p(s(bool,'F_0')) | p(s(bool,X166)) | ~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X166)) ) | $spl57),
  inference(resolution,[],[f3291,f482])).
fof(f12031,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3868,f568])).
fof(f3868,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl57),
  inference(resolution,[],[f3291,f544])).
fof(f544,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f316])).
fof(f11527,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3859,f568])).
fof(f11513,plain,(
  ( ! [X8] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3844,f568])).
fof(f3844,plain,(
  ( ! [X161] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X161) | ~p(s(bool,X161))) ) | $spl57),
  inference(resolution,[],[f3291,f471])).
fof(f11490,plain,(
  ( ! [X8] : (sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3843,f568])).
fof(f3843,plain,(
  ( ! [X160] : (sP11(X160,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X160))) ) | $spl57),
  inference(resolution,[],[f3291,f468])).
fof(f11463,plain,(
  ( ! [X8] : (~sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3842,f568])).
fof(f3842,plain,(
  ( ! [X159] : (~sP11(X159,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X159))) ) | $spl57),
  inference(resolution,[],[f3291,f466])).
fof(f11411,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f11280,f568])).
fof(f11280,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(subsumption_resolution,[],[f11279,f508])).
fof(f11279,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(forward_demodulation,[],[f11264,f1677])).
fof(f11264,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(resolution,[],[f3841,f1661])).
fof(f1661,plain,(
  ( ! [X160] : (sP11(X160,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X160))) )),
  inference(resolution,[],[f324,f468])).
fof(f3841,plain,(
  ( ! [X158] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X158) | p(s(bool,X158))) ) | $spl57),
  inference(resolution,[],[f3291,f466])).
fof(f11403,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f11280,f570])).
fof(f11304,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl57),
  inference(resolution,[],[f11280,f375])).
fof(f11303,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f11280,f374])).
fof(f11301,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f11280,f373])).
fof(f11299,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl57),
  inference(resolution,[],[f11280,f371])).
fof(f11297,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f11280,f369])).
fof(f11272,plain,(
  ( ! [X8] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3841,f568])).
fof(f10788,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl42),
  inference(superposition,[],[f9558,f568])).
fof(f4482,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl117),
  inference(cnf_transformation,[],[f4482_D])).
fof(f4482_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f10733,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl42),
  inference(superposition,[],[f9556,f568])).
fof(f9556,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl42),
  inference(resolution,[],[f2527,f540])).
fof(f10717,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9535,f568])).
fof(f10692,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9534,f568])).
fof(f10669,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9533,f568])).
fof(f10648,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f9532,f568])).
fof(f10624,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f10520,f3062])).
fof(f10611,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X9))) ) | $spl57),
  inference(resolution,[],[f10519,f3823])).
fof(f10587,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f10444,f3823])).
fof(f10578,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X9))) ) | $spl57),
  inference(resolution,[],[f10572,f3823])).
fof(f10382,plain,(
  ( ! [X10,X11] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3827,f568])).
fof(f10374,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X4)) ) | $spl57),
  inference(resolution,[],[f3827,f432])).
fof(f10364,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3826,f568])).
fof(f10348,plain,(
  ( ! [X10,X11] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3824,f568])).
fof(f10341,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3824,f433])).
fof(f10332,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3823,f568])).
fof(f10316,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3811,f568])).
fof(f10287,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3810,f568])).
fof(f10270,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3797,f568])).
fof(f10253,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3781,f568])).
fof(f10236,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3780,f568])).
fof(f10199,plain,(
  ( ! [X10,X11] : (sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3071,f568])).
fof(f10182,plain,(
  ( ! [X10,X11] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3062,f568])).
fof(f10165,plain,(
  ( ! [X10,X11] : (sP6(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3061,f568])).
fof(f10148,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3060,f568])).
fof(f10118,plain,(
  ( ! [X10,X11] : (sP7(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3059,f568])).
fof(f10102,plain,(
  ( ! [X10,X11] : (sP8(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3056,f568])).
fof(f10086,plain,(
  ( ! [X10,X11] : (sP4(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3043,f568])).
fof(f10070,plain,(
  ( ! [X10,X11] : (sP2(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3027,f568])).
fof(f10054,plain,(
  ( ! [X10,X11] : (sP2(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3026,f568])).
fof(f10026,plain,(
  ( ! [X10,X11] : (sP0(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3013,f568])).
fof(f9972,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f570])).
fof(f9875,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl18),
  inference(resolution,[],[f1720,f375])).
fof(f9872,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f374])).
fof(f9871,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f373])).
fof(f9870,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl18),
  inference(resolution,[],[f1720,f371])).
fof(f9868,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f369])).
fof(f9859,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3083,f568])).
fof(f9845,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3082,f568])).
fof(f1856,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | $spl24),
  inference(cnf_transformation,[],[f1856_D])).
fof(f1856_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1858,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | $spl26),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9710,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f9651,f582])).
fof(f9651,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f568,f582])).
fof(f9666,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f9606,f582])).
fof(f9606,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f582,f568])).
fof(f9560,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f570])).
fof(f9463,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl42),
  inference(resolution,[],[f2527,f375])).
fof(f9460,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f374])).
fof(f9459,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f373])).
fof(f9458,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl42),
  inference(resolution,[],[f2527,f371])).
fof(f9456,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f369])).
fof(f9402,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) )),
  inference(superposition,[],[f321,f568])).
fof(f2515,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))))))) ) | $spl38),
  inference(cnf_transformation,[],[f2515_D])).
fof(f2515_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9395,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0))))) )),
  inference(superposition,[],[f321,f568])).
fof(f9393,plain,(
  ( ! [X296,X295,X297] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X295),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X296)) | ~p(s(bool,X297)) | s(bool,X297) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f570])).
fof(f9294,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f321,f375])).
fof(f9293,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f374])).
fof(f9291,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f373])).
fof(f9289,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f321,f371])).
fof(f9287,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f369])).
fof(f9209,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl231),
  inference(cnf_transformation,[],[f9209_D])).
fof(f9209_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f9252,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl120),
  inference(subsumption_resolution,[],[f4795,f510])).
fof(f4795,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X28) | p(s(bool,X28))) ) | $spl120),
  inference(superposition,[],[f4751,f510])).
fof(f9250,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl120),
  inference(subsumption_resolution,[],[f4797,f510])).
fof(f4797,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X30) | p(s(bool,X30))) ) | $spl120),
  inference(superposition,[],[f4751,f541])).
fof(f541,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f9249,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl120),
  inference(subsumption_resolution,[],[f4798,f510])).
fof(f4798,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X31) | p(s(bool,X31))) ) | $spl120),
  inference(superposition,[],[f4751,f543])).
fof(f543,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.list.LIST_REL_trans','thm.bool.EQ_CLAUSES|split|2')).
fof(f1171,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) | $spl12),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f9248,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X35)) ) | $spl120),
  inference(subsumption_resolution,[],[f4802,f4751])).
fof(f4802,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl120),
  inference(superposition,[],[f4751,f569])).
fof(f4823,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X36)) ) | $spl120),
  inference(trivial_inequality_removal,[],[f4803])).
fof(f4803,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,X36) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X36)) ) | $spl120),
  inference(superposition,[],[f4751,f569])).
fof(f4822,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X37))) ) | $spl120),
  inference(trivial_inequality_removal,[],[f4804])).
fof(f4804,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,X37) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X37))) ) | $spl120),
  inference(superposition,[],[f4751,f547])).
fof(f4821,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | $spl120),
  inference(trivial_inequality_removal,[],[f4805])).
fof(f4805,plain,(
  ( ! [X38] : (s(bool,X38) != s(bool,X38) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | $spl120),
  inference(superposition,[],[f4751,f545])).
fof(f6175,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f544,f3859])).
fof(f6179,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f546,f3859])).
fof(f546,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f317])).
fof(f6194,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ($spl57 | $spl120)),
  inference(superposition,[],[f4751,f3859])).
fof(f6438,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3868,f510])).
fof(f6440,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3868,f541])).
fof(f6441,plain,(
  s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3868,f543])).
fof(f5460,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl124),
  inference(cnf_transformation,[],[f5460_D])).
fof(f5460_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6445,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl57),
  inference(superposition,[],[f3868,f569])).
fof(f9211,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl120),
  inference(subsumption_resolution,[],[f8537,f496])).
fof(f8537,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X48) != s(bool,X49) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl120),
  inference(superposition,[],[f4766,f545])).
fof(f4766,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | ~p(s(bool,X2))) ) | $spl120),
  inference(superposition,[],[f4751,f547])).
fof(f9210,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl120),
  inference(subsumption_resolution,[],[f8538,f496])).
fof(f8538,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | ~p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl120),
  inference(superposition,[],[f4766,f547])).
fof(f9207,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X52) != s(bool,X53) | ~p(s(bool,X53))) ) | $spl120),
  inference(subsumption_resolution,[],[f8539,f4766])).
fof(f8539,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | s(bool,'F_0') = s(bool,X52)) ) | $spl120),
  inference(superposition,[],[f4766,f569])).
fof(f9158,plain,(
  $false | ($spl18 | $spl24)),
  inference(subsumption_resolution,[],[f9030,f559])).
fof(f9030,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')))) | ($spl18 | $spl24)),
  inference(backward_demodulation,[],[f8809,f1720])).
fof(f8809,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | $spl24),
  inference(backward_demodulation,[],[f1905,f582])).
fof(f1905,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) | $spl24),
  inference(backward_demodulation,[],[f1856,f582])).
fof(f8808,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f2233,f582])).
fof(f2233,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f582,f568])).
fof(f8807,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f2277,f582])).
fof(f2277,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f568,f582])).
fof(f8805,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | $spl228),
  inference(cnf_transformation,[],[f8805_D])).
fof(f8805_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8806,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | $spl16),
  inference(forward_demodulation,[],[f2549,f1717])).
fof(f2549,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | $spl16),
  inference(superposition,[],[f1717,f568])).
fof(f8803,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | $spl16),
  inference(forward_demodulation,[],[f2594,f1717])).
fof(f2594,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))))) | $spl16),
  inference(superposition,[],[f568,f1717])).
fof(f8762,plain,(
  $false | ($spl26 | $spl43)),
  inference(subsumption_resolution,[],[f8761,f555])).
fof(f8761,plain,(
  p(s(bool,'F_0')) | ($spl26 | $spl43)),
  inference(forward_demodulation,[],[f8706,f3413])).
fof(f3413,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl43),
  inference(resolution,[],[f2528,f510])).
fof(f8706,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl26),
  inference(superposition,[],[f560,f1858])).
fof(f8693,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))) | $spl26),
  inference(superposition,[],[f1858,f568])).
fof(f7920,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3836,f568])).
fof(f7900,plain,(
  ( ! [X10,X11] : (~sP10('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3835,f568])).
fof(f7861,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3834,f568])).
fof(f7841,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3833,f568])).
fof(f7804,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3815,f568])).
fof(f7766,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3814,f568])).
fof(f7745,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3813,f568])).
fof(f7707,plain,(
  ( ! [X10,X11] : (sP5('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3812,f568])).
fof(f7679,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3806,f568])).
fof(f7651,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3805,f568])).
fof(f7610,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3790,f568])).
fof(f7586,plain,(
  ( ! [X10,X11] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3789,f568])).
fof(f7549,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3784,f568])).
fof(f7520,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3776,f568])).
fof(f7494,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3775,f568])).
fof(f7463,plain,(
  ( ! [X10,X11] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3774,f568])).
fof(f7447,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl190),
  inference(cnf_transformation,[],[f7447_D])).
fof(f7447_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7439,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3390,f568])).
fof(f3390,plain,(
  ( ! [X151,X150] : (~sP10(X150,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X151) | p(s(bool,X150))) ) | $spl43),
  inference(resolution,[],[f2528,f457])).
fof(f7426,plain,(
  ( ! [X10,X11] : (~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) | $spl188),
  inference(cnf_transformation,[],[f7426_D])).
fof(f7426_D,plain,(
  ( ! [X10,X11] : (~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7419,plain,(
  ( ! [X10,X11] : (~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3389,f568])).
fof(f3389,plain,(
  ( ! [X149,X148] : (~sP10('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X148,X149) | p(s(bool,X148))) ) | $spl43),
  inference(resolution,[],[f2528,f457])).
fof(f7404,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11))) ) | $spl186),
  inference(cnf_transformation,[],[f7404_D])).
fof(f7404_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7396,plain,(
  ( ! [X10,X11] : (~sP10(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3388,f568])).
fof(f3388,plain,(
  ( ! [X146,X147] : (~sP10(X147,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X146) | p(s(bool,X146))) ) | $spl43),
  inference(resolution,[],[f2528,f456])).
fof(f7383,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11))) ) | $spl184),
  inference(cnf_transformation,[],[f7383_D])).
fof(f7383_D,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7376,plain,(
  ( ! [X10,X11] : (~sP10(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3387,f568])).
fof(f3387,plain,(
  ( ! [X144,X145] : (~sP10(X145,X144,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X144))) ) | $spl43),
  inference(resolution,[],[f2528,f456])).
fof(f6887,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) | $spl182),
  inference(cnf_transformation,[],[f6887_D])).
fof(f6887_D,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6870,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3369,f568])).
fof(f3369,plain,(
  ( ! [X109,X108] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X108,X109) | p(s(bool,X108))) ) | $spl43),
  inference(resolution,[],[f2528,f422])).
fof(f6855,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) | $spl180),
  inference(cnf_transformation,[],[f6855_D])).
fof(f6855_D,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6850,plain,(
  ( ! [X10,X11] : (sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3368,f568])).
fof(f3368,plain,(
  ( ! [X107,X106] : (sP5(X106,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X107) | p(s(bool,X106))) ) | $spl43),
  inference(resolution,[],[f2528,f422])).
fof(f6835,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) | $spl178),
  inference(cnf_transformation,[],[f6835_D])).
fof(f6835_D,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6830,plain,(
  ( ! [X10,X11] : (sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3367,f568])).
fof(f3367,plain,(
  ( ! [X105,X104] : (sP5(X104,X105,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X104))) ) | $spl43),
  inference(resolution,[],[f2528,f421])).
fof(f6815,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl176),
  inference(cnf_transformation,[],[f6815_D])).
fof(f6815_D,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6798,plain,(
  ( ! [X10,X11] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3366,f568])).
fof(f3366,plain,(
  ( ! [X103,X102] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X102,X103) | p(s(bool,X103))) ) | $spl43),
  inference(resolution,[],[f2528,f421])).
fof(f6783,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) | $spl174),
  inference(cnf_transformation,[],[f6783_D])).
fof(f6783_D,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6776,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3360,f568])).
fof(f3360,plain,(
  ( ! [X90,X91] : (~sP5(X90,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X91) | ~p(s(bool,X90))) ) | $spl43),
  inference(resolution,[],[f2528,f412])).
fof(f6762,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl172),
  inference(cnf_transformation,[],[f6762_D])).
fof(f6762_D,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f6755,plain,(
  ( ! [X10,X11] : (~sP5(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3359,f568])).
fof(f3359,plain,(
  ( ! [X88,X89] : (~sP5(X88,X89,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X88))) ) | $spl43),
  inference(resolution,[],[f2528,f411])).
fof(f6716,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl170),
  inference(cnf_transformation,[],[f6716_D])).
fof(f6716_D,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f6706,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3344,f568])).
fof(f3344,plain,(
  ( ! [X59,X58] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X58,X59) | ~p(s(bool,X58))) ) | $spl43),
  inference(resolution,[],[f2528,f394])).
fof(f6692,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) | $spl168),
  inference(cnf_transformation,[],[f6692_D])).
fof(f6692_D,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6682,plain,(
  ( ! [X10,X11] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3343,f568])).
fof(f3343,plain,(
  ( ! [X57,X56] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X57,X56) | ~p(s(bool,X56))) ) | $spl43),
  inference(resolution,[],[f2528,f393])).
fof(f6668,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) | $spl166),
  inference(cnf_transformation,[],[f6668_D])).
fof(f6668_D,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f6662,plain,(
  ( ! [X10,X11] : (sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3338,f568])).
fof(f3338,plain,(
  ( ! [X47,X46] : (sP1(X46,X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X46))) ) | $spl43),
  inference(resolution,[],[f2528,f388])).
fof(f6646,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) | $spl164),
  inference(cnf_transformation,[],[f6646_D])).
fof(f6646_D,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6636,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | ~p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3330,f568])).
fof(f3330,plain,(
  ( ! [X30,X31] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X30,X31) | ~p(s(bool,X30))) ) | $spl43),
  inference(resolution,[],[f2528,f377])).
fof(f6622,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) | $spl162),
  inference(cnf_transformation,[],[f6622_D])).
fof(f6622_D,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6611,plain,(
  ( ! [X10,X11] : (~sP1(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X10)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3329,f568])).
fof(f3329,plain,(
  ( ! [X28,X29] : (~sP1(X28,X29,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X28))) ) | $spl43),
  inference(resolution,[],[f2528,f376])).
fof(f6597,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) | $spl160),
  inference(cnf_transformation,[],[f6597_D])).
fof(f6597_D,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6583,plain,(
  ( ! [X10,X11] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | p(s(bool,X11)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3328,f568])).
fof(f3328,plain,(
  ( ! [X26,X27] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X27,X26) | p(s(bool,X26))) ) | $spl43),
  inference(resolution,[],[f2528,f376])).
fof(f6563,plain,(
  ( ! [X8] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0')),X8) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2930,f568])).
fof(f6568,plain,(
  ( ! [X0] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0)) ) | $spl156),
  inference(cnf_transformation,[],[f6568_D])).
fof(f6568_D,plain,(
  ( ! [X0] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6558,plain,(
  ( ! [X0] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X0) | s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) ) | $spl18),
  inference(superposition,[],[f2930,f568])).
fof(f6547,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl18),
  inference(superposition,[],[f2929,f568])).
fof(f6552,plain,(
  ( ! [X0] : (sP12(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl152),
  inference(cnf_transformation,[],[f6552_D])).
fof(f6552_D,plain,(
  ( ! [X0] : (sP12(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6542,plain,(
  ( ! [X0] : (sP12(X0,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13)))))) ) | $spl18),
  inference(superposition,[],[f2929,f568])).
fof(f6500,plain,(
  ( ! [X8] : (~sP12(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3910,f568])).
fof(f6475,plain,(
  ( ! [X8] : (~sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3909,f568])).
fof(f6422,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3868,f568])).
fof(f5910,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3859,f568])).
fof(f5884,plain,(
  ( ! [X8] : (sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3844,f568])).
fof(f5862,plain,(
  ( ! [X8] : (sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3843,f568])).
fof(f5840,plain,(
  ( ! [X8] : (~sP11(X8,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3842,f568])).
fof(f5820,plain,(
  ( ! [X8] : (~sP11('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3841,f568])).
fof(f5807,plain,(
  ( ! [X8] : (~sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) | $spl134),
  inference(cnf_transformation,[],[f5807_D])).
fof(f5807_D,plain,(
  ( ! [X8] : (~sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f5796,plain,(
  ( ! [X8] : (~sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3463,f568])).
fof(f3463,plain,(
  ( ! [X169] : (~sP12(X169,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X169))) ) | $spl43),
  inference(subsumption_resolution,[],[f3404,f527])).
fof(f3404,plain,(
  ( ! [X169] : (p(s(bool,'F_0')) | p(s(bool,X169)) | ~sP12(X169,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f482])).
fof(f5781,plain,(
  ( ! [X8] : (~sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) | $spl132),
  inference(cnf_transformation,[],[f5781_D])).
fof(f5781_D,plain,(
  ( ! [X8] : (~sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5771,plain,(
  ( ! [X8] : (~sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3462,f568])).
fof(f3462,plain,(
  ( ! [X168] : (~sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X168) | p(s(bool,X168))) ) | $spl43),
  inference(subsumption_resolution,[],[f3403,f527])).
fof(f3403,plain,(
  ( ! [X168] : (p(s(bool,'F_0')) | p(s(bool,X168)) | ~sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X168)) ) | $spl43),
  inference(resolution,[],[f2528,f482])).
fof(f5755,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | $spl131),
  inference(cnf_transformation,[],[f5755_D])).
fof(f5755_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f5748,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl43),
  inference(resolution,[],[f4935,f797])).
fof(f797,plain,(
  ( ! [X142,X141] : (sP6('const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)),X141,X142)) )),
  inference(resolution,[],[f323,f445])).
fof(f4935,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl43),
  inference(subsumption_resolution,[],[f4929,f437])).
fof(f4929,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl43),
  inference(resolution,[],[f4902,f449])).
fof(f4902,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl43),
  inference(subsumption_resolution,[],[f4900,f555])).
fof(f4900,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl43),
  inference(resolution,[],[f4897,f432])).
fof(f4897,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) ) | $spl43),
  inference(subsumption_resolution,[],[f4889,f439])).
fof(f4889,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) ) | $spl43),
  inference(resolution,[],[f4861,f437])).
fof(f4861,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl43),
  inference(subsumption_resolution,[],[f4857,f440])).
fof(f4857,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl43),
  inference(resolution,[],[f4851,f428])).
fof(f4851,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl43),
  inference(subsumption_resolution,[],[f4850,f545])).
fof(f4850,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl43),
  inference(subsumption_resolution,[],[f4847,f4837])).
fof(f4837,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl43),
  inference(subsumption_resolution,[],[f4832,f443])).
fof(f4832,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl43),
  inference(resolution,[],[f4723,f427])).
fof(f4723,plain,(
  ( ! [X8,X7] : (sP9(X8,'F_0',X7) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) ) | $spl43),
  inference(superposition,[],[f4676,f446])).
fof(f4676,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl43),
  inference(superposition,[],[f3422,f3413])).
fof(f3422,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | $spl43),
  inference(resolution,[],[f2528,f544])).
fof(f4847,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl43),
  inference(resolution,[],[f4838,f431])).
fof(f4838,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl43),
  inference(subsumption_resolution,[],[f4833,f439])).
fof(f4833,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl43),
  inference(resolution,[],[f4723,f428])).
fof(f5733,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5603,f719])).
fof(f5603,plain,(
  ( ! [X285] : (p(s(bool,X285))) ) | $spl6),
  inference(resolution,[],[f719,f522])).
fof(f5732,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5601,f719])).
fof(f5601,plain,(
  ( ! [X283] : (p(s(bool,X283))) ) | $spl6),
  inference(resolution,[],[f719,f519])).
fof(f5731,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5599,f719])).
fof(f5599,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(resolution,[],[f719,f516])).
fof(f5730,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f5598,f719])).
fof(f5598,plain,(
  ( ! [X281] : (p(s(bool,X281))) ) | $spl6),
  inference(resolution,[],[f719,f513])).
fof(f5606,plain,(
  $false | $spl6),
  inference(resolution,[],[f719,f556])).
fof(f5607,plain,(
  $false | $spl6),
  inference(resolution,[],[f719,f557])).
fof(f5608,plain,(
  $false | $spl6),
  inference(resolution,[],[f719,f575])).
fof(f5616,plain,(
  $false | ($spl6 | $spl18)),
  inference(resolution,[],[f719,f1720])).
fof(f5618,plain,(
  $false | $spl6),
  inference(resolution,[],[f719,f560])).
fof(f5639,plain,(
  $false | $spl6),
  inference(resolution,[],[f719,f323])).
fof(f5481,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | ($spl43 | $spl57)),
  inference(resolution,[],[f4934,f3823])).
fof(f4934,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl43),
  inference(subsumption_resolution,[],[f4928,f443])).
fof(f4928,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl43),
  inference(resolution,[],[f4902,f448])).
fof(f5486,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | $spl126),
  inference(cnf_transformation,[],[f5486_D])).
fof(f5486_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f5480,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X9))) ) | $spl43),
  inference(resolution,[],[f4934,f1641])).
fof(f1641,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)),X123)) )),
  inference(resolution,[],[f324,f436])).
fof(f5454,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | ($spl43 | $spl57)),
  inference(resolution,[],[f4976,f3823])).
fof(f4976,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl43),
  inference(subsumption_resolution,[],[f4975,f443])).
fof(f4975,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl43),
  inference(subsumption_resolution,[],[f4974,f508])).
fof(f4974,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl43),
  inference(subsumption_resolution,[],[f4967,f441])).
fof(f4967,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) ) | $spl43),
  inference(resolution,[],[f4725,f432])).
fof(f4725,plain,(
  ( ! [X12,X11] : (~sP9(X12,'F_0',X11) | ~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X11)) ) | $spl43),
  inference(superposition,[],[f4676,f448])).
fof(f5459,plain,(
  s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | $spl123),
  inference(cnf_transformation,[],[f5459_D])).
fof(f5459_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f5453,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | ~p(s(bool,X9))) ) | $spl43),
  inference(resolution,[],[f4976,f1641])).
fof(f4751,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl120),
  inference(cnf_transformation,[],[f4751_D])).
fof(f4751_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4750,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f4745])).
fof(f4745,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl43),
  inference(resolution,[],[f4722,f449])).
fof(f4722,plain,(
  ( ! [X6,X5] : (sP9(X6,X5,'F_0') | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl43),
  inference(superposition,[],[f4676,f446])).
fof(f4701,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl119),
  inference(cnf_transformation,[],[f4701_D])).
fof(f4701_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f4673,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl43),
  inference(superposition,[],[f3422,f568])).
fof(f4483,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl116),
  inference(cnf_transformation,[],[f4483_D])).
fof(f4483_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4174,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl43),
  inference(superposition,[],[f3413,f568])).
fof(f4166,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) | $spl114),
  inference(cnf_transformation,[],[f4166_D])).
fof(f4166_D,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4161,plain,(
  ( ! [X8] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3398,f568])).
fof(f3398,plain,(
  ( ! [X163] : (sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X163) | ~p(s(bool,X163))) ) | $spl43),
  inference(resolution,[],[f2528,f471])).
fof(f4147,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) | $spl112),
  inference(cnf_transformation,[],[f4147_D])).
fof(f4147_D,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4142,plain,(
  ( ! [X8] : (sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3397,f568])).
fof(f3397,plain,(
  ( ! [X162] : (sP11(X162,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X162))) ) | $spl43),
  inference(resolution,[],[f2528,f468])).
fof(f4128,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) | $spl110),
  inference(cnf_transformation,[],[f4128_D])).
fof(f4128_D,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4123,plain,(
  ( ! [X8] : (~sP11(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3396,f568])).
fof(f3396,plain,(
  ( ! [X161] : (~sP11(X161,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X161))) ) | $spl43),
  inference(resolution,[],[f2528,f466])).
fof(f4111,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) | $spl108),
  inference(cnf_transformation,[],[f4111_D])).
fof(f4111_D,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4106,plain,(
  ( ! [X8] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X8) | p(s(bool,X8)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3395,f568])).
fof(f3395,plain,(
  ( ! [X160] : (~sP11('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X160) | p(s(bool,X160))) ) | $spl43),
  inference(resolution,[],[f2528,f466])).
fof(f4088,plain,(
  ( ! [X10,X11] : (sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3827,f568])).
fof(f4071,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3826,f568])).
fof(f4054,plain,(
  ( ! [X10,X11] : (sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3824,f568])).
fof(f4030,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3823,f568])).
fof(f4015,plain,(
  ( ! [X10,X11] : (sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3811,f568])).
fof(f3999,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3810,f568])).
fof(f3982,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3797,f568])).
fof(f3966,plain,(
  ( ! [X10,X11] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3781,f568])).
fof(f3942,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl57),
  inference(superposition,[],[f3780,f568])).
fof(f3880,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl57),
  inference(superposition,[],[f3291,f568])).
fof(f3872,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f570])).
fof(f3773,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl57),
  inference(resolution,[],[f3291,f375])).
fof(f3772,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f374])).
fof(f3770,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f373])).
fof(f3768,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl57),
  inference(resolution,[],[f3291,f371])).
fof(f3766,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl57),
  inference(resolution,[],[f3291,f369])).
fof(f3765,plain,(
  ( ! [X10,X11] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl84),
  inference(cnf_transformation,[],[f3765_D])).
fof(f3765_D,plain,(
  ( ! [X10,X11] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3759,plain,(
  ( ! [X10,X11] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3209,f568])).
fof(f3209,plain,(
  ( ! [X132,X133] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X132,X133)) ) | $spl43),
  inference(resolution,[],[f2528,f441])).
fof(f3748,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl82),
  inference(cnf_transformation,[],[f3748_D])).
fof(f3748_D,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3742,plain,(
  ( ! [X10,X11] : (sP7(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3208,f568])).
fof(f3208,plain,(
  ( ! [X130,X131] : (sP7(X130,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X131)) ) | $spl43),
  inference(resolution,[],[f2528,f440])).
fof(f3731,plain,(
  ( ! [X10,X11] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl80),
  inference(cnf_transformation,[],[f3731_D])).
fof(f3731_D,plain,(
  ( ! [X10,X11] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3725,plain,(
  ( ! [X10,X11] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3206,f568])).
fof(f3206,plain,(
  ( ! [X127,X126] : (sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X126,X127)) ) | $spl43),
  inference(resolution,[],[f2528,f437])).
fof(f3714,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl78),
  inference(cnf_transformation,[],[f3714_D])).
fof(f3714_D,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3709,plain,(
  ( ! [X10,X11] : (sP8(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3205,f568])).
fof(f3205,plain,(
  ( ! [X125,X124] : (sP8(X124,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X125)) ) | $spl43),
  inference(resolution,[],[f2528,f436])).
fof(f3659,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f570])).
fof(f3560,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X23) != s(bool,X25)) ) | $spl29),
  inference(resolution,[],[f2503,f375])).
fof(f3559,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f374])).
fof(f3557,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f373])).
fof(f3555,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X8) != s(bool,X10)) ) | $spl29),
  inference(resolution,[],[f2503,f371])).
fof(f3553,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl29),
  inference(resolution,[],[f2503,f369])).
fof(f3550,plain,(
  ( ! [X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl74),
  inference(cnf_transformation,[],[f3550_D])).
fof(f3550_D,plain,(
  ( ! [X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3546,plain,(
  ( ! [X10,X11] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3193,f568])).
fof(f3193,plain,(
  ( ! [X101,X100] : (sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X100,X101)) ) | $spl43),
  inference(resolution,[],[f2528,f420])).
fof(f3534,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) | $spl72),
  inference(cnf_transformation,[],[f3534_D])).
fof(f3534_D,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3530,plain,(
  ( ! [X10,X11] : (sP4(X10,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3192,f568])).
fof(f3192,plain,(
  ( ! [X99,X98] : (sP4(X98,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X99)) ) | $spl43),
  inference(resolution,[],[f2528,f419])).
fof(f3518,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl70),
  inference(cnf_transformation,[],[f3518_D])).
fof(f3518_D,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3513,plain,(
  ( ! [X10,X11] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3179,f568])).
fof(f3179,plain,(
  ( ! [X72,X73] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X72,X73)) ) | $spl43),
  inference(resolution,[],[f2528,f403])).
fof(f3501,plain,(
  ( ! [X10,X11] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl68),
  inference(cnf_transformation,[],[f3501_D])).
fof(f3501_D,plain,(
  ( ! [X10,X11] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3497,plain,(
  ( ! [X10,X11] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X10,X11) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3163,f568])).
fof(f3163,plain,(
  ( ! [X41,X40] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),X40,X41)) ) | $spl43),
  inference(resolution,[],[f2528,f386])).
fof(f3485,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl66),
  inference(cnf_transformation,[],[f3485_D])).
fof(f3485_D,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3481,plain,(
  ( ! [X10,X11] : (sP0(X10,X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(superposition,[],[f3162,f568])).
fof(f3162,plain,(
  ( ! [X39,X38] : (sP0(X38,X39,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f384])).
fof(f3293,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl60),
  inference(cnf_transformation,[],[f3293_D])).
fof(f3293_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3426,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f570])).
fof(f3327,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X23) != s(bool,X25)) ) | $spl43),
  inference(resolution,[],[f2528,f375])).
fof(f3326,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f374])).
fof(f3324,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f373])).
fof(f3322,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X8) != s(bool,X10)) ) | $spl43),
  inference(resolution,[],[f2528,f371])).
fof(f3320,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f369])).
fof(f3292,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),X1)))) ) | $spl58),
  inference(cnf_transformation,[],[f3292_D])).
fof(f3292_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),X1)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3319,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(resolution,[],[f2528,f348])).
fof(f3310,plain,(
  ( ! [X8] : (sP12(X8,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl42),
  inference(superposition,[],[f3082,f568])).
fof(f3254,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f570])).
fof(f3155,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X23) != s(bool,X25)) ) | $spl43),
  inference(resolution,[],[f2528,f375])).
fof(f3154,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f374])).
fof(f3152,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f373])).
fof(f3150,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X8) != s(bool,X10)) ) | $spl43),
  inference(resolution,[],[f2528,f371])).
fof(f3148,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl43),
  inference(resolution,[],[f2528,f369])).
fof(f3147,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),X0),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X1)))) | s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))) != s('type.list.list'(sK13),'const.list.TAKE_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl43),
  inference(resolution,[],[f2528,f348])).
fof(f3102,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f570])).
fof(f3005,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl42),
  inference(resolution,[],[f2527,f375])).
fof(f3002,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f374])).
fof(f3001,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f373])).
fof(f3000,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl42),
  inference(resolution,[],[f2527,f371])).
fof(f2998,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl42),
  inference(resolution,[],[f2527,f369])).
fof(f2953,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(sK18(sK17,sK15,sK14,sK13,sK13))))) | $spl18),
  inference(superposition,[],[f1720,f568])).
fof(f2949,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f570])).
fof(f2852,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X29) != s(bool,X30)) ) | $spl18),
  inference(resolution,[],[f1720,f375])).
fof(f2849,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f374])).
fof(f2848,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f373])).
fof(f2847,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X14) != s(bool,X15)) ) | $spl18),
  inference(resolution,[],[f1720,f371])).
fof(f2845,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl18),
  inference(resolution,[],[f1720,f369])).
fof(f2546,plain,(
  $false | $spl17),
  inference(subsumption_resolution,[],[f2341,f1718])).
fof(f2341,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(forward_demodulation,[],[f731,f582])).
fof(f731,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f323,f325])).
fof(f2465,plain,(
  ( ! [X19] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X19))))) )),
  inference(superposition,[],[f321,f568])).
fof(f2457,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK13),sK17)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30(X0))))) )),
  inference(superposition,[],[f321,f568])).
fof(f2455,plain,(
  ( ! [X296,X295,X297] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X295),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X296)) | ~p(s(bool,X297)) | s(bool,X297) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X295),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f570])).
fof(f2356,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X39),s('type.list.list'(sK13),sK17))))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f321,f375])).
fof(f2355,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X35),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f374])).
fof(f2353,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X27),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f373])).
fof(f2351,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X19),s('type.list.list'(sK13),sK17))))) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f321,f371])).
fof(f2349,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK16))))),s(sK13,'const.list.EL_2'(s('type.num.num',X11),s('type.list.list'(sK13),sK17)))))) )),
  inference(resolution,[],[f321,f369])).
fof(f2501,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(duplicate_literal_removal,[],[f2500])).
fof(f2500,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(forward_demodulation,[],[f2343,f582])).
fof(f2343,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK17)))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK29(sK17,sK16,sK14,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))))),
  inference(resolution,[],[f321,f347])).
fof(f2224,plain,(
  $false | ($spl17 | $spl24)),
  inference(subsumption_resolution,[],[f2223,f1856])).
fof(f2223,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | ($spl17 | $spl24)),
  inference(forward_demodulation,[],[f1718,f2116])).
fof(f2116,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | $spl24),
  inference(forward_demodulation,[],[f731,f1905])).
fof(f1899,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f1839,f582])).
fof(f1839,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f568,f582])).
fof(f1854,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f1795,f582])).
fof(f1795,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK30('const.list.LENGTH_1'(s('type.list.list'(sK13),sK16))))))),
  inference(superposition,[],[f582,f568])).
fof(f1782,plain,(
  $false | $spl14),
  inference(equality_resolution,[],[f1174])).
fof(f1723,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) ) | $spl22),
  inference(cnf_transformation,[],[f1723_D])).
fof(f1723_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1690,plain,(
  ( ! [X178,X179] : (p(s(bool,X178)) | ~p(s(bool,X179)) | s(bool,X179) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f324,f570])).
fof(f1591,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f324,f375])).
fof(f1590,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f324,f374])).
fof(f1588,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f324,f373])).
fof(f1586,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f324,f371])).
fof(f1584,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f324,f369])).
fof(f1718,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) | $spl17),
  inference(cnf_transformation,[],[f1718_D])).
fof(f1718_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1722,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) | $spl21),
  inference(cnf_transformation,[],[f1722_D])).
fof(f1722_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1583,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK15))))),s(sK13,'const.list.EL_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.list.list'(sK13),sK17)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f324,f328])).
fof(f1720,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | $spl18),
  inference(cnf_transformation,[],[f1720_D])).
fof(f1720_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1582,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK18(sK17,sK15,sK14,sK13,sK13)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK15))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK13),sK17)))),
  inference(resolution,[],[f324,f327])).
fof(f1567,plain,(
  $false | $spl12),
  inference(equality_resolution,[],[f1171])).
fof(f1506,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1378,f1248])).
fof(f1248,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(equality_resolution,[],[f718])).
fof(f1378,plain,(
  ( ! [X285] : (p(s(bool,X285))) ) | $spl4),
  inference(resolution,[],[f1248,f522])).
fof(f1505,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1376,f1248])).
fof(f1376,plain,(
  ( ! [X283] : (p(s(bool,X283))) ) | $spl4),
  inference(resolution,[],[f1248,f519])).
fof(f1504,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1374,f1248])).
fof(f1374,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(resolution,[],[f1248,f516])).
fof(f1503,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1373,f1248])).
fof(f1373,plain,(
  ( ! [X281] : (p(s(bool,X281))) ) | $spl4),
  inference(resolution,[],[f1248,f513])).
fof(f1381,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f556])).
fof(f1382,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f557])).
fof(f1383,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f575])).
fof(f1392,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f560])).
fof(f1413,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f322])).
fof(f1414,plain,(
  $false | $spl4),
  inference(resolution,[],[f1248,f323])).
fof(f1257,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | ~p(s(bool,X54))) ) | $spl4),
  inference(subsumption_resolution,[],[f1203,f718])).
fof(f1203,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | ~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X54)) ) | $spl4),
  inference(superposition,[],[f718,f569])).
fof(f1256,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52))) ) | $spl4),
  inference(subsumption_resolution,[],[f1202,f718])).
fof(f1202,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | s(bool,'F_0') = s(bool,X52)) ) | $spl4),
  inference(superposition,[],[f718,f569])).
fof(f1250,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1201])).
fof(f1201,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50)) | ~p(s(bool,X50))) ) | $spl4),
  inference(superposition,[],[f718,f547])).
fof(f1251,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1200])).
fof(f1200,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48)) | ~p(s(bool,X48))) ) | $spl4),
  inference(superposition,[],[f718,f545])).
fof(f1175,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X54))) ) | $spl2),
  inference(subsumption_resolution,[],[f1118,f717])).
fof(f1118,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | p(s(bool,X54)) | s(bool,'T_0') = s(bool,X54)) ) | $spl2),
  inference(superposition,[],[f717,f569])).
fof(f1173,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X52))) ) | $spl2),
  inference(subsumption_resolution,[],[f1117,f717])).
fof(f1117,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X52)) ) | $spl2),
  inference(superposition,[],[f717,f569])).
fof(f1164,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1114])).
fof(f1114,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46)) | p(s(bool,X46))) ) | $spl2),
  inference(superposition,[],[f717,f543])).
fof(f1165,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1113])).
fof(f1113,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | p(s(bool,X44))) ) | $spl2),
  inference(superposition,[],[f717,f541])).
fof(f1166,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1111])).
fof(f1111,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl2),
  inference(superposition,[],[f717,f510])).
fof(f1097,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f983,f720])).
fof(f983,plain,(
  ( ! [X293] : (~p(s(bool,X293))) ) | $spl8),
  inference(resolution,[],[f720,f508])).
fof(f989,plain,(
  $false | $spl8),
  inference(resolution,[],[f720,f555])).
fof(f994,plain,(
  $false | $spl8),
  inference(resolution,[],[f720,f559])).
fof(f865,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) | $spl10),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f837,plain,(
  ( ! [X184,X183] : (p(s(bool,X183)) | p(s(bool,X184)) | s(bool,X184) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f323,f570])).
fof(f740,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f323,f375])).
fof(f737,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f323,f374])).
fof(f736,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f323,f373])).
fof(f735,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f323,f371])).
fof(f733,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK17)))) )),
  inference(resolution,[],[f323,f369])).
fof(f716,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) | $spl0),
  inference(cnf_transformation,[],[f716_D])).
fof(f716_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f688,plain,(
  ( ! [X184,X183] : (p(s(bool,X183)) | p(s(bool,X184)) | s(bool,X184) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) )),
  inference(resolution,[],[f322,f570])).
fof(f591,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f322,f375])).
fof(f588,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) )),
  inference(resolution,[],[f322,f374])).
fof(f587,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) )),
  inference(resolution,[],[f322,f373])).
fof(f586,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f322,f371])).
fof(f584,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK13,fun(sK13,bool)),sK14),s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK16)))) )),
  inference(resolution,[],[f322,f369])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1038
% Passive clauses: 12180
% Generated clauses: 37855
% Final active clauses: 524
% Final passive clauses: 4790
% Input formulas: 69
% Initial clauses: 249
% 
% Selected by SInE selection: 60
% SInE iterations: 1
% 
% Duplicate literals: 1026
% Trivial inequalities: 200
% Fw subsumption resolutions: 4361
% Fw demodulations: 4369
% Bw demodulations: 205
% 
% Simple tautologies: 2196
% Equational tautologies: 3
% Forward subsumptions: 12510
% 
% Binary resolution: 13070
% Forward superposition: 7713
% Backward superposition: 5918
% Equality resolution: 19
% 
% Split clauses: 1050
% Split components: 2220
% SAT solver clauses: 1132
% SAT solver unit clauses: 26
% SAT solver binary clauses: 109
% 
% Sat splits: 1050
% Sat splitting refutations: 58
% 
% Memory used [KB]: 11129
% Time elapsed: 0.670 s
% ------------------------------
% ------------------------------
% Success in time 12.399 s
