Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.pair
% (31061)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 4316
% Generated clauses: 28254
% Final active clauses: 133
% Final passive clauses: 3721
% Input formulas: 37
% Initial clauses: 55
% 
% Duplicate literals: 1021
% Trivial inequalities: 728
% Fw subsumption resolutions: 3595
% Fw demodulations: 4719
% Bw demodulations: 15
% 
% Simple tautologies: 646
% Equational tautologies: 55
% Forward subsumptions: 12339
% 
% Binary resolution: 203
% Factoring: 1
% Forward superposition: 8575
% Backward superposition: 8845
% Self superposition: 228
% Equality factoring: 144
% Equality resolution: 43
% 
% Split clauses: 515
% Split components: 1034
% SAT solver clauses: 889
% SAT solver unit clauses: 48
% SAT solver binary clauses: 313
% 
% Sat splits: 515
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 877
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pair
% (31119)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 232
% Passive clauses: 1781
% Generated clauses: 22764
% Final active clauses: 119
% Final passive clauses: 1024
% Input formulas: 37
% Initial clauses: 55
% 
% Duplicate literals: 772
% Trivial inequalities: 351
% Fw subsumption resolutions: 1643
% Fw demodulations: 2887
% Bw demodulations: 485
% Condensations: 1635
% 
% Simple tautologies: 247
% Equational tautologies: 31
% Forward subsumptions: 10615
% Bw demodulations to eq. taut.: 55
% 
% Binary resolution: 331
% Factoring: 1
% Forward superposition: 6462
% Backward superposition: 5266
% Self superposition: 148
% Equality factoring: 69
% Equality resolution: 31
% 
% Split clauses: 112
% Split components: 224
% SAT solver clauses: 2741
% SAT solver unit clauses: 64
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 4
% 
% Sat splits: 112
% Sat splitting refutations: 39
% 
% TWLsolver clauses: 2737
% TWLsolver calls for satisfiability: 406
% 
% Memory used [KB]: 19061
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pair
% SZS output start Proof for thm.pair
fof(f4693,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f138,f196,f194,f137,f197,f193,f430,f435,f433,f1588,f1591,f1593,f1590,f1597,f1598,f1594,f1607,f1608,f4607,f4609,f4613,f4615,f4617,f4619,f4639,f4640,f4642,f4643,f4645,f4646,f4656,f4657,f4692])).
fof(f4692,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4691,f4663])).
fof(f4663,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.C_3'(s(fun(sK1,fun(sK0,bool)),sK2),s(sK0,X0),s(sK1,X1))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f193,f209])).
fof(f209,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | p(s(bool,'const.combin.C_3'(s(fun(X0,fun(X1,bool)),X2),s(X1,X4),s(X0,X3))))) )),
  inference(backward_demodulation,[],[f188,f174])).
fof(f174,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | ~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK10(X2,X1,X0)))),s(X1,sK11(X2,X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((? [X5,X6] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(rectify,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((? [X3,X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))) | ! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : (! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ! [X2] : (! [X3,X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X0,X1),bool),'const.pair.UNCURRY_1'(s(fun(X0,fun(X1,bool)),X2)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X6] : ! [X8] : (! [X9,X10] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9))),s(X6,X10)))) <=> p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.pair.PFORALL_THM')).
fof(f188,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.C_3'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.C_3'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.C_3'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11,X7,X6] : ! [X5,X19,X23] : s(X11,'const.combin.C_3'(s(fun(X7,fun(X6,X11)),X5),s(X6,X19),s(X7,X23))) = s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X5),s(X7,X23))),s(X6,X19)))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.combin.C_DEF')).
fof(f4691,plain,(
  ~p(s(bool,'const.combin.C_3'(s(fun(sK1,fun(sK0,bool)),sK2),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0)),s(sK1,sK3)))) | $spl3),
  inference(forward_demodulation,[],[f4677,f188])).
fof(f4677,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3))),s(sK0,sK4(happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)),sK0))))) | $spl3),
  inference(unit_resulting_resolution,[],[f196,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X1,X0))))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7] : ! [X12] : (p(s(bool,'|dollar|forall_1'(s(fun(X7,bool),X12)))) <=> ! [X9] : p(s(bool,happ(s(fun(X7,bool),X12),s(X7,X9)))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.FORALL_THM')).
fof(f4657,plain,(
  $false | ($spl4 | $spl11)),
  inference(subsumption_resolution,[],[f4629,f141])).
fof(f141,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pair.ELIM_PFORALL_EVAL',aHL_FALSITY)).
fof(f4629,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl11)),
  inference(superposition,[],[f4597,f2355])).
fof(f2355,plain,(
  s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | $spl11),
  inference(unit_resulting_resolution,[],[f1590,f144])).
fof(f144,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pair.ELIM_PFORALL_EVAL',aHL_BOOL_CASES)).
fof(f4597,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0))))) ) | $spl4),
  inference(subsumption_resolution,[],[f4569,f139])).
fof(f139,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.TRUTH')).
fof(f4569,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0))))) ) | $spl4),
  inference(superposition,[],[f210,f403])).
fof(f403,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.C_3'(s(fun(sK1,fun(sK0,bool)),sK2),s(sK0,X0),s(sK1,X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f397,f146])).
fof(f146,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X18] : (s(bool,'T_0') = s(bool,X18) <=> p(s(bool,X18)))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.EQ_CLAUSES|split|0')).
fof(f397,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.combin.C_3'(s(fun(sK1,fun(sK0,bool)),sK2),s(sK0,X2),s(sK1,X1))))) ) | $spl4),
  inference(forward_demodulation,[],[f396,f188])).
fof(f396,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X1))),s(sK0,X2))))) ) | $spl4),
  inference(subsumption_resolution,[],[f394,f139])).
fof(f394,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X1))),s(sK0,X2))))) ) | $spl4),
  inference(superposition,[],[f155,f222])).
fof(f222,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X0)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f197,f146])).
fof(f155,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f210,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.combin.C_3'(s(fun(X0,fun(X1,bool)),X2),s(X1,sK9(X2,X1,X0)),s(X0,sK8(X2,X1,X0))))) | p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(backward_demodulation,[],[f188,f203])).
fof(f203,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X0,fun(X1,bool)),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X2,X1,X0)))),s(X1,sK9(X2,X1,X0)))))) )),
  inference(backward_demodulation,[],[f185,f172])).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X2,X1,X0)))),s(X1,sK9(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),sK7(X2,X1,X0)))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),sK7(X2,X1,X0))))))) | ! [X6,X7] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK8(X2,X1,X0)))),s(X1,sK9(X2,X1,X0))))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X8))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X8)))))) | ! [X6,X7] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X7))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(rectify,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) | ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))) <=> ! [X4,X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X7,X6] : ! [X8] : (! [X15] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X15))))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X15)))))) <=> ! [X16,X17] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X16))),s(X6,X17)))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.pair.ELIM_PFORALL')).
fof(f185,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X4))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X4)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X7,X6] : ! [X12,X5] : s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),X12),s('type.pair.prod'(X7,X6),X5))) = s(X11,happ(s(fun(X6,X11),happ(s(fun(X7,fun(X6,X11)),X12),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X6),X5))))),s(X6,'const.pair.SND_1'(s('type.pair.prod'(X7,X6),X5)))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.pair.ELIM_UNCURRY')).
fof(f4656,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f4627,f141])).
fof(f4627,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4)),
  inference(superposition,[],[f4597,f1776])).
fof(f1776,plain,(
  s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1622,f150])).
fof(f150,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X18] : (s(bool,'F_0') = s(bool,X18) <=> ~p(s(bool,X18)))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.EQ_CLAUSES|split|2')).
fof(f1622,plain,(
  ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl1),
  inference(forward_demodulation,[],[f1610,f184])).
fof(f184,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X7,X6] : ! [X13,X14] : s(X11,'const.pair.UNCURRY_2'(s(fun(X7,fun(X6,X11)),X13),s('type.pair.prod'(X7,X6),X14))) = s(X11,happ(s(fun('type.pair.prod'(X7,X6),X11),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,X11)),X13))),s('type.pair.prod'(X7,X6),X14)))),
  file('thm.pair.ELIM_PFORALL_EVAL','ARITY_1_2_const.pair.UNCURRY')).
fof(f1610,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f156])).
fof(f4646,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f4606,f139])).
fof(f4606,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1777,f4597,f168])).
fof(f168,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X25,X26] : ((p(s(bool,X25)) => p(s(bool,X26))) => ((p(s(bool,X26)) => p(s(bool,X25))) => s(bool,X25) = s(bool,X26)))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.IMP_ANTISYM_AX')).
fof(f1777,plain,(
  s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK4('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f1622,f145])).
fof(f145,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f4645,plain,(
  $false | ($spl4 | $spl11)),
  inference(subsumption_resolution,[],[f4608,f139])).
fof(f4608,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl11)),
  inference(unit_resulting_resolution,[],[f1590,f4597,f168])).
fof(f4643,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f4612,f139])).
fof(f4612,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1777,f4597,f168])).
fof(f4642,plain,(
  $false | ($spl4 | $spl11)),
  inference(subsumption_resolution,[],[f4614,f139])).
fof(f4614,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl11)),
  inference(unit_resulting_resolution,[],[f1590,f4597,f168])).
fof(f4640,plain,(
  $false | ($spl4 | $spl11)),
  inference(subsumption_resolution,[],[f4638,f215])).
fof(f215,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f141,f145])).
fof(f4638,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl11)),
  inference(backward_demodulation,[],[f4618,f2355])).
fof(f4618,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),X0)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f4597,f146])).
fof(f4639,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f4637,f215])).
fof(f4637,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f4618,f1776])).
fof(f4619,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1777,f4597,f146])).
fof(f4617,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f1776,f4597,f149])).
fof(f149,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f4615,plain,(
  $false | ($spl4 | $spl11)),
  inference(unit_resulting_resolution,[],[f139,f1590,f4597,f168])).
fof(f4613,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f139,f1777,f4597,f168])).
fof(f4609,plain,(
  $false | ($spl4 | $spl11)),
  inference(unit_resulting_resolution,[],[f139,f1590,f4597,f168])).
fof(f4607,plain,(
  $false | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f139,f1777,f4597,f168])).
fof(f1608,plain,(
  $false | ($spl1 | $spl14)),
  inference(subsumption_resolution,[],[f194,f1594])).
fof(f1607,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f1606,f139])).
fof(f1606,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f196,f222])).
fof(f1594,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1598,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))) | s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X6))))) ) | $spl16),
  inference(cnf_transformation,[],[f1598_D])).
fof(f1598_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))) | s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X6))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1597,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X6))))) ) | $spl1),
  inference(forward_demodulation,[],[f1596,f142])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.marker.Abbrev_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : s(bool,X0) = s(bool,'const.marker.Abbrev_1'(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : s(bool,X9) = s(bool,'const.marker.Abbrev_1'(s(bool,X9)))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.marker.Abbrev_def')).
fof(f1596,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X6)))) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7)))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1577,f145])).
fof(f1577,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X6)))) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7))))) | s(bool,'T_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f158,f683])).
fof(f683,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl1),
  inference(superposition,[],[f592,f531])).
fof(f531,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f522,f141])).
fof(f522,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f298,f234])).
fof(f234,plain,(
  s(bool,'F_0') = s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f194,f150])).
fof(f298,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f288,f141])).
fof(f288,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f165,f144])).
fof(f165,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f102])).
fof(f592,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,X9)))) ) | $spl1),
  inference(superposition,[],[f531,f142])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X1,X0)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : ! [X5] : (p(s(bool,'|dollar|forall_1'(s(fun(X7,bool),X5)))) <=> ! [X19] : s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X5),s(X7,X19))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.bool.FORALL_DEF')).
fof(f1590,plain,(
  s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | $spl11),
  inference(cnf_transformation,[],[f1590_D])).
fof(f1590_D,plain,(
  s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1593,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | p(s(bool,X1))) ) | $spl1),
  inference(forward_demodulation,[],[f1592,f184])).
fof(f1592,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1573,f145])).
fof(f1573,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(superposition,[],[f158,f531])).
fof(f1591,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f1591_D])).
fof(f1591_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1588,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0)))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1587,f146])).
fof(f1587,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | p(s(bool,X0))) ) | $spl1),
  inference(forward_demodulation,[],[f1586,f142])).
fof(f1586,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK0,bool)),sK2),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f1572,f184])).
fof(f1572,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))),s('type.pair.prod'(sK1,sK0),sK5('const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)),'type.pair.prod'(sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.marker.Abbrev_1'(s(bool,X0)))) ) | $spl1),
  inference(superposition,[],[f158,f592])).
fof(f433,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun(bool,bool),'const.marker.Abbrev_0')))) | $spl6),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun(bool,bool),'const.marker.Abbrev_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f435,plain,(
  ~p(s(bool,sK4('const.marker.Abbrev_0',bool))) | $spl9),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ~p(s(bool,sK4('const.marker.Abbrev_0',bool))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f430,plain,(
  ~p(s(bool,sK4('const.marker.Abbrev_0',bool))) | p(s(bool,'|dollar|forall_1'(s(fun(bool,bool),'const.marker.Abbrev_0'))))),
  inference(superposition,[],[f156,f198])).
fof(f198,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.marker.Abbrev_0'),s(bool,X0)))) )),
  inference(forward_demodulation,[],[f143,f142])).
fof(f143,plain,(
  ( ! [X0] : (s(bool,'const.marker.Abbrev_1'(s(bool,X0))) = s(bool,happ(s(fun(bool,bool),'const.marker.Abbrev_0'),s(bool,X0)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s(bool,'const.marker.Abbrev_1'(s(bool,X0))) = s(bool,happ(s(fun(bool,bool),'const.marker.Abbrev_0'),s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : s(bool,'const.marker.Abbrev_1'(s(bool,X13))) = s(bool,happ(s(fun(bool,bool),'const.marker.Abbrev_0'),s(bool,X13)))),
  file('thm.pair.ELIM_PFORALL_EVAL','ARITY_0_1_const.marker.Abbrev')).
fof(f193,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl0),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f197,plain,(
  ( ! [X4] : (p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X4))))))) ) | $spl4),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  ( ! [X4] : (p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X4))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f137,plain,(
  ( ! [X4] : (p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X4)))))) | p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ! [X4] : p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,X4))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ! [X4] : p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ? [X3] : ~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ! [X3] : p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))) & (~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) | ? [X3] : ~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))))))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <~> ! [X3] : p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <=> ! [X3] : p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X1,X0),bool),'const.pair.UNCURRY_1'(s(fun(X1,fun(X0,bool)),X2)))))) <=> ! [X3] : p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))) <=> ! [X9] : p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : (p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(X7,X6),bool),'const.pair.UNCURRY_1'(s(fun(X7,fun(X6,bool)),X8)))))) <=> ! [X9] : p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X8),s(X7,X9)))))))),
  file('thm.pair.ELIM_PFORALL_EVAL','thm.pair.ELIM_PFORALL_EVAL')).
fof(f194,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f196,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | $spl3),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f138,plain,(
  ~p(s(bool,'|dollar|forall_1'(s(fun(sK0,bool),happ(s(fun(sK1,fun(sK0,bool)),sK2),s(sK1,sK3)))))) | ~p(s(bool,'|dollar|forall_1'(s(fun('type.pair.prod'(sK1,sK0),bool),'const.pair.UNCURRY_1'(s(fun(sK1,fun(sK0,bool)),sK2))))))),
  inference(cnf_transformation,[],[f106])).
% SZS output end Proof for thm.pair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 77
% Passive clauses: 347
% Generated clauses: 4543
% Final active clauses: 37
% Final passive clauses: 109
% Input formulas: 37
% Initial clauses: 55
% 
% Duplicate literals: 123
% Trivial inequalities: 54
% Fw subsumption resolutions: 257
% Fw demodulations: 901
% Bw demodulations: 38
% Condensations: 170
% 
% Simple tautologies: 56
% Equational tautologies: 45
% Forward subsumptions: 2567
% 
% Unit resulting resolution: 229
% Factoring: 1
% Forward superposition: 1206
% Backward superposition: 1353
% Self superposition: 87
% Equality factoring: 34
% Equality resolution: 8
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 378
% SAT solver unit clauses: 35
% SAT solver binary clauses: 181
% 
% Sat splits: 6
% Sat splitting refutations: 17
% 
% Memory used [KB]: 4349
% Time elapsed: 0.156 s
% ------------------------------
% ------------------------------
% Success in time 1.469 s
