Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (117134)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 1834
% Generated clauses: 20188
% Final active clauses: 126
% Final passive clauses: 1555
% Input formulas: 25
% Initial clauses: 67
% 
% Duplicate literals: 1022
% Trivial inequalities: 386
% Fw subsumption resolutions: 3303
% Fw demodulations: 2172
% Bw demodulations: 13
% 
% Simple tautologies: 556
% Equational tautologies: 58
% Forward subsumptions: 8987
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 369
% Forward superposition: 5577
% Backward superposition: 6684
% Self superposition: 199
% Equality factoring: 163
% Equality resolution: 37
% 
% Split clauses: 1197
% Split components: 2507
% SAT solver clauses: 1891
% SAT solver unit clauses: 30
% SAT solver binary clauses: 172
% 
% Sat splits: 1197
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1888
% TWLsolver calls for satisfiability: 184
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f3898,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3896,f107,f189,f108,f196,f151,f212,f216,f220,f104,f194,f113,f200,f217,f222,f105,f192,f3760,f1319,f3762,f703,f859,f241,f244,f106,f190,f112,f201,f252,f279,f111,f199,f1411,f110,f198,f3658,f3652,f227,f232,f245,f247,f240,f243,f117,f202,f248,f250,f2293,f2294,f344,f352,f150,f213])).
fof(f213,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl24),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,sK4(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X3,X2,X1,X0)),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK4(X3,X2,X1,X0))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f64])).
fof(f64,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME','thm.real_sigma.real_sigma_9_prove_ltb')).
fof(f352,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(X13,X16),s(fun(X13,bool),X15)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X13,'type.realax.real'),X14),s(fun(X13,bool),X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X13,bool),X15))))),s('type.realax.real',happ(s(fun(X13,'type.realax.real'),X14),s(X13,X16))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X13,bool),X15)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X13,sK4(X16,X14,X15,X13)),s(fun(X13,bool),X15)))) ) | $spl72),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,'const.bool.IN_2'(s(X13,X16),s(fun(X13,bool),X15)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X13,'type.realax.real'),X14),s(fun(X13,bool),X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X13,bool),X15))))),s('type.realax.real',happ(s(fun(X13,'type.realax.real'),X14),s(X13,X16))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X13,bool),X15)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X13,sK4(X16,X14,X15,X13)),s(fun(X13,bool),X15)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f344,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X13,'type.realax.real'),X14),s(fun(X13,bool),X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X13,bool),X15))))),s('type.realax.real',happ(s(fun(X13,'type.realax.real'),X14),s(X13,X16))))) | ~p(s(bool,'const.bool.IN_2'(s(X13,X16),s(fun(X13,bool),X15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X13,bool),X15)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X13,sK4(X16,X14,X15,X13)),s(fun(X13,bool),X15)))) ) | ($spl14 | $spl24)),
  inference(resolution,[],[f213,f200])).
fof(f2294,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl282),
  inference(cnf_transformation,[],[f2294_D])).
fof(f2294_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f2293,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl8 | $spl30 | $spl32 | $spl72)),
  inference(subsumption_resolution,[],[f2292,f196])).
fof(f2292,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl8 | $spl30 | $spl32 | $spl72)),
  inference(forward_demodulation,[],[f2291,f220])).
fof(f2291,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl8 | $spl32 | $spl72)),
  inference(subsumption_resolution,[],[f2257,f196])).
fof(f2257,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X0),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X0),s(sK0,sK3))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,X0,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl32 | $spl72)),
  inference(superposition,[],[f352,f222])).
fof(f250,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | $spl42),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f248,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | ($spl1 | $spl36)),
  inference(superposition,[],[f189,f243])).
fof(f202,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f202_D])).
fof(f202_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME','thm.bool.EQ_CLAUSES|split|3')).
fof(f243,plain,(
  ( ! [X0] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | $spl36),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ( ! [X0] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f240,plain,(
  ( ! [X0] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl18)),
  inference(resolution,[],[f190,f202])).
fof(f247,plain,(
  ( ! [X0,X1] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | $spl40),
  inference(cnf_transformation,[],[f247_D])).
fof(f247_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f245,plain,(
  ( ! [X0,X1] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X0))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) ) | $spl36),
  inference(superposition,[],[f243,f243])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl35),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f227,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl11 | $spl16)),
  inference(resolution,[],[f201,f198])).
fof(f3652,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl406),
  inference(cnf_transformation,[],[f3652_D])).
fof(f3652_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f3658,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl35 | $spl40 | $spl42 | $spl282)),
  inference(subsumption_resolution,[],[f3657,f232])).
fof(f3657,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl40 | $spl42 | $spl282)),
  inference(forward_demodulation,[],[f3656,f222])).
fof(f3656,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl40 | $spl42 | $spl282)),
  inference(condensation,[],[f3655])).
fof(f3655,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) | ($spl40 | $spl42 | $spl282)),
  inference(subsumption_resolution,[],[f3640,f250])).
fof(f3640,plain,(
  ( ! [X3] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X3))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) ) | ($spl40 | $spl282)),
  inference(superposition,[],[f2294,f247])).
fof(f198,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f110,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME',aHL_FALSITY)).
fof(f1411,plain,(
  $false | ($spl8 | $spl11 | $spl180)),
  inference(subsumption_resolution,[],[f1350,f196])).
fof(f1350,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl180)),
  inference(backward_demodulation,[],[f1319,f198])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f111,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME',aHL_BOOL_CASES)).
fof(f279,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl44),
  inference(cnf_transformation,[],[f279_D])).
fof(f279_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f252,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f199,f199])).
fof(f201,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME','thm.bool.EQ_CLAUSES|split|1')).
fof(f190,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) ) | $spl2),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f106,plain,(
  ( ! [X4] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) & (p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),sK1)))) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X4)))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f57])).
fof(f57,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4)))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) & s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X10))))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME','thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME')).
fof(f244,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1)))) ) | $spl38),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f241,plain,(
  ( ! [X1] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl16)),
  inference(resolution,[],[f190,f201])).
fof(f859,plain,(
  ( ! [X72,X75] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X72))) | s(bool,X75) = s(bool,'const.bool.IN_2'(s(sK0,X72),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X75)) ) | $spl124),
  inference(cnf_transformation,[],[f859_D])).
fof(f859_D,plain,(
  ( ! [X72,X75] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X72))) | s(bool,X75) = s(bool,'const.bool.IN_2'(s(sK0,X72),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f703,plain,(
  ( ! [X80,X79] : (s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X79))) | s(bool,X80) = s(bool,'const.bool.IN_2'(s(sK0,X79),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl38 | $spl44)),
  inference(trivial_inequality_removal,[],[f605])).
fof(f605,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') != s(bool,'T_0') | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X79))) | s(bool,X80) = s(bool,'const.bool.IN_2'(s(sK0,X79),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl38 | $spl44)),
  inference(superposition,[],[f244,f279])).
fof(f3762,plain,(
  s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK4(sK3,sK2,sK1,sK0)))) | $spl412),
  inference(cnf_transformation,[],[f3762_D])).
fof(f3762_D,plain,(
  s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK4(sK3,sK2,sK1,sK0)))) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f1319,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl180),
  inference(cnf_transformation,[],[f1319_D])).
fof(f1319_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f3760,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK4(sK3,sK2,sK1,sK0))))) ) | ($spl124 | $spl406)),
  inference(duplicate_literal_removal,[],[f3669])).
fof(f3669,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) = s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK4(sK3,sK2,sK1,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl124 | $spl406)),
  inference(superposition,[],[f3652,f859])).
fof(f192,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f105,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f70])).
fof(f222,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl32),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f217,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | ($spl4 | $spl14)),
  inference(resolution,[],[f200,f192])).
fof(f200,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f194,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl6),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f104,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f70])).
fof(f220,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl30),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f216,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | ($spl6 | $spl14)),
  inference(resolution,[],[f200,f194])).
fof(f212,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK4(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl22),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK4(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK4(X3,X2,X1,X0)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f196,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f108,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_FINITE_SAME',aHL_TRUTH)).
fof(f189,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))))) | $spl1),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f107,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3)))))),
  inference(cnf_transformation,[],[f70])).
fof(f3896,plain,(
  $false | ($spl1 | $spl8 | $spl22 | $spl30 | $spl32 | $spl412)),
  inference(subsumption_resolution,[],[f3895,f196])).
fof(f3895,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8 | $spl22 | $spl30 | $spl32 | $spl412)),
  inference(forward_demodulation,[],[f3894,f220])).
fof(f3894,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl1 | $spl8 | $spl22 | $spl32 | $spl412)),
  inference(subsumption_resolution,[],[f3893,f196])).
fof(f3893,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl1 | $spl22 | $spl32 | $spl412)),
  inference(forward_demodulation,[],[f3892,f222])).
fof(f3892,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl1 | $spl22 | $spl412)),
  inference(subsumption_resolution,[],[f3891,f189])).
fof(f3891,plain,(
  s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl22 | $spl412)),
  inference(trivial_inequality_removal,[],[f3889])).
fof(f3889,plain,(
  s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) != s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),sK2),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl22 | $spl412)),
  inference(superposition,[],[f212,f3762])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 99
% Passive clauses: 255
% Generated clauses: 3751
% Final active clauses: 78
% Final passive clauses: 113
% Input formulas: 25
% Initial clauses: 67
% 
% Duplicate literals: 291
% Trivial inequalities: 106
% Fw subsumption resolutions: 409
% Fw demodulations: 144
% Bw demodulations: 62
% Condensations: 372
% Global subsumptions: 8
% 
% Simple tautologies: 153
% Equational tautologies: 19
% Forward subsumptions: 1625
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 104
% Forward superposition: 1182
% Backward superposition: 636
% Self superposition: 45
% Equality factoring: 32
% Equality resolution: 36
% 
% Split clauses: 79
% Split components: 165
% SAT solver clauses: 432
% SAT solver unit clauses: 71
% SAT solver binary clauses: 27
% SAT solver learnt clauses: 17
% 
% Sat splits: 79
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 399
% TWLsolver calls for satisfiability: 296
% 
% Memory used [KB]: 3582
% Time elapsed: 0.183 s
% ------------------------------
% ------------------------------
% Success in time 0.802 s
