Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.complex
% (69341)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 128
% Passive clauses: 932
% Generated clauses: 29215
% Final active clauses: 123
% Final passive clauses: 709
% Input formulas: 13
% Initial clauses: 22
% 
% Duplicate literals: 2445
% Trivial inequalities: 47
% Fw subsumption resolutions: 3871
% Fw demodulations: 4710
% Bw demodulations: 2
% 
% Simple tautologies: 254
% Equational tautologies: 28
% Forward subsumptions: 13879
% 
% Binary resolution: 12
% Forward superposition: 10534
% Backward superposition: 7164
% Self superposition: 153
% Equality factoring: 36
% Equality resolution: 27
% 
% Split clauses: 1444
% Split components: 2911
% SAT solver clauses: 2418
% SAT solver unit clauses: 4
% SAT solver binary clauses: 45
% 
% Sat splits: 1444
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2400
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 14328
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.complex
% (69359)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 107
% Passive clauses: 388
% Generated clauses: 10584
% Final active clauses: 91
% Final passive clauses: 259
% Input formulas: 13
% Initial clauses: 22
% 
% Duplicate literals: 1130
% Trivial inequalities: 57
% Fw subsumption resolutions: 562
% Fw demodulations: 741
% Bw demodulations: 32
% Condensations: 2042
% 
% Simple tautologies: 148
% Equational tautologies: 44
% Forward subsumptions: 4385
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 8
% Forward superposition: 2591
% Backward superposition: 2764
% Self superposition: 144
% Equality factoring: 110
% Equality resolution: 41
% 
% Split clauses: 88
% Split components: 181
% SAT solver clauses: 1136
% SAT solver unit clauses: 21
% SAT solver binary clauses: 74
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 88
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1135
% TWLsolver calls for satisfiability: 169
% 
% Memory used [KB]: 9083
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.complex
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.complex
% SZS output start Proof for thm.complex
fof(f11555,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f44,f69,f67,f43,f68,f66,f71,f77,f80,f78,f852,f853,f857,f854,f855,f863,f856,f865,f862,f1043,f1054,f1052,f1065,f1066,f1074,f1067,f1068,f1077,f1069,f1079,f1619,f1620,f1621,f1622,f1623,f1624,f1078,f1625,f1626,f1627,f1708,f1076,f1628,f1073,f1072,f2820,f2822,f2825,f2828,f2856,f2857,f4443,f4423,f4409,f4459,f4404,f4403,f4394,f4452,f4392,f4460,f4358,f4474,f4357,f4475,f4476,f4352,f4527,f4528,f4290,f4529,f4275,f4531,f4308,f4532,f4309,f4533,f4320,f4534,f4321,f4535,f4322,f4536,f4323,f4537,f4324,f4538,f4325,f4477,f4539,f4345,f4541,f4346,f4542,f4552,f4555,f4556,f3074,f4557,f1246,f860,f858,f4953,f4954,f4982,f4990,f4994,f4995,f4996,f4997,f4999,f5003,f5004,f5005,f4998,f5007,f5010,f5011,f4992,f5027,f5028,f5030,f5031,f5021,f5032,f5020,f5033,f1373,f1370,f5006,f5038,f2238,f2235,f5002,f5042,f2538,f2541,f5039,f3747,f3744,f4993,f3707,f4991,f5044,f4451,f5035,f5327,f7109,f7110,f7127,f7126,f11545])).
fof(f11545,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f55,f11539,f8928])).
fof(f8928,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)))))) ) | $spl0),
  inference(superposition,[],[f932,f66])).
fof(f932,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X28),s('type.realax.real',X29))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),X27))) | s('type.realax.real',X28) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X26))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X27)))))) )),
  inference(superposition,[],[f60,f55])).
fof(f60,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) | s(X0,X3) = s(X0,X5)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) | (s(X0,X3) = s(X0,X5) & s(X1,X2) = s(X1,X4))) & (s(X0,X3) != s(X0,X5) | s(X1,X2) != s(X1,X4) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4)))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) | (s(X0,X3) = s(X0,X5) & s(X1,X2) = s(X1,X4))) & ((s(X0,X3) != s(X0,X5) | s(X1,X2) != s(X1,X4)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4)))))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) <=> (s(X0,X3) = s(X0,X5) & s(X1,X2) = s(X1,X4)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) <=> (s(X0,X3) = s(X0,X5) & s(X1,X2) = s(X1,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9,X13] : ! [X12,X11,X14,X15] : (s('type.pair.prod'(X9,X13),'const.pair.,_2'(s(X9,X11),s(X13,X12))) = s('type.pair.prod'(X9,X13),'const.pair.,_2'(s(X9,X15),s(X13,X14))) <=> (s(X9,X11) = s(X9,X15) & s(X13,X12) = s(X13,X14)))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.pair.PAIR_EQ')).
fof(f11539,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)))))) ) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f11265,f63])).
fof(f63,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2)) & (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2))) <=> s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11,X12,X6] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X6))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X12))) <=> s('type.realax.real',X6) = s('type.realax.real',X12))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.real.REAL_EQ_LADD')).
fof(f11265,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f69,f10915,f58])).
fof(f58,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & ((s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),X7) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.complex.COMPLEX_RE_IM_EQ')).
fof(f10915,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f10829,f63])).
fof(f10829,plain,(
  s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))))) | $spl0),
  inference(unit_resulting_resolution,[],[f55,f8525])).
fof(f8525,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)))))) ) | $spl0),
  inference(superposition,[],[f930,f66])).
fof(f930,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X20),s('type.realax.real',X21))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X18),s('type.pair.prod'('type.realax.real','type.realax.real'),X19))) | s('type.realax.real',X21) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X18))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X19)))))) )),
  inference(superposition,[],[f61,f55])).
fof(f61,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4))) | s(X1,X2) = s(X1,X4)) )),
  inference(cnf_transformation,[],[f41])).
fof(f55,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.complex.complex_add')).
fof(f7126,plain,(
  ( ! [X28,X27] : (s(fun(fun(bool,bool),bool),X27) = s(fun(fun(bool,bool),bool),X28) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X27,X28,bool,fun(bool,bool)))))) ) | $spl92),
  inference(cnf_transformation,[],[f7126_D])).
fof(f7126_D,plain,(
  ( ! [X28,X27] : (s(fun(fun(bool,bool),bool),X27) = s(fun(fun(bool,bool),bool),X28) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X27,X28,bool,fun(bool,bool)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7127,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X29)))) ) | $spl94),
  inference(cnf_transformation,[],[f7127_D])).
fof(f7127_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X29)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f7110,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X32))) | s(fun(fun(bool,bool),bool),X30) = s(fun(fun(bool,bool),bool),X31) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X30,X31,bool,fun(bool,bool)))))) ) | $spl20),
  inference(superposition,[],[f534,f1052])).
fof(f534,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK3(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK3(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f59,f48])).
fof(f48,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.complex.COMPLEX_EQ_LADD',aHL_BOOL_CASES)).
fof(f59,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.complex.COMPLEX_EQ_LADD',aHL_EXT)).
fof(f7109,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X29))) | s(fun(fun(bool,bool),bool),X27) = s(fun(fun(bool,bool),bool),X28) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X27,X28,bool,fun(bool,bool)))))) ) | $spl20),
  inference(superposition,[],[f534,f1052])).
fof(f5327,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f45,f5153])).
fof(f5153,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(superposition,[],[f47,f78])).
fof(f47,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.complex.COMPLEX_EQ_LADD',aHL_FALSITY)).
fof(f45,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.bool.TRUTH')).
fof(f5035,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) ) | $spl88),
  inference(cnf_transformation,[],[f5035_D])).
fof(f5035_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4451,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) | $spl36),
  inference(cnf_transformation,[],[f4451_D])).
fof(f4451_D,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5044,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'T_0'),s(X74,X76))) != s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'F_0'),s(X74,X75)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f5043])).
fof(f5043,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'T_0'),s(X74,X76))) != s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'F_0'),s(X74,X75))) | s(bool,'T_0') = s(bool,X77)) ) | $spl28),
  inference(forward_demodulation,[],[f3708,f1076])).
fof(f3708,plain,(
  ( ! [X76,X74,X77,X75] : (s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'T_0'),s(X74,X76))) != s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,'F_0'),s(X74,X75))) | s(bool,X77) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X77)) ) | $spl28),
  inference(superposition,[],[f289,f1076])).
fof(f289,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,X31))) != s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X32),s(X30,X33))) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f60,f48])).
fof(f4991,plain,(
  ( ! [X69,X73] : (s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X73)) ) | $spl74),
  inference(cnf_transformation,[],[f4991_D])).
fof(f4991_D,plain,(
  ( ! [X69,X73] : (s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X73)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3707,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X72))) != s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'F_0'),s(X70,X71))) | s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X69)) ) | $spl14),
  inference(superposition,[],[f289,f1230])).
fof(f1230,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f862,f48])).
fof(f4993,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'T_0'),s(X51,X53))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f3703,f4293])).
fof(f4293,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4285])).
fof(f4285,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl28)),
  inference(equality_factoring,[],[f1864])).
fof(f1864,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1722,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f48,f48])).
fof(f1722,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl22 | $spl28)),
  inference(backward_demodulation,[],[f1076,f1053])).
fof(f1053,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl22),
  inference(cnf_transformation,[],[f1053_D])).
fof(f1053_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3703,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'T_0'),s(X51,X53))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) | s(bool,X54) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f289,f1722])).
fof(f3744,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) )),
  inference(condensation,[],[f3743])).
fof(f3743,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f3742])).
fof(f3742,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f3698])).
fof(f3698,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'T_0'),s(X24,X26))) != s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,'F_0'),s(X24,X25))) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X23) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f289,f72])).
fof(f3747,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) )),
  inference(condensation,[],[f3746])).
fof(f3746,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f3745])).
fof(f3745,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f3695])).
fof(f3695,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,'T_0'),s(X6,X8))) != s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,'F_0'),s(X6,X7))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X5) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f289,f72])).
fof(f5039,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0')))) ) | $spl90),
  inference(cnf_transformation,[],[f5039_D])).
fof(f5039_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0')))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2541,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,'F_0'))) != s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X38),s(bool,'T_0'))) | s(bool,X39) = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X35)) ) | $spl14),
  inference(superposition,[],[f418,f1230])).
fof(f418,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X31),s(bool,'F_0'))) != s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X32),s(bool,X33))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f61,f48])).
fof(f2538,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'T_0'),s(X21,X23))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'F_0'),s(X21,X22))) | s(bool,X24) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X20)) ) | $spl14),
  inference(superposition,[],[f289,f1230])).
fof(f5042,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') = s(bool,X61) | s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,'F_0'))) != s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X60),s(bool,'T_0')))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f5041])).
fof(f5041,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') = s(bool,X61) | s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,'F_0'))) != s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X60),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X61)) ) | $spl28),
  inference(forward_demodulation,[],[f2227,f1076])).
fof(f2227,plain,(
  ( ! [X61,X59,X60,X58] : (s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,'F_0'))) != s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X60),s(bool,'T_0'))) | s(bool,X61) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X61)) ) | $spl28),
  inference(superposition,[],[f418,f1076])).
fof(f5002,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X51),s(bool,'F_0'))) != s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X52),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2225,f4293])).
fof(f2225,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X51),s(bool,'F_0'))) != s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X52),s(bool,'T_0'))) | s(bool,X53) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f418,f1722])).
fof(f2235,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0')))) )),
  inference(condensation,[],[f2234])).
fof(f2234,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f2233])).
fof(f2233,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f2217])).
fof(f2217,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'F_0'))) != s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X9),s(bool,'T_0'))) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f418,f72])).
fof(f2238,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0')))) )),
  inference(condensation,[],[f2237])).
fof(f2237,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f2236])).
fof(f2236,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f2216])).
fof(f2216,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f418,f72])).
fof(f5038,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,X14) | s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'F_0'),s(X11,X12)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f5037])).
fof(f5037,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,X14) | s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'F_0'),s(X11,X12))) | s(bool,'T_0') = s(bool,X14)) ) | $spl28),
  inference(forward_demodulation,[],[f2075,f1076])).
fof(f2075,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'F_0'),s(X11,X12))) | s(bool,X14) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X14)) ) | $spl28),
  inference(superposition,[],[f289,f1076])).
fof(f5006,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'F_0'),s(X11,X12))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1880,f4293])).
fof(f1880,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'F_0'),s(X11,X12))) | s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f289,f1722])).
fof(f1370,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) )),
  inference(condensation,[],[f1369])).
fof(f1369,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f1368])).
fof(f1368,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f1337])).
fof(f1337,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X9))) != s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'F_0'),s(X7,X8))) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,X6) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f289,f72])).
fof(f1373,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3)))) )),
  inference(condensation,[],[f1372])).
fof(f1372,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(duplicate_literal_removal,[],[f1371])).
fof(f1371,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f1336])).
fof(f1336,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X3))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f289,f72])).
fof(f5033,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'F_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90) | s(bool,'T_0') != s(bool,X87)) ) | $spl86),
  inference(cnf_transformation,[],[f5033_D])).
fof(f5033_D,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'F_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90) | s(bool,'T_0') != s(bool,X87)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5020,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'F_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(bool,'T_0') != s(bool,X87) | s(X88,X89) != s(X88,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl4 | $spl14)),
  inference(backward_demodulation,[],[f78,f2554])).
fof(f2554,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,'T_0') != s(bool,X87) | s(X88,X89) != s(X88,X90) | s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,g(s(bool,X86))),s(X88,X89))) | s(bool,'T_0') = s(bool,X86)) ) | $spl14),
  inference(superposition,[],[f62,f1230])).
fof(f62,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X3) != s(X0,X5) | s(X1,X2) != s(X1,X4) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X5),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f5032,plain,(
  ( ! [X227,X225,X226] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'F_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) | s(bool,'T_0') != s(bool,X225)) ) | $spl84),
  inference(cnf_transformation,[],[f5032_D])).
fof(f5032_D,plain,(
  ( ! [X227,X225,X226] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'F_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) | s(bool,'T_0') != s(bool,X225)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5021,plain,(
  ( ! [X227,X225,X226,X224] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'F_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) | s(bool,'T_0') != s(bool,X225) | s(bool,'T_0') = s(bool,X224)) ) | ($spl4 | $spl14)),
  inference(backward_demodulation,[],[f78,f2588])).
fof(f2588,plain,(
  ( ! [X227,X225,X226,X224] : (s(bool,'T_0') != s(bool,X225) | s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,g(s(bool,X224))))) | s(bool,'T_0') = s(bool,X224)) ) | $spl14),
  inference(superposition,[],[f685,f1230])).
fof(f685,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | s('type.pair.prod'(X3,X0),'const.pair.,_2'(s(X3,X4),s(X0,X1))) = s('type.pair.prod'(X3,X0),'const.pair.,_2'(s(X3,X4),s(X0,X2)))) )),
  inference(equality_resolution,[],[f62])).
fof(f5031,plain,(
  ( ! [X227,X225,X226] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'T_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225)))) ) | $spl82),
  inference(cnf_transformation,[],[f5031_D])).
fof(f5031_D,plain,(
  ( ! [X227,X225,X226] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'T_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5030,plain,(
  ( ! [X227,X225,X226,X224] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'T_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) | s(bool,'T_0') = s(bool,X224)) ) | ($spl4 | $spl12 | $spl14 | $spl48)),
  inference(forward_demodulation,[],[f5029,f4849])).
fof(f4849,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl12 | $spl48)),
  inference(superposition,[],[f859,f4477])).
fof(f859,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl12),
  inference(cnf_transformation,[],[f859_D])).
fof(f859_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f5029,plain,(
  ( ! [X227,X225,X226,X224] : (s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,'F_0'))) = s('type.pair.prod'(X226,bool),'const.pair.,_2'(s(X226,X227),s(bool,X225))) | s(bool,'T_0') = s(bool,X224)) ) | ($spl4 | $spl12 | $spl14 | $spl48)),
  inference(subsumption_resolution,[],[f5021,f4985])).
fof(f4985,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X24))) = s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X25))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl12 | $spl48)),
  inference(global_subsumption,[],[f44,f43,f45,f48,f55,f58,f57,f56,f59,f62,f61,f60,f64,f63,f72,f74,f77,f110,f142,f285,f287,f414,f416,f525,f527,f668,f670,f673,f685,f826,f828,f831,f852,f853,f854,f855,f856,f930,f932,f934,f939,f2502,f530,f532,f830,f825,f672,f667,f534,f529,f418,f413,f289,f284,f47,f859,f4477,f4641,f4841,f4842,f4849,f4967,f4968,f4969,f4970,f4971,f4983,f4973,f4974,f4984,f4975,f4976])).
fof(f4976,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,X25) | s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X24))) = s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X25))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f830])).
fof(f4975,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'T_0'))) = s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X1))) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f825])).
fof(f4984,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X29),s(X31,X32))) = s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X30),s(X31,X33))) | s(bool,'T_0') = s(bool,X29) | s(X31,X32) != s(X31,X33)) ) | ($spl12 | $spl48)),
  inference(global_subsumption,[],[f44,f43,f45,f48,f55,f58,f57,f56,f59,f62,f61,f60,f64,f63,f72,f74,f77,f110,f142,f285,f287,f414,f416,f525,f527,f668,f670,f673,f685,f826,f828,f831,f852,f853,f854,f855,f856,f930,f932,f934,f939,f2502,f530,f532,f830,f825,f672,f667,f534,f529,f418,f413,f289,f284,f47,f859,f4477,f4641,f4841,f4842,f4849,f4967,f4968,f4969,f4970,f4971,f4983,f4973,f4974])).
fof(f4974,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') != s(bool,X30) | s(X31,X32) != s(X31,X33) | s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X29),s(X31,X32))) = s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X30),s(X31,X33))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f672])).
fof(f4973,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'T_0'),s(X2,X3))) = s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X4))) | s(bool,X0) != s(bool,X1) | s(X2,X3) != s(X2,X4) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f667])).
fof(f4983,plain,(
  ( ! [X17,X18,X16] : (s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK3(X17,X18,bool,X16))))) ) | ($spl12 | $spl48)),
  inference(subsumption_resolution,[],[f4972,f2502])).
fof(f4972,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,f(s(X16,sK3(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK3(X17,X18,bool,X16))))) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f534])).
fof(f4971,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'T_0'))) != s(X18,g(s(bool,'T_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK3(X16,X17,X18,bool))) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f529])).
fof(f4970,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X31),s(bool,'T_0'))) != s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X32),s(bool,X33))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X29)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f418])).
fof(f4969,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X4) | s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f413])).
fof(f4968,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'T_0'),s(X30,X31))) != s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X32),s(X30,X33))) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X29)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f289])).
fof(f4967,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X3) | s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X3),s(X1,X4))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f284])).
fof(f4842,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,f(s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl48),
  inference(superposition,[],[f4477,f72])).
fof(f4841,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) ) | $spl48),
  inference(superposition,[],[f4477,f48])).
fof(f4641,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f859,f48])).
fof(f284,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X3),s(X1,X4))) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f60,f48])).
fof(f413,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f61,f48])).
fof(f529,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK3(X16,X17,X18,bool))) )),
  inference(superposition,[],[f59,f48])).
fof(f667,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(X2,X3) != s(X2,X4) | s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'F_0'),s(X2,X3))) = s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f62,f48])).
fof(f672,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'F_0') != s(bool,X30) | s(X31,X32) != s(X31,X33) | s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X29),s(X31,X32))) = s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X30),s(X31,X33))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f62,f48])).
fof(f825,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'F_0'))) = s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f685,f48])).
fof(f830,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X24))) = s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X25))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f685,f48])).
fof(f532,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK3(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK3(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f59,f72])).
fof(f530,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f59,f72])).
fof(f2502,plain,(
  ( ! [X21,X19,X20] : (s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,'T_0') = s(bool,f(s(X19,sK3(X20,X21,bool,X19))))) )),
  inference(trivial_inequality_removal,[],[f2501])).
fof(f2501,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK3(X20,X21,bool,X19)))) | s(bool,'T_0') = s(bool,f(s(X19,sK3(X20,X21,bool,X19))))) )),
  inference(superposition,[],[f534,f48])).
fof(f939,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s('type.pair.prod'(X66,'type.pair.prod'('type.realax.real','type.realax.real')),'const.pair.,_2'(s(X66,X68),s('type.pair.prod'('type.realax.real','type.realax.real'),X69))) != s('type.pair.prod'(X66,'type.pair.prod'('type.realax.real','type.realax.real')),'const.pair.,_2'(s(X66,X67),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X64),s('type.pair.prod'('type.realax.real','type.realax.real'),X65))))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X69) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X64))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X65))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X64))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X65)))))))) )),
  inference(superposition,[],[f61,f55])).
fof(f934,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.pair.prod'('type.pair.prod'('type.realax.real','type.realax.real'),X36),'const.pair.,_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X38),s(X36,X39))) != s('type.pair.prod'('type.pair.prod'('type.realax.real','type.realax.real'),X36),'const.pair.,_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X34),s('type.pair.prod'('type.realax.real','type.realax.real'),X35))),s(X36,X37))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X38) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X34))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X35))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X34))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X35)))))))) )),
  inference(superposition,[],[f60,f55])).
fof(f831,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,X28) != s(bool,X29) | s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X31),s(bool,'T_0'))) = s('type.pair.prod'(X30,bool),'const.pair.,_2'(s(X30,X31),s(bool,X29))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f685,f72])).
fof(f828,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) != s(bool,X16) | s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X14))) = s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f685,f72])).
fof(f826,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X4))) = s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X5))) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f685,f72])).
fof(f673,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,X34) != s(bool,X35) | s(X36,X37) != s(X36,X38) | s('type.pair.prod'(bool,X36),'const.pair.,_2'(s(bool,'T_0'),s(X36,X37))) = s('type.pair.prod'(bool,X36),'const.pair.,_2'(s(bool,X35),s(X36,X38))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f62,f72])).
fof(f670,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,X18) != s(bool,X19) | s(X20,X21) != s(X20,X22) | s('type.pair.prod'(bool,X20),'const.pair.,_2'(s(bool,X17),s(X20,X21))) = s('type.pair.prod'(bool,X20),'const.pair.,_2'(s(bool,X19),s(X20,X22))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f62,f72])).
fof(f668,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s(bool,'T_0') != s(bool,X6) | s(X7,X8) != s(X7,X9) | s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X5),s(X7,X8))) = s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X6),s(X7,X9))) | s(bool,X5) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f62,f72])).
fof(f527,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK3(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f59,f72])).
fof(f525,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK3(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f59,f72])).
fof(f416,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.pair.prod'(X19,bool),'const.pair.,_2'(s(X19,X20),s(bool,X18))) != s('type.pair.prod'(X19,bool),'const.pair.,_2'(s(X19,X21),s(bool,X22))) | s(bool,X17) = s(bool,X22) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f61,f72])).
fof(f414,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,'T_0'))) != s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X8),s(bool,X9))) | s(bool,X5) = s(bool,X9) | s(bool,X5) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f61,f72])).
fof(f287,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20))) != s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X21),s(X19,X22))) | s(bool,X17) = s(bool,X21) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f60,f72])).
fof(f285,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,'T_0'),s(X6,X7))) != s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X8),s(X6,X9))) | s(bool,X5) = s(bool,X8) | s(bool,X5) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f60,f72])).
fof(f142,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f45,f72])).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f72,f72])).
fof(f74,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f47,f48])).
fof(f64,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f56,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f57,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f5028,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'T_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90)) ) | $spl80),
  inference(cnf_transformation,[],[f5028_D])).
fof(f5028_D,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'T_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5027,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'T_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl4 | $spl12 | $spl14 | $spl48)),
  inference(forward_demodulation,[],[f5026,f4849])).
fof(f5026,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,'F_0'),s(X88,X89))) = s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X87),s(X88,X90))) | s(X88,X89) != s(X88,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl4 | $spl12 | $spl14 | $spl48)),
  inference(subsumption_resolution,[],[f5020,f4984])).
fof(f4992,plain,(
  ( ! [X70,X72,X71] : (s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X71))) != s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X72)))) ) | $spl76),
  inference(cnf_transformation,[],[f4992_D])).
fof(f4992_D,plain,(
  ( ! [X70,X72,X71] : (s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X71))) != s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X72)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5011,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f1373,f4849])).
fof(f5010,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f1370,f4849])).
fof(f5007,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X12))) != s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,'T_0'),s(X11,X13))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl12 | $spl22 | $spl28 | $spl48)),
  inference(forward_demodulation,[],[f5006,f4849])).
fof(f4998,plain,(
  ( ! [X37,X38,X36] : (s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,'T_0'))) != s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X38),s(bool,'T_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f4998_D])).
fof(f4998_D,plain,(
  ( ! [X37,X38,X36] : (s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,'T_0'))) != s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X38),s(bool,'T_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5005,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f2238,f4849])).
fof(f5004,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f2235,f4849])).
fof(f5003,plain,(
  ( ! [X52,X50,X53,X51] : (s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X51),s(bool,'T_0'))) != s('type.pair.prod'(X50,bool),'const.pair.,_2'(s(X50,X52),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl12 | $spl22 | $spl28 | $spl48)),
  inference(forward_demodulation,[],[f5002,f4849])).
fof(f4999,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'T_0'),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,'T_0'),s(X21,X23))) | s(bool,X24) = s(bool,g(s(bool,X20))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X20)) ) | ($spl12 | $spl14 | $spl48)),
  inference(forward_demodulation,[],[f2538,f4849])).
fof(f4997,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,'T_0'))) != s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X38),s(bool,'T_0'))) | s(bool,X39) = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X35)) ) | ($spl12 | $spl14 | $spl48)),
  inference(forward_demodulation,[],[f2541,f4849])).
fof(f4996,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f3747,f4849])).
fof(f4995,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X3))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl48)),
  inference(forward_demodulation,[],[f3744,f4849])).
fof(f4994,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'T_0'),s(X51,X52))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'T_0'),s(X51,X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl12 | $spl22 | $spl28 | $spl48)),
  inference(forward_demodulation,[],[f4993,f4849])).
fof(f4990,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X71))) != s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,'T_0'),s(X70,X72))) | s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X69)) ) | ($spl12 | $spl14 | $spl48)),
  inference(forward_demodulation,[],[f3707,f4849])).
fof(f4982,plain,(
  $false | ($spl12 | $spl48)),
  inference(subsumption_resolution,[],[f4966,f45])).
fof(f4966,plain,(
  ~p(s(bool,'T_0')) | ($spl12 | $spl48)),
  inference(backward_demodulation,[],[f4849,f47])).
fof(f4954,plain,(
  $false | ($spl7 | $spl12 | $spl48)),
  inference(subsumption_resolution,[],[f4849,f80])).
fof(f4953,plain,(
  $false | ($spl7 | $spl12 | $spl48)),
  inference(subsumption_resolution,[],[f4840,f80])).
fof(f4840,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl12 | $spl48)),
  inference(superposition,[],[f4477,f859])).
fof(f858,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl10),
  inference(cnf_transformation,[],[f858_D])).
fof(f858_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f860,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl13),
  inference(cnf_transformation,[],[f860_D])).
fof(f860_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1246,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool))) ) | $spl14),
  inference(superposition,[],[f529,f862])).
fof(f4557,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X42) = s(fun(bool,bool),X43) | s(bool,'T_0') = s(bool,sK3(X42,X43,bool,bool))) ) | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f2544,f80])).
fof(f2544,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X42) = s(fun(bool,bool),X43) | s(bool,'T_0') = s(bool,sK3(X42,X43,bool,bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl14),
  inference(superposition,[],[f529,f1230])).
fof(f3074,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl14),
  inference(superposition,[],[f529,f862])).
fof(f4556,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool))) ) | ($spl7 | $spl14)),
  inference(subsumption_resolution,[],[f3075,f80])).
fof(f3075,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl14),
  inference(superposition,[],[f529,f1230])).
fof(f4555,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | ($spl14 | $spl22 | $spl28)),
  inference(condensation,[],[f4554])).
fof(f4554,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool))) ) | ($spl14 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4553])).
fof(f4553,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X8)) ) | ($spl14 | $spl22 | $spl28)),
  inference(forward_demodulation,[],[f3077,f4544])).
fof(f4544,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl14 | $spl22 | $spl28)),
  inference(global_subsumption,[],[f4543,f44,f43,f45,f48,f55,f58,f57,f56,f59,f62,f61,f60,f64,f63,f72,f74,f77,f110,f142,f285,f287,f414,f416,f525,f527,f668,f670,f673,f685,f826,f828,f831,f852,f853,f854,f855,f856,f930,f932,f934,f939,f1230,f2546,f2551,f2588,f2554,f530,f532,f862,f830,f825,f672,f667,f534,f529,f418,f413,f289,f284,f47])).
fof(f2551,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s('type.pair.prod'(X72,bool),'const.pair.,_2'(s(X72,X73),s(bool,'T_0'))) != s('type.pair.prod'(X72,bool),'const.pair.,_2'(s(X72,X74),s(bool,X75))) | s(bool,X75) = s(bool,g(s(bool,X71))) | s(bool,'T_0') = s(bool,X71)) ) | $spl14),
  inference(superposition,[],[f61,f1230])).
fof(f2546,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.pair.prod'(bool,X47),'const.pair.,_2'(s(bool,'T_0'),s(X47,X48))) != s('type.pair.prod'(bool,X47),'const.pair.,_2'(s(bool,X49),s(X47,X50))) | s(bool,X49) = s(bool,g(s(bool,X46))) | s(bool,'T_0') = s(bool,X46)) ) | $spl14),
  inference(superposition,[],[f60,f1230])).
fof(f4543,plain,(
  ( ! [X105] : (s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl14 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f4100,f4294])).
fof(f4294,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | ($spl22 | $spl28)),
  inference(condensation,[],[f4284])).
fof(f4284,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl22 | $spl28)),
  inference(equality_factoring,[],[f1864])).
fof(f4100,plain,(
  ( ! [X105] : (s(bool,'T_0') = s(bool,X105) | s(bool,X105) = s(bool,f(s(bool,X105))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl14 | $spl22 | $spl28)),
  inference(superposition,[],[f862,f1864])).
fof(f3077,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f529,f72])).
fof(f4552,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | ($spl14 | $spl22 | $spl28)),
  inference(condensation,[],[f4551])).
fof(f4551,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK3(X15,X16,bool,bool))) ) | ($spl14 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4550])).
fof(f4550,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK3(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X17)) ) | ($spl14 | $spl22 | $spl28)),
  inference(forward_demodulation,[],[f3080,f4544])).
fof(f3080,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK3(X15,X16,bool,bool)) | s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f529,f72])).
fof(f4542,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) = s(bool,X38) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) | $spl72),
  inference(cnf_transformation,[],[f4542_D])).
fof(f4542_D,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) = s(bool,X38) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4346,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,f(s(bool,X38))) | s(bool,X36) = s(bool,X38)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4154])).
fof(f4154,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,f(s(bool,X38))) | s(bool,X36) = s(bool,X38)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f142,f1864])).
fof(f4541,plain,(
  ( ! [X52,X50,X55,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,X50),s(X51,X53))) | s(bool,X50) = s(bool,X55) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) | $spl70),
  inference(cnf_transformation,[],[f4541_D])).
fof(f4541_D,plain,(
  ( ! [X52,X50,X55,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,X50),s(X51,X53))) | s(bool,X50) = s(bool,X55) | s(bool,X55) = s(bool,f(s(bool,X55)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4345,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,X50),s(X51,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,f(s(bool,X55))) | s(bool,X50) = s(bool,X55)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4157])).
fof(f4157,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))) != s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,X50),s(X51,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,f(s(bool,X55))) | s(bool,X50) = s(bool,X55)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f289,f1864])).
fof(f4539,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1)))) ) | $spl68),
  inference(cnf_transformation,[],[f4539_D])).
fof(f4539_D,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4477,plain,(
  ( ! [X38] : (s(bool,X38) = s(bool,f(s(bool,X38)))) ) | $spl48),
  inference(cnf_transformation,[],[f4477_D])).
fof(f4477_D,plain,(
  ( ! [X38] : (s(bool,X38) = s(bool,f(s(bool,X38)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4325,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1)))) ) | ($spl7 | $spl22 | $spl28)),
  inference(condensation,[],[f4246])).
fof(f4246,plain,(
  ( ! [X163,X161,X162] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X162))) | s(bool,X163) = s(bool,f(s(bool,X163))) | s(bool,X161) = s(bool,X163)) ) | ($spl7 | $spl22 | $spl28)),
  inference(superposition,[],[f283,f1864])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X2)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f80,f60])).
fof(f4538,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X1)))) ) | $spl66),
  inference(cnf_transformation,[],[f4538_D])).
fof(f4538_D,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X1)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4324,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X1)))) ) | ($spl7 | $spl22 | $spl28)),
  inference(condensation,[],[f4247])).
fof(f4247,plain,(
  ( ! [X165,X166,X164] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X165))) | s(bool,X166) = s(bool,f(s(bool,X166))) | s(bool,X164) = s(bool,X166)) ) | ($spl7 | $spl22 | $spl28)),
  inference(superposition,[],[f283,f1864])).
fof(f4537,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) | $spl64),
  inference(cnf_transformation,[],[f4537_D])).
fof(f4537_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4323,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | ($spl22 | $spl28)),
  inference(condensation,[],[f4248])).
fof(f4248,plain,(
  ( ! [X167,X171,X169,X170,X168] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X168),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X169),s(bool,X170))) | s(bool,'F_0') = s(bool,X169) | s(bool,'T_0') = s(bool,X168) | s(bool,X171) = s(bool,f(s(bool,X171))) | s(bool,X167) = s(bool,X171)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f284,f1864])).
fof(f4536,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | $spl62),
  inference(cnf_transformation,[],[f4536_D])).
fof(f4536_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4322,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X2),s(bool,X3))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | ($spl22 | $spl28)),
  inference(condensation,[],[f4249])).
fof(f4249,plain,(
  ( ! [X175,X173,X176,X174,X172] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X175),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X173),s(bool,X174))) | s(bool,'F_0') = s(bool,X175) | s(bool,'T_0') = s(bool,X173) | s(bool,X176) = s(bool,f(s(bool,X176))) | s(bool,X172) = s(bool,X176)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f284,f1864])).
fof(f4535,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3)) ) | $spl60),
  inference(cnf_transformation,[],[f4535_D])).
fof(f4535_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4321,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl22 | $spl28)),
  inference(condensation,[],[f4250])).
fof(f4250,plain,(
  ( ! [X177,X180,X178,X181,X179] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X178),s(bool,X179))) | s(bool,X178) = s(bool,X180) | s(bool,'T_0') = s(bool,X180) | s(bool,X181) = s(bool,f(s(bool,X181))) | s(bool,X177) = s(bool,X181)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f289,f1864])).
fof(f4534,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3)) ) | $spl58),
  inference(cnf_transformation,[],[f4534_D])).
fof(f4534_D,plain,(
  ( ! [X2,X3,X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4320,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl22 | $spl28)),
  inference(condensation,[],[f4251])).
fof(f4251,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X184),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X183))) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,X186) = s(bool,f(s(bool,X186))) | s(bool,X182) = s(bool,X186)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f289,f1864])).
fof(f4533,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'F_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0')))) ) | $spl56),
  inference(cnf_transformation,[],[f4533_D])).
fof(f4533_D,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'F_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0')))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4309,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'F_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'T_0')))) ) | ($spl7 | $spl22 | $spl28)),
  inference(condensation,[],[f4262])).
fof(f4262,plain,(
  ( ! [X243,X241,X242] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'F_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X242),s(bool,'T_0'))) | s(bool,X243) = s(bool,f(s(bool,X243))) | s(bool,X241) = s(bool,X243)) ) | ($spl7 | $spl22 | $spl28)),
  inference(superposition,[],[f403,f1864])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X2),s(bool,'T_0')))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f80,f61])).
fof(f4532,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'F_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f4532_D])).
fof(f4532_D,plain,(
  ( ! [X1] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'F_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4308,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'F_0')))) ) | ($spl7 | $spl22 | $spl28)),
  inference(condensation,[],[f4263])).
fof(f4263,plain,(
  ( ! [X246,X244,X245] : (s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))) != s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X245),s(bool,'F_0'))) | s(bool,X246) = s(bool,f(s(bool,X246))) | s(bool,X244) = s(bool,X246)) ) | ($spl7 | $spl22 | $spl28)),
  inference(superposition,[],[f403,f1864])).
fof(f4531,plain,(
  ( ! [X314] : (s(bool,X314) = s(bool,f(s(bool,X314))) | s(bool,X314) = s(bool,g(s(bool,'F_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f4531_D])).
fof(f4531_D,plain,(
  ( ! [X314] : (s(bool,X314) = s(bool,f(s(bool,X314))) | s(bool,X314) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4275,plain,(
  ( ! [X313,X312,X314] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X312) = s(fun(bool,bool),X313) | s(bool,'T_0') = s(bool,sK3(X312,X313,bool,bool)) | s(bool,X314) = s(bool,f(s(bool,X314))) | s(bool,X314) = s(bool,g(s(bool,'F_0')))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f529,f1864])).
fof(f4529,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | $spl50),
  inference(cnf_transformation,[],[f4529_D])).
fof(f4529_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4290,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl28)),
  inference(equality_factoring,[],[f1864])).
fof(f4528,plain,(
  $false | ($spl13 | $spl36)),
  inference(subsumption_resolution,[],[f4512,f4451])).
fof(f4512,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl13 | $spl36)),
  inference(backward_demodulation,[],[f4451,f860])).
fof(f4527,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f4495,f45])).
fof(f4495,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(backward_demodulation,[],[f4451,f47])).
fof(f4352,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl13 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4073])).
fof(f4073,plain,(
  ( ! [X43,X44] : (s(bool,X44) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,f(s(bool,X44))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl13 | $spl22 | $spl28)),
  inference(superposition,[],[f1447,f1864])).
fof(f1447,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f1147,f48])).
fof(f1147,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl13),
  inference(unit_resulting_resolution,[],[f860,f48])).
fof(f4476,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) | ($spl13 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f4356,f4384])).
fof(f4384,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl13 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4383])).
fof(f4383,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl13 | $spl22 | $spl28)),
  inference(condensation,[],[f3935])).
fof(f3935,plain,(
  ( ! [X50,X49] : (s(bool,X49) = s(bool,f(s(bool,X49))) | s(bool,'T_0') = s(bool,f(s(bool,X49))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl13 | $spl22 | $spl28)),
  inference(superposition,[],[f1447,f1864])).
fof(f4356,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,f(s(bool,X38)))) ) | ($spl13 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4070])).
fof(f4070,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl13 | $spl22 | $spl28)),
  inference(superposition,[],[f1155,f1864])).
fof(f1155,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f860,f48])).
fof(f4475,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X20),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X23),s(X21,X24))) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'F_0') = s(bool,X23)) ) | $spl46),
  inference(cnf_transformation,[],[f4475_D])).
fof(f4475_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X20),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X23),s(X21,X24))) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4357,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X20),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X23),s(X21,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4067])).
fof(f4067,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X20),s(X21,X22))) != s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X23),s(X21,X24))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f284,f1864])).
fof(f4474,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl44),
  inference(cnf_transformation,[],[f4474_D])).
fof(f4474_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4358,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f4065])).
fof(f4065,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f74,f1864])).
fof(f4460,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,'F_0'))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) | $spl42),
  inference(cnf_transformation,[],[f4460_D])).
fof(f4460_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,'F_0'))) | s(bool,X15) = s(bool,f(s(bool,X15)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4392,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,f(s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3915])).
fof(f3915,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl13 | $spl22 | $spl28)),
  inference(superposition,[],[f1864,f1167])).
fof(f1167,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(trivial_inequality_removal,[],[f1161])).
fof(f1161,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f860,f72])).
fof(f4452,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) | $spl38),
  inference(cnf_transformation,[],[f4452_D])).
fof(f4452_D,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,X27)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4394,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12)))) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3912])).
fof(f3912,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,X14))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1864,f1864])).
fof(f4403,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20)))) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3884])).
fof(f3884,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1864,f1864])).
fof(f4404,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,f(s(bool,'F_0'))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl13 | $spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3882])).
fof(f3882,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,f(s(bool,'F_0'))) | s(bool,X15) = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl13 | $spl22 | $spl28)),
  inference(superposition,[],[f1864,f1167])).
fof(f4459,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) | $spl40),
  inference(cnf_transformation,[],[f4459_D])).
fof(f4459_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4409,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3877])).
fof(f3877,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1864,f48])).
fof(f4423,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3864])).
fof(f3864,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1864,f1864])).
fof(f4443,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl22 | $spl28)),
  inference(duplicate_literal_removal,[],[f3834])).
fof(f3834,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,f(s(bool,X27))) | s(bool,X25) = s(bool,X27)) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1864,f1864])).
fof(f2857,plain,(
  $false | ($spl13 | $spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2846,f2808])).
fof(f2808,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl28)),
  inference(duplicate_literal_removal,[],[f2636])).
fof(f2636,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl24 | $spl28)),
  inference(superposition,[],[f1072,f1076])).
fof(f2846,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl13 | $spl24 | $spl28)),
  inference(backward_demodulation,[],[f2808,f860])).
fof(f2856,plain,(
  $false | ($spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2829,f45])).
fof(f2829,plain,(
  ~p(s(bool,'T_0')) | ($spl24 | $spl28)),
  inference(backward_demodulation,[],[f2808,f47])).
fof(f2828,plain,(
  $false | ($spl13 | $spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2827,f860])).
fof(f2827,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl13 | $spl24 | $spl28)),
  inference(condensation,[],[f2826])).
fof(f2826,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl13 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f2813,f1076])).
fof(f2813,plain,(
  ( ! [X10] : (s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl13 | $spl24)),
  inference(duplicate_literal_removal,[],[f2624])).
fof(f2624,plain,(
  ( ! [X10] : (s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl13 | $spl24)),
  inference(superposition,[],[f1072,f1167])).
fof(f2825,plain,(
  $false | ($spl7 | $spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2824,f80])).
fof(f2824,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl24 | $spl28)),
  inference(condensation,[],[f2823])).
fof(f2823,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl24 | $spl28)),
  inference(forward_demodulation,[],[f2818,f1076])).
fof(f2818,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f2619])).
fof(f2619,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f1072,f48])).
fof(f2822,plain,(
  $false | ($spl7 | $spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2821,f80])).
fof(f2821,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f2617,f1076])).
fof(f2617,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl7 | $spl24)),
  inference(unit_resulting_resolution,[],[f80,f1072])).
fof(f2820,plain,(
  $false | ($spl13 | $spl24 | $spl28)),
  inference(subsumption_resolution,[],[f2819,f860])).
fof(f2819,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl13 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f2618,f1076])).
fof(f2618,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | ($spl13 | $spl24)),
  inference(unit_resulting_resolution,[],[f860,f1072])).
fof(f1072,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f1072_D])).
fof(f1072_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1073,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl26),
  inference(cnf_transformation,[],[f1073_D])).
fof(f1073_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1628,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(superposition,[],[f72,f1053])).
fof(f1076,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl28),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1708,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f1708_D])).
fof(f1708_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1627,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl22),
  inference(superposition,[],[f72,f1053])).
fof(f1626,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl22),
  inference(superposition,[],[f72,f1053])).
fof(f1625,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f72,f1053])).
fof(f1078,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl32),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1624,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl22),
  inference(superposition,[],[f48,f1053])).
fof(f1623,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl22),
  inference(superposition,[],[f1053,f48])).
fof(f1622,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl22),
  inference(superposition,[],[f1053,f72])).
fof(f1621,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl22),
  inference(superposition,[],[f1053,f72])).
fof(f1620,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(superposition,[],[f1053,f72])).
fof(f1619,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f1053,f72])).
fof(f1079,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl33),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1069,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl23),
  inference(superposition,[],[f1054,f48])).
fof(f1077,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1068,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl23),
  inference(superposition,[],[f1054,f72])).
fof(f1067,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl23),
  inference(superposition,[],[f1054,f72])).
fof(f1074,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl27),
  inference(cnf_transformation,[],[f1074_D])).
fof(f1074_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1066,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl23),
  inference(superposition,[],[f1054,f72])).
fof(f1065,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f1054,f72])).
fof(f1052,plain,(
  ( ! [X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18)) ) | $spl20),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  ( ! [X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1054,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl23),
  inference(cnf_transformation,[],[f1054_D])).
fof(f1054_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1043,plain,(
  ( ! [X17,X18] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f1006])).
fof(f1006,plain,(
  ( ! [X17,X18] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18)) ) | $spl10),
  inference(superposition,[],[f59,f858])).
fof(f862,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl14),
  inference(cnf_transformation,[],[f862_D])).
fof(f862_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f865,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f856,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK3(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f529,f48])).
fof(f863,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl16),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f855,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK3(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f529,f72])).
fof(f854,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK3(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f529,f72])).
fof(f857,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f853,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK3(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f529,f72])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f529,f72])).
fof(f78,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f80,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f77,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f48])).
fof(f71,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f70])).
fof(f70,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f67,f68])).
fof(f66,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | $spl0),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f68,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) | $spl2),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f43,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)) & (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f31])).
fof(f31,plain,(
  ? [X0,X1,X2] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2)) & (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ? [X0,X1,X2] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) <~> s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),X8))),
  file('thm.complex.COMPLEX_EQ_LADD','thm.complex.COMPLEX_EQ_LADD')).
fof(f67,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) | $spl1),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f69,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) | $spl3),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f44,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK2) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK2)))),
  inference(cnf_transformation,[],[f32])).
% SZS output end Proof for thm.complex
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 106
% Passive clauses: 456
% Generated clauses: 11655
% Final active clauses: 69
% Final passive clauses: 173
% Input formulas: 13
% Initial clauses: 22
% 
% Duplicate literals: 506
% Trivial inequalities: 77
% Fw subsumption resolutions: 195
% Fw demodulations: 1613
% Bw demodulations: 113
% Condensations: 747
% Global subsumptions: 5
% 
% Simple tautologies: 122
% Equational tautologies: 92
% Forward subsumptions: 6975
% Fw demodulations to eq. taut.: 772
% Bw demodulations to eq. taut.: 17
% 
% Unit resulting resolution: 2221
% Forward superposition: 2584
% Backward superposition: 3065
% Self superposition: 101
% Equality factoring: 53
% Equality resolution: 27
% 
% Split clauses: 86
% Split components: 178
% SAT solver clauses: 649
% SAT solver unit clauses: 8
% SAT solver binary clauses: 154
% 
% Sat splits: 86
% Sat splitting refutations: 14
% 
% Memory used [KB]: 7803
% Time elapsed: 0.363 s
% ------------------------------
% ------------------------------
% Success in time 1.681 s
