Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (60676)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 643
% Passive clauses: 9841
% Generated clauses: 18081
% Final active clauses: 567
% Final passive clauses: 8685
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 278
% Trivial inequalities: 87
% Fw subsumption resolutions: 1502
% Fw demodulations: 106
% Bw demodulations: 75
% 
% Simple tautologies: 226
% Equational tautologies: 65
% Forward subsumptions: 5765
% 
% Binary resolution: 11269
% Factoring: 984
% Forward superposition: 1439
% Backward superposition: 1880
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 31
% 
% Split clauses: 294
% Split components: 590
% SAT solver clauses: 619
% SAT solver unit clauses: 24
% SAT solver binary clauses: 143
% 
% Sat splits: 294
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 617
% TWLsolver calls for satisfiability: 354
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (60702)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 443
% Passive clauses: 2991
% Generated clauses: 12862
% Final active clauses: 419
% Final passive clauses: 2131
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 223
% Trivial inequalities: 40
% Fw subsumption resolutions: 1054
% Fw demodulations: 612
% Bw demodulations: 438
% Condensations: 369
% Global subsumptions: 13
% 
% Simple tautologies: 289
% Equational tautologies: 18
% Forward subsumptions: 4240
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 4118
% Factoring: 215
% Forward superposition: 1312
% Backward superposition: 1206
% Self superposition: 26
% Equality factoring: 30
% Equality resolution: 10
% 
% Split clauses: 43
% Split components: 87
% SAT solver clauses: 3156
% SAT solver unit clauses: 182
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 43
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 3104
% TWLsolver calls for satisfiability: 1806
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (60743)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2490
% Passive clauses: 9497
% Generated clauses: 127621
% Final active clauses: 720
% Final passive clauses: 821
% Input formulas: 45
% Initial clauses: 202
% 
% Duplicate literals: 6771
% Trivial inequalities: 228
% Fw subsumption resolutions: 7153
% Fw demodulations: 15385
% Bw demodulations: 3405
% Condensations: 11348
% Global subsumptions: 439
% 
% Simple tautologies: 397
% Equational tautologies: 1756
% Forward subsumptions: 73458
% Bw demodulations to eq. taut.: 482
% 
% Unit resulting resolution: 27332
% Factoring: 12
% Forward superposition: 23839
% Backward superposition: 18564
% Self superposition: 271
% Equality factoring: 206
% Equality resolution: 8
% 
% Split clauses: 451
% Split components: 908
% SAT solver clauses: 11456
% SAT solver unit clauses: 71
% SAT solver binary clauses: 1776
% 
% Sat splits: 451
% Sat splitting refutations: 616
% 
% Memory used [KB]: 57568
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f83663,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f264,f470,f468,f263,f262,f469,f467,f1001,f1105,f1101,f1170,f1171,f1172,f1173,f1765,f4215,f4216,f4609,f4610,f5275,f6189,f6298,f6195,f11864,f11865,f10654,f10247,f11900,f11901,f10239,f10238,f11903,f10253,f10254,f10255,f1003,f11938,f11940,f11939,f5344,f5343,f5342,f1220,f1222,f1221,f13014,f13016,f13739,f13970,f13973,f13972,f12997,f12995,f14031,f13022,f14032,f13004,f14041,f14053,f14075,f14078,f14083,f14084,f14338,f14339,f14669,f13023,f14670,f13006,f26343,f26352,f26353,f26488,f26489,f26490,f26491,f26492,f26905,f26908,f26909,f26910,f26906,f26911,f26913,f26916,f26917,f26914,f26921,f26922,f26902,f26901,f26925,f26923,f26993,f26992,f26996,f26997,f27180,f27181,f27182,f27183,f27184,f27179,f28511,f28513,f28516,f28518,f28514,f28521,f28641,f28640,f28734,f28735,f28736,f28737,f28738,f28739,f28744,f28745,f28802,f28801,f28805,f28806,f29117,f29118,f29119,f29120,f29116,f29122,f29846,f29848,f30055,f30057,f30059,f30060,f30058,f30063,f31175,f31176,f31177,f31179,f31180,f31178,f31181,f31182,f31168,f31167,f31184,f31183,f31525,f31527,f31528,f31529,f31530,f31531,f31532,f31534,f31535,f31533,f31536,f31537,f31522,f31521,f31539,f31538,f36125,f32086,f32085,f31526,f31975,f31390,f31014,f30634,f36175,f36176,f29762,f29457,f27350,f26591,f26568,f36217,f36219,f36220,f36221,f36223,f36224,f36225,f36226,f36227,f36228,f36262,f36263,f20090,f19611,f19257,f36278,f36279,f36281,f36282,f13605,f36323,f36324,f36422,f36395,f14077,f36506,f36508,f36509,f36510,f36577,f36580,f36602,f36603,f36983,f36985,f36986,f36987,f36995,f37035,f37038,f37048,f37057,f36997,f37058,f37068,f36269,f37916,f37919,f37918,f38100,f6302,f38101,f38107,f38184,f38192,f39337,f39737,f39798,f41592,f44694,f44695,f44729,f44700,f44728,f44702,f44730,f45103,f45670,f28512,f45676,f28522,f45663,f45662,f45678,f45677,f46361,f30056,f46367,f30064,f47107,f47114,f47115,f47665,f47671,f47672,f48444,f48471,f48640,f48696,f48721,f49491,f49516,f49550,f49599,f49948,f45130,f50340,f50360,f50945,f50946,f50947,f50949,f50950,f50958,f50960,f51028,f51037,f51016,f51015,f51023,f51130,f51131,f51132,f51134,f51135,f51143,f51145,f51212,f51027,f51221,f51201,f51200,f51208,f52541,f52558,f53261,f53275,f53262,f53267,f53268,f53269,f53271,f53272,f53273,f53433,f53500,f53522,f53537,f53539,f53540,f53541,f53564,f53545,f53567,f53619,f53622,f53637,f53646,f53566,f53647,f53661,f53563,f53684,f53663,f53856,f53857,f53858,f53860,f53861,f53948,f53929,f53928,f53936,f54380,f54381,f54382,f54383,f54873,f54897,f55086,f55063,f55128,f55175,f55291,f55332,f55565,f55758,f55984,f56062,f56591,f56551,f56791,f56815,f58790,f58792,f58791,f58788,f58910,f58911,f58909,f60081,f48457,f60089,f53276,f60118,f60103,f60108,f60109,f60110,f60111,f60098,f60113,f60114,f60115,f60116,f60119,f59904,f59902,f59896,f59894,f59725,f59723,f59719,f59715,f58000,f57999,f57948,f55380,f54197,f60176,f54186,f60180,f50820,f50818,f50812,f50810,f50665,f50663,f50659,f50655,f49624,f49623,f49574,f48092,f60223,f47889,f36340,f36405,f47434,f36169,f47221,f47131,f36193,f46480,f46129,f60246,f46144,f46143,f46142,f60247,f60248,f45936,f45857,f45819,f29961,f28515,f29959,f29960,f41378,f40487,f40395,f36157,f60297,f60298,f60299,f60300,f38893,f60302,f38486,f60312,f37227,f37141,f36218,f60329,f60330,f36280,f60331,f60332,f36261,f36260,f26907,f60333,f60334,f32087,f60335,f60336,f60338,f60339,f60340,f31463,f60341,f60342,f60337,f30242,f30162,f60343,f29036,f29016,f28894,f60344,f27102,f27084,f60346,f20167,f60348,f60349,f26915,f17417,f28517,f12201,f12202,f12203,f12210,f59511,f13829,f47936,f47698,f47359,f47325,f47146,f47145,f42712,f36614,f42531,f42389,f42189,f42195,f42192,f60440,f42127,f36592,f41499,f36522,f39241,f36521,f60486,f60496,f11888,f60908,f60910,f60911,f60912,f60927,f61101,f61104,f13839,f61123,f13003,f61124,f61134,f1772,f61138,f61139,f58793,f61140,f58912,f61141,f61142,f61143,f61144,f61145,f61146,f61147,f61148,f61149,f61150,f61151,f61152,f61153,f61154,f61155,f61156,f61157,f61158,f61159,f61160,f11902,f61161,f61162,f61163,f61164,f61165,f61166,f471,f61167,f472,f61815,f61819,f61821,f61837,f61822,f61826,f4226,f61827,f11886,f61828,f61927,f61930,f61941,f62103,f62115,f62104,f61839,f62106,f62116,f4225,f72014,f72015,f72950,f71806,f72959,f71715,f72960,f68647,f68444,f72977,f72979,f72993,f72995,f72962,f73002,f73003,f72961,f35101,f34991,f73012,f73120,f73242,f73351,f73349,f73558,f73560,f73561,f73562,f73582,f73566,f73583,f73764,f73767,f73777,f73786,f13005,f73787,f73797,f73581,f74623,f74624,f74625,f74626,f75855,f75857,f75858,f75859,f75879,f75863,f75882,f76920,f76923,f76933,f77169,f75881,f77170,f77180,f75878,f80836,f80856,f80857,f80860,f80880,f80881,f80884,f80902,f80903,f80923,f80943,f80944,f80947,f80967,f80968,f80970,f81008,f81009,f80985,f80988,f80989,f80995,f81015,f80996,f81012,f81016,f81055,f81056,f81060,f81087,f81088,f81110,f81136,f81137,f81121,f81014,f81140,f81159,f81160,f80855,f81163,f81182,f81183,f80879,f81186,f81203,f81204,f80901,f81207,f81226,f81227,f80942,f81255,f81274,f81275,f80966,f81277,f81314,f81315,f81007,f81292,f81295,f81296,f81302,f81321,f81303,f81318,f81322,f81360,f81361,f81054,f81365,f81391,f81392,f81086,f81395,f81420,f81421,f81135,f81406,f81320,f81654,f81657,f81679,f81658,f81680,f81664,f81681,f81665,f81678,f82027,f82154,f82171,f82153,f82172,f82299,f82301,f82300,f82472,f82474,f82473,f82861,f81319,f82862,f81013,f82947,f82948,f82973,f14054,f82953,f82974,f82954,f82955,f82975,f83060,f83065,f83066,f83067,f83072,f83073,f73350,f83271,f83274,f83275,f83278,f83298,f11937,f83283,f83300,f83284,f83297,f83369,f83299,f1002,f83567,f72009,f26344,f72092,f26345,f26912,f82854,f29121,f83579,f82601,f82570,f36524,f72589,f72200,f60438,f4227])).
fof(f4227,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl16),
  inference(cnf_transformation,[],[f4227_D])).
fof(f4227_D,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f60438,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl334),
  inference(cnf_transformation,[],[f60438_D])).
fof(f60438_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f72200,plain,(
  ( ! [X180,X181,X179] : (~p(s(bool,X179)) | p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl94)),
  inference(duplicate_literal_removal,[],[f72163])).
fof(f72163,plain,(
  ( ! [X180,X181,X179] : (~p(s(bool,X179)) | p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X180) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X179)) | ~p(s(bool,X181))) ) | ($spl2 | $spl94)),
  inference(resolution,[],[f61754,f41496])).
fof(f41496,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X2) | ~p(s(bool,X2)) | ~p(s(bool,X3))) ) | $spl2),
  inference(resolution,[],[f36516,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f36516,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1,X0) | ~p(s(bool,X1)) | sP9(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f36500,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f36500,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1) | ~p(s(bool,X1)) | ~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X1,X0) | ~sP6(X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl2),
  inference(resolution,[],[f524,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f524,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(resolution,[],[f469,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f188])).
fof(f188,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f61754,plain,(
  ( ! [X10,X8,X11,X9] : (~sP9(X11,X9,X8) | ~p(s(bool,X8)) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl94),
  inference(superposition,[],[f49123,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f56,f152,f151,f150,f149])).
fof(f56,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUBSET_INTER','thm.sat.sat_2|49|_prove_ltb')).
fof(f49123,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl94),
  inference(resolution,[],[f36851,f36863])).
fof(f36863,plain,(
  ( ! [X118,X117] : (sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X117,X118) | s(bool,X117) != s(bool,X118)) ) | $spl94),
  inference(resolution,[],[f14077,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f36851,plain,(
  ( ! [X94,X93] : (~sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X93,X94) | p(s(bool,X94)) | ~p(s(bool,X93))) ) | $spl94),
  inference(resolution,[],[f14077,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f186])).
fof(f72589,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X18))) ) | ($spl2 | $spl94)),
  inference(resolution,[],[f72212,f524])).
fof(f72212,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X2,X0,X3) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0))) ) | $spl94),
  inference(subsumption_resolution,[],[f72211,f334])).
fof(f72211,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP6(X0,X3,X2)) ) | $spl94),
  inference(subsumption_resolution,[],[f72210,f328])).
fof(f72210,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl94),
  inference(duplicate_literal_removal,[],[f72110])).
fof(f72110,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0)) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl94),
  inference(resolution,[],[f61754,f320])).
fof(f36524,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl182),
  inference(cnf_transformation,[],[f36524_D])).
fof(f36524_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f82570,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X127))) ) | ($spl2 | $spl8 | $spl40)),
  inference(duplicate_literal_removal,[],[f82533])).
fof(f82533,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X127)) | ~p(s(bool,X127)) | ~p(s(bool,X126))) ) | ($spl2 | $spl8 | $spl40)),
  inference(resolution,[],[f82464,f41496])).
fof(f82464,plain,(
  ( ! [X47,X45,X46] : (~sP9(X47,X46,X45) | ~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X45))) ) | ($spl8 | $spl40)),
  inference(trivial_inequality_removal,[],[f82436])).
fof(f82436,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X46) | ~p(s(bool,X47)) | ~sP9(X47,X46,X45) | ~p(s(bool,X45))) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82215,f418])).
fof(f418,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.EQ_CLAUSES|split|1')).
fof(f82215,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP9(X11,X9,X10)) ) | ($spl8 | $spl40)),
  inference(forward_demodulation,[],[f82176,f13003])).
fof(f82176,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP9(X11,X9,X10)) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f62401,f337])).
fof(f62401,plain,(
  ( ! [X317] : (s(bool,X317) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X317)) ) | ($spl8 | $spl40)),
  inference(trivial_inequality_removal,[],[f62367])).
fof(f62367,plain,(
  ( ! [X317] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X317) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X317)) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f12160,f13003])).
fof(f12160,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1003,f442])).
fof(f442,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_INTER',aHL_BOOL_CASES)).
fof(f82601,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl2 | $spl8 | $spl40)),
  inference(resolution,[],[f82584,f524])).
fof(f82584,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X2,X1) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X0))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82583,f334])).
fof(f82583,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP6(X2,X1,X0)) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82582,f328])).
fof(f82582,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) ) | ($spl8 | $spl40)),
  inference(duplicate_literal_removal,[],[f82480])).
fof(f82480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X2)) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f82464,f320])).
fof(f83579,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X127))) ) | ($spl2 | $spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82818,f411])).
fof(f411,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.EQ_CLAUSES|split|3')).
fof(f82818,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X127) | ~p(s(bool,X127))) ) | ($spl2 | $spl8 | $spl40)),
  inference(duplicate_literal_removal,[],[f82785])).
fof(f82785,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X127) | ~p(s(bool,X127)) | ~p(s(bool,X126))) ) | ($spl2 | $spl8 | $spl40)),
  inference(resolution,[],[f82463,f41496])).
fof(f82463,plain,(
  ( ! [X50,X48,X49] : (~sP9(X50,X49,X48) | ~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) ) | ($spl8 | $spl40)),
  inference(trivial_inequality_removal,[],[f82437])).
fof(f82437,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X49) | ~p(s(bool,X50)) | ~sP9(X50,X49,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82215,f442])).
fof(f29121,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) | $spl126),
  inference(cnf_transformation,[],[f29121_D])).
fof(f29121_D,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f82854,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X12) | ~p(s(bool,X13))) ) | ($spl2 | $spl8 | $spl40)),
  inference(resolution,[],[f82833,f524])).
fof(f82833,plain,(
  ( ! [X4,X5,X3] : (~sP8(X3,X5,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X3))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82832,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f194])).
fof(f82832,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~sP8(X3,X5,X4) | ~sP6(X5,X4,X3)) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82831,f328])).
fof(f82831,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | ~sP8(X3,X5,X4) | ~sP7(X4,X5,X3) | ~sP6(X5,X4,X3)) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82733,f416])).
fof(f416,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.EQ_CLAUSES|split|0')).
fof(f82733,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X4)) | ~sP8(X3,X5,X4) | ~sP7(X4,X5,X3) | ~sP6(X5,X4,X3)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f82463,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f26912,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl106),
  inference(cnf_transformation,[],[f26912_D])).
fof(f26912_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f26345,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) | $spl98),
  inference(cnf_transformation,[],[f26345_D])).
fof(f26345_D,plain,(
  ( ! [X65] : (s(bool,'F_0') != s(bool,X65)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f72092,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl94),
  inference(subsumption_resolution,[],[f72042,f411])).
fof(f72042,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl94),
  inference(superposition,[],[f61753,f442])).
fof(f61753,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X4)) | p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl94),
  inference(superposition,[],[f49123,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f26344,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl96),
  inference(cnf_transformation,[],[f26344_D])).
fof(f26344_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f72009,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl94),
  inference(subsumption_resolution,[],[f71959,f411])).
fof(f71959,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl94),
  inference(superposition,[],[f61752,f442])).
fof(f61752,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X0)) | p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl94),
  inference(superposition,[],[f49123,f335])).
fof(f83567,plain,(
  $false | ($spl1 | $spl182 | $spl407)),
  inference(subsumption_resolution,[],[f83566,f73349])).
fof(f83566,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl1 | $spl182)),
  inference(subsumption_resolution,[],[f83543,f402])).
fof(f402,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_CLAUSES|split|3')).
fof(f83543,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl1 | $spl182)),
  inference(superposition,[],[f62899,f63218])).
fof(f63218,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl1),
  inference(resolution,[],[f60499,f416])).
fof(f60499,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl1),
  inference(resolution,[],[f468,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_INTER','thm.pred_set.SUBSET_DEF')).
fof(f62899,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))) ) | $spl182),
  inference(subsumption_resolution,[],[f62653,f403])).
fof(f403,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f62653,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl182),
  inference(superposition,[],[f434,f36524])).
fof(f434,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f1002,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl6),
  inference(cnf_transformation,[],[f1002_D])).
fof(f1002_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f83299,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)))) ) | $spl496),
  inference(cnf_transformation,[],[f83299_D])).
fof(f83299_D,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f83369,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl21 | $spl40)),
  inference(resolution,[],[f81113,f13485])).
fof(f13485,plain,(
  ( ! [X4,X5,X3] : (~sP7(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13463])).
fof(f13463,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP7(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl8),
  inference(resolution,[],[f13026,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f13026,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X8,X6) | ~p(s(bool,X6)) | ~sP7(X8,X6,X7) | p(s(bool,X7))) ) | $spl8),
  inference(resolution,[],[f12093,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f189])).
fof(f12093,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f12087,f416])).
fof(f12087,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl8),
  inference(resolution,[],[f11558,f320])).
fof(f11558,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl8),
  inference(subsumption_resolution,[],[f11557,f332])).
fof(f11557,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f11522,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f194])).
fof(f11522,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl8),
  inference(resolution,[],[f11482,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f186])).
fof(f11482,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl8),
  inference(resolution,[],[f11435,f326])).
fof(f11435,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,X6) | sP9(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) ) | $spl8),
  inference(subsumption_resolution,[],[f11434,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f11434,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f11424,f412])).
fof(f412,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f11424,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl8),
  inference(resolution,[],[f11408,f321])).
fof(f11408,plain,(
  ( ! [X6,X8,X7] : (sP6(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(subsumption_resolution,[],[f11376,f332])).
fof(f11376,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X6,X7,X8)) ) | $spl8),
  inference(resolution,[],[f5323,f316])).
fof(f5323,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f1003,f335])).
fof(f81113,plain,(
  ( ! [X14,X12,X13] : (sP7('const.bool.IN_2'(s(sK12,X12),s(fun(sK12,bool),sK13)),X13,X14) | p(s(bool,'const.bool.IN_2'(s(sK12,X12),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl40)),
  inference(resolution,[],[f79863,f38026])).
fof(f38026,plain,(
  ( ! [X130] : (~sP11('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X130) | p(s(bool,X130))) ) | $spl21),
  inference(subsumption_resolution,[],[f37991,f427])).
fof(f427,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_INTER',aHL_FALSITY)).
fof(f37991,plain,(
  ( ! [X130] : (p(s(bool,'F_0')) | p(s(bool,X130)) | ~sP11('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X130)) ) | $spl21),
  inference(resolution,[],[f37914,f375])).
fof(f375,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f37914,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl21),
  inference(trivial_inequality_removal,[],[f37910])).
fof(f37910,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl21),
  inference(superposition,[],[f6195,f418])).
fof(f79863,plain,(
  ( ! [X227,X225,X226,X224] : (sP11(X224,'const.bool.IN_2'(s(sK12,X225),s(fun(sK12,bool),sK15))) | sP7('const.bool.IN_2'(s(sK12,X225),s(fun(sK12,bool),sK13)),X226,X227)) ) | $spl40),
  inference(resolution,[],[f72908,f330])).
fof(f72908,plain,(
  ( ! [X200,X201] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),sK13)))) | sP11(X201,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),sK15)))) ) | $spl40),
  inference(resolution,[],[f62415,f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f211])).
fof(f62415,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl40),
  inference(subsumption_resolution,[],[f62169,f403])).
fof(f62169,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | $spl40),
  inference(superposition,[],[f434,f13003])).
fof(f83297,plain,(
  ( ! [X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) | $spl492),
  inference(cnf_transformation,[],[f83297_D])).
fof(f83297_D,plain,(
  ( ! [X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f83284,plain,(
  ( ! [X68,X67] : (p(s(bool,'const.bool.IN_2'(s(sK12,X67),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X67),s(fun(sK12,bool),sK13))) | ~p(s(bool,X68))) ) | ($spl4 | $spl21 | $spl40)),
  inference(resolution,[],[f80973,f61709])).
fof(f61709,plain,(
  ( ! [X12,X13] : (~sP8(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl4),
  inference(subsumption_resolution,[],[f61708,f332])).
fof(f61708,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f61707,f397])).
fof(f397,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_CLAUSES|split|4')).
fof(f61707,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f61654,f330])).
fof(f61654,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP7('F_0',X13,X12) | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(resolution,[],[f60618,f321])).
fof(f60618,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl4),
  inference(superposition,[],[f1821,f337])).
fof(f1821,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(backward_demodulation,[],[f660,f658])).
fof(f658,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl4),
  inference(resolution,[],[f471,f411])).
fof(f660,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl4),
  inference(resolution,[],[f471,f416])).
fof(f80973,plain,(
  ( ! [X14,X12,X13] : (sP8(X12,'const.bool.IN_2'(s(sK12,X13),s(fun(sK12,bool),sK13)),X14) | p(s(bool,'const.bool.IN_2'(s(sK12,X13),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl40)),
  inference(resolution,[],[f79859,f38026])).
fof(f79859,plain,(
  ( ! [X210,X208,X211,X209] : (sP11(X208,'const.bool.IN_2'(s(sK12,X209),s(fun(sK12,bool),sK15))) | sP8(X210,'const.bool.IN_2'(s(sK12,X209),s(fun(sK12,bool),sK13)),X211)) ) | $spl40),
  inference(resolution,[],[f72908,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f189])).
fof(f83300,plain,(
  ( ! [X66,X64] : (p(s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') != s(bool,X66) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK13)))) ) | $spl498),
  inference(cnf_transformation,[],[f83300_D])).
fof(f83300_D,plain,(
  ( ! [X66,X64] : (p(s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') != s(bool,X66) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK13)))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f83283,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X65)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X66)) ) | ($spl8 | $spl21 | $spl40)),
  inference(resolution,[],[f80973,f21063])).
fof(f21063,plain,(
  ( ! [X21,X22,X20] : (~sP8(X21,X22,'F_0') | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,'T_0') != s(bool,X20)) ) | $spl8),
  inference(subsumption_resolution,[],[f21062,f397])).
fof(f21062,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | p(s(bool,'F_0')) | ~sP8(X21,X22,'F_0')) ) | $spl8),
  inference(subsumption_resolution,[],[f21061,f20950])).
fof(f20950,plain,(
  ( ! [X14,X12,X13] : (sP6(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) ) | $spl8),
  inference(subsumption_resolution,[],[f20931,f332])).
fof(f20931,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP6(X13,'F_0',X14)) ) | $spl8),
  inference(resolution,[],[f12164,f316])).
fof(f12164,plain,(
  ( ! [X10,X11,X9] : (sP9(X11,'F_0',X9) | s(bool,'T_0') != s(bool,X10) | s(bool,X9) != s(bool,X10) | p(s(bool,X11))) ) | $spl8),
  inference(superposition,[],[f1003,f335])).
fof(f21061,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | p(s(bool,'F_0')) | ~sP8(X21,X22,'F_0') | ~sP6(X22,'F_0',X21)) ) | $spl8),
  inference(subsumption_resolution,[],[f21029,f330])).
fof(f21029,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | p(s(bool,'F_0')) | ~sP8(X21,X22,'F_0') | ~sP7('F_0',X22,X21) | ~sP6(X22,'F_0',X21)) ) | $spl8),
  inference(resolution,[],[f12166,f321])).
fof(f12166,plain,(
  ( ! [X17,X15,X16] : (~sP9(X17,'F_0',X15) | s(bool,'T_0') != s(bool,X16) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X16)) ) | $spl8),
  inference(superposition,[],[f1003,f337])).
fof(f11937,plain,(
  ( ! [X43] : (s(bool,'T_0') = s(bool,X43)) ) | $spl36),
  inference(cnf_transformation,[],[f11937_D])).
fof(f11937_D,plain,(
  ( ! [X43] : (s(bool,'T_0') = s(bool,X43)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f83298,plain,(
  ( ! [X38] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13)))) ) | $spl494),
  inference(cnf_transformation,[],[f83298_D])).
fof(f83298_D,plain,(
  ( ! [X38] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13)))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f83278,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK15)))) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK13))) | ~p(s(bool,X54))) ) | ($spl8 | $spl21 | $spl40)),
  inference(resolution,[],[f80973,f82833])).
fof(f83275,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK15)))) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13))) | ~p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,X44))) ) | ($spl21 | $spl40 | $spl152)),
  inference(resolution,[],[f80973,f71583])).
fof(f71583,plain,(
  ( ! [X6,X4,X7,X5] : (~sP8(X7,X6,X4) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X5)) | p(s(bool,X4))) ) | $spl152),
  inference(subsumption_resolution,[],[f71582,f332])).
fof(f71582,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X5)) | ~sP8(X7,X6,X4) | ~sP6(X6,X4,X7)) ) | $spl152),
  inference(subsumption_resolution,[],[f71581,f330])).
fof(f71581,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X5)) | ~sP8(X7,X6,X4) | ~sP7(X4,X6,X7) | ~sP6(X6,X4,X7)) ) | $spl152),
  inference(duplicate_literal_removal,[],[f71480])).
fof(f71480,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | s(bool,X5) != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X5)) | p(s(bool,X4)) | ~sP8(X7,X6,X4) | ~sP7(X4,X6,X7) | ~sP6(X6,X4,X7)) ) | $spl152),
  inference(resolution,[],[f32031,f321])).
fof(f32031,plain,(
  ( ! [X6,X4,X7,X5] : (~sP9(X7,X6,X5) | p(s(bool,X6)) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X7)) | ~p(s(bool,X4))) ) | $spl152),
  inference(superposition,[],[f31538,f418])).
fof(f83274,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X39)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13))) | p(s(bool,X40))) ) | ($spl21 | $spl40 | $spl152)),
  inference(resolution,[],[f80973,f40490])).
fof(f40490,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X5,X3) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X3))) ) | $spl152),
  inference(subsumption_resolution,[],[f40489,f332])).
fof(f40489,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X5) | ~sP8(X4,X5,X3) | ~sP6(X5,X3,X4)) ) | $spl152),
  inference(subsumption_resolution,[],[f40488,f330])).
fof(f40488,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X5) | ~sP8(X4,X5,X3) | ~sP7(X3,X5,X4) | ~sP6(X5,X3,X4)) ) | $spl152),
  inference(duplicate_literal_removal,[],[f40428])).
fof(f40428,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X3)) | ~sP8(X4,X5,X3) | ~sP7(X3,X5,X4) | ~sP6(X5,X3,X4)) ) | $spl152),
  inference(resolution,[],[f32078,f321])).
fof(f32078,plain,(
  ( ! [X50,X48,X49] : (~sP9(X50,X49,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X48)) ) | $spl152),
  inference(trivial_inequality_removal,[],[f32056])).
fof(f32056,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X49)) | ~sP9(X50,X49,X48) | ~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X48)) ) | $spl152),
  inference(superposition,[],[f31538,f442])).
fof(f83271,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl21 | $spl40)),
  inference(resolution,[],[f80973,f21264])).
fof(f21264,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f21263,f328])).
fof(f21263,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f21262,f411])).
fof(f21262,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f21203,f11558])).
fof(f21203,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4) | ~sP6(X3,X5,X4)) ) | $spl8),
  inference(resolution,[],[f12179,f321])).
fof(f12179,plain,(
  ( ! [X6,X8,X7] : (~sP9(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl8),
  inference(superposition,[],[f1003,f337])).
fof(f73350,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl408),
  inference(cnf_transformation,[],[f73350_D])).
fof(f73350_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f83073,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X2))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f83069,f427])).
fof(f83069,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X2))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f79922,f375])).
fof(f79922,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)),X0)) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f72909,f60499])).
fof(f72909,plain,(
  ( ! [X202,X203] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),sK13)))) | sP11('const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),sK15)),X203)) ) | $spl40),
  inference(resolution,[],[f62415,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f83072,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X1))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f83068,f427])).
fof(f83068,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X1))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f79922,f379])).
fof(f379,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f123,f156])).
fof(f123,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X17,X6] : ((~(p(s(bool,X6)) | p(s(bool,X17))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X17)) => p(s(bool,'F_0')))))),
  file('thm.pred_set.SUBSET_INTER','thm.sat.sat_9_prove_ltb')).
fof(f83067,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X0))) ) | ($spl1 | $spl40 | $spl256)),
  inference(resolution,[],[f79922,f51027])).
fof(f83066,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f83062,f427])).
fof(f83062,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f79805,f375])).
fof(f79805,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f72908,f60499])).
fof(f83065,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f83061,f427])).
fof(f83061,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f79805,f379])).
fof(f83060,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl1 | $spl40 | $spl256)),
  inference(resolution,[],[f79805,f51027])).
fof(f82975,plain,(
  ( ! [X47] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK15)))) ) | $spl490),
  inference(cnf_transformation,[],[f82975_D])).
fof(f82975_D,plain,(
  ( ! [X47] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK15)))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f82955,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X50),s(fun(sK12,bool),sK13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X50),s(fun(sK12,bool),sK15))) | s(bool,'F_0') = s(bool,X51) | ~p(s(bool,X52))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f72891,f82833])).
fof(f72891,plain,(
  ( ! [X158,X159,X157] : (sP8(X158,X159,'const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),sK15))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),sK13))))) ) | $spl40),
  inference(resolution,[],[f62415,f324])).
fof(f82954,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),sK15))) | ~p(s(bool,X48)) | ~p(s(bool,X49))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f72891,f82584])).
fof(f82974,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13)))) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK15)))) ) | $spl488),
  inference(cnf_transformation,[],[f82974_D])).
fof(f82974_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13)))) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK15)))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f82953,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13)))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK15))) | ~p(s(bool,X46))) ) | ($spl40 | $spl94)),
  inference(resolution,[],[f72891,f72212])).
fof(f14054,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14)) ) | $spl90),
  inference(cnf_transformation,[],[f14054_D])).
fof(f14054_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f82973,plain,(
  ( ! [X24] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK13)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK15)))) ) | $spl486),
  inference(cnf_transformation,[],[f82973_D])).
fof(f82973_D,plain,(
  ( ! [X24] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK13)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK15)))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f82948,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,X29)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f72891,f21264])).
fof(f82947,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK13)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK15))) | ~p(s(bool,X25)) | ~p(s(bool,X26))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f72891,f21261])).
fof(f21261,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f21260,f328])).
fof(f21260,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f21259,f416])).
fof(f21259,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f21202,f11558])).
fof(f21202,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1) | ~sP6(X0,X2,X1)) ) | $spl8),
  inference(resolution,[],[f12179,f320])).
fof(f81013,plain,(
  ( ! [X52,X51] : (sP11(X51,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK15))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK13)))) ) | $spl450),
  inference(cnf_transformation,[],[f81013_D])).
fof(f81013_D,plain,(
  ( ! [X52,X51] : (sP11(X51,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK15))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK13)))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f82862,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),sK13))) | ~p(s(bool,X32)) | sP11(X33,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),sK15)))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f82833,f79859])).
fof(f81319,plain,(
  ( ! [X52,X51] : (sP11('const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK15)),X52) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK13)))) ) | $spl464),
  inference(cnf_transformation,[],[f81319_D])).
fof(f81319_D,plain,(
  ( ! [X52,X51] : (sP11('const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK15)),X52) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK13)))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f82861,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13))) | ~p(s(bool,X28)) | sP11('const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK15)),X29)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f82833,f79976])).
fof(f79976,plain,(
  ( ! [X210,X208,X211,X209] : (sP11('const.bool.IN_2'(s(sK12,X208),s(fun(sK12,bool),sK15)),X209) | sP8(X210,'const.bool.IN_2'(s(sK12,X208),s(fun(sK12,bool),sK13)),X211)) ) | $spl40),
  inference(resolution,[],[f72909,f325])).
fof(f82473,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP9(X31,X30,X28) | ~p(s(bool,X31))) ) | $spl484),
  inference(cnf_transformation,[],[f82473_D])).
fof(f82473_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP9(X31,X30,X28) | ~p(s(bool,X31))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f82474,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP9(X35,X34,X32) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82432,f416])).
fof(f82432,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP9(X35,X34,X32) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82215,f391])).
fof(f391,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_ANTISYM_AX')).
fof(f82472,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP9(X31,X30,X28) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82431,f416])).
fof(f82431,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP9(X31,X30,X28) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82215,f391])).
fof(f82300,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP9(X31,X28,X30) | ~p(s(bool,X31))) ) | $spl482),
  inference(cnf_transformation,[],[f82300_D])).
fof(f82300_D,plain,(
  ( ! [X30,X28,X31] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X28)) | ~sP9(X31,X28,X30) | ~p(s(bool,X31))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f82301,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP9(X35,X32,X34) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82259,f416])).
fof(f82259,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X34) | ~p(s(bool,X35)) | ~sP9(X35,X32,X34) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82214,f391])).
fof(f82214,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'T_0') = s(bool,X6) | ~p(s(bool,X8)) | ~sP9(X8,X7,X6)) ) | ($spl8 | $spl40)),
  inference(forward_demodulation,[],[f82175,f13003])).
fof(f82175,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X6) | ~p(s(bool,X8)) | ~sP9(X8,X7,X6)) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f62401,f337])).
fof(f82299,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP9(X31,X28,X30) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl40)),
  inference(subsumption_resolution,[],[f82258,f416])).
fof(f82258,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X30) | ~p(s(bool,X31)) | ~sP9(X31,X28,X30) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | ($spl8 | $spl40)),
  inference(superposition,[],[f82214,f391])).
fof(f82172,plain,(
  ( ! [X347,X348] : (p(s(bool,X347)) | p(s(bool,X348)) | sP9(X348,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X347)) ) | $spl480),
  inference(cnf_transformation,[],[f82172_D])).
fof(f82172_D,plain,(
  ( ! [X347,X348] : (p(s(bool,X347)) | p(s(bool,X348)) | sP9(X348,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X347)) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f82153,plain,(
  ( ! [X346,X347,X348] : (~p(s(bool,X346)) | p(s(bool,X347)) | sP9(X348,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X347) | p(s(bool,X348))) ) | ($spl1 | $spl94)),
  inference(subsumption_resolution,[],[f81976,f416])).
fof(f81976,plain,(
  ( ! [X346,X347,X348] : (s(bool,'T_0') != s(bool,X346) | ~p(s(bool,X346)) | p(s(bool,X347)) | sP9(X348,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X347) | p(s(bool,X348))) ) | ($spl1 | $spl94)),
  inference(superposition,[],[f61728,f63218])).
fof(f61728,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl94),
  inference(superposition,[],[f49123,f335])).
fof(f82171,plain,(
  ( ! [X354,X353] : (p(s(bool,X353)) | p(s(bool,X354)) | sP9(X354,X353,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl478),
  inference(cnf_transformation,[],[f82171_D])).
fof(f82171_D,plain,(
  ( ! [X354,X353] : (p(s(bool,X353)) | p(s(bool,X354)) | sP9(X354,X353,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f82154,plain,(
  ( ! [X352,X354,X353] : (~p(s(bool,X352)) | p(s(bool,X353)) | sP9(X354,X353,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X354))) ) | ($spl1 | $spl94)),
  inference(subsumption_resolution,[],[f81978,f416])).
fof(f81978,plain,(
  ( ! [X352,X354,X353] : (s(bool,'T_0') != s(bool,X352) | ~p(s(bool,X352)) | p(s(bool,X353)) | sP9(X354,X353,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X354))) ) | ($spl1 | $spl94)),
  inference(superposition,[],[f61729,f63218])).
fof(f61729,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X6)) | p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl94),
  inference(superposition,[],[f49123,f335])).
fof(f82027,plain,(
  $false | ($spl1 | $spl266 | $spl407)),
  inference(subsumption_resolution,[],[f82026,f73349])).
fof(f82026,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl1 | $spl266)),
  inference(subsumption_resolution,[],[f81757,f402])).
fof(f81757,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | ($spl1 | $spl266)),
  inference(superposition,[],[f53276,f63218])).
fof(f81678,plain,(
  ( ! [X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK14)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) | $spl470),
  inference(cnf_transformation,[],[f81678_D])).
fof(f81678_D,plain,(
  ( ! [X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK14)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f81665,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.bool.IN_2'(s(sK12,X61),s(fun(sK12,bool),sK14)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X61),s(fun(sK12,bool),sK13))) | ~p(s(bool,X62))) ) | ($spl4 | $spl266)),
  inference(resolution,[],[f72769,f61709])).
fof(f72769,plain,(
  ( ! [X156,X154,X155] : (sP8(X155,'const.bool.IN_2'(s(sK12,X154),s(fun(sK12,bool),sK13)),X156) | p(s(bool,'const.bool.IN_2'(s(sK12,X154),s(fun(sK12,bool),sK14))))) ) | $spl266),
  inference(resolution,[],[f53276,f325])).
fof(f81681,plain,(
  ( ! [X60,X58] : (p(s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK14)))) | s(bool,'T_0') != s(bool,X60) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK13)))) ) | $spl476),
  inference(cnf_transformation,[],[f81681_D])).
fof(f81681_D,plain,(
  ( ! [X60,X58] : (p(s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK14)))) | s(bool,'T_0') != s(bool,X60) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK13)))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f81664,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X59)) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X60)) ) | ($spl8 | $spl266)),
  inference(resolution,[],[f72769,f21063])).
fof(f81680,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)))) ) | $spl474),
  inference(cnf_transformation,[],[f81680_D])).
fof(f81680_D,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X42)) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f81658,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK14)))) | s(bool,X42) != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13))) | ~p(s(bool,X43)) | ~p(s(bool,X42)) | p(s(bool,X44))) ) | ($spl152 | $spl266)),
  inference(resolution,[],[f72769,f71583])).
fof(f81679,plain,(
  ( ! [X38] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13)))) ) | $spl472),
  inference(cnf_transformation,[],[f81679_D])).
fof(f81679_D,plain,(
  ( ! [X38] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK14)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13)))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f81657,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X39)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X38),s(fun(sK12,bool),sK13))) | p(s(bool,X40))) ) | ($spl152 | $spl266)),
  inference(resolution,[],[f72769,f40490])).
fof(f81654,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl266)),
  inference(resolution,[],[f72769,f21264])).
fof(f81320,plain,(
  ( ! [X57,X56,X55] : (sP11('const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK15)),X56) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK13)))) ) | $spl466),
  inference(cnf_transformation,[],[f81320_D])).
fof(f81320_D,plain,(
  ( ! [X57,X56,X55] : (sP11('const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK15)),X56) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK13)))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f81406,plain,(
  ( ! [X26,X24,X23,X25] : (sP11('const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),sK15)),X24) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79980,f13485])).
fof(f79980,plain,(
  ( ! [X227,X225,X226,X224] : (sP11('const.bool.IN_2'(s(sK12,X224),s(fun(sK12,bool),sK15)),X225) | sP7('const.bool.IN_2'(s(sK12,X224),s(fun(sK12,bool),sK13)),X226,X227)) ) | $spl40),
  inference(resolution,[],[f72909,f330])).
fof(f81135,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl460),
  inference(cnf_transformation,[],[f81135_D])).
fof(f81135_D,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f81421,plain,(
  ( ! [X10,X8,X11,X9] : (sP7('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81397,f427])).
fof(f81397,plain,(
  ( ! [X10,X8,X11,X9] : (sP7('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79980,f375])).
fof(f81420,plain,(
  ( ! [X6,X4,X7,X5] : (sP7('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81396,f427])).
fof(f81396,plain,(
  ( ! [X6,X4,X7,X5] : (sP7('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79980,f379])).
fof(f81395,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79980,f51027])).
fof(f81086,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | $spl458),
  inference(cnf_transformation,[],[f81086_D])).
fof(f81086_D,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f81392,plain,(
  ( ! [X10,X8,X11,X9] : (sP7(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81367,f427])).
fof(f81367,plain,(
  ( ! [X10,X8,X11,X9] : (sP7(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79979,f375])).
fof(f79979,plain,(
  ( ! [X222,X220,X223,X221] : (sP11('const.bool.IN_2'(s(sK12,X220),s(fun(sK12,bool),sK15)),X221) | sP7(X222,'const.bool.IN_2'(s(sK12,X220),s(fun(sK12,bool),sK13)),X223)) ) | $spl40),
  inference(resolution,[],[f72909,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f81391,plain,(
  ( ! [X6,X4,X7,X5] : (sP7(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81366,f427])).
fof(f81366,plain,(
  ( ! [X6,X4,X7,X5] : (sP7(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79979,f379])).
fof(f81365,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79979,f51027])).
fof(f81054,plain,(
  ( ! [X2,X0,X1] : (sP8('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl456),
  inference(cnf_transformation,[],[f81054_D])).
fof(f81054_D,plain,(
  ( ! [X2,X0,X1] : (sP8('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f81361,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81324,f427])).
fof(f81324,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79977,f375])).
fof(f79977,plain,(
  ( ! [X212,X215,X213,X214] : (sP11('const.bool.IN_2'(s(sK12,X212),s(fun(sK12,bool),sK15)),X213) | sP8('const.bool.IN_2'(s(sK12,X212),s(fun(sK12,bool),sK13)),X214,X215)) ) | $spl40),
  inference(resolution,[],[f72909,f326])).
fof(f81360,plain,(
  ( ! [X6,X4,X7,X5] : (sP8('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81323,f427])).
fof(f81323,plain,(
  ( ! [X6,X4,X7,X5] : (sP8('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79977,f379])).
fof(f81322,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79977,f51027])).
fof(f81318,plain,(
  ( ! [X39,X40] : (sP11('const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK15)),X40) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK13)))) ) | $spl462),
  inference(cnf_transformation,[],[f81318_D])).
fof(f81318_D,plain,(
  ( ! [X39,X40] : (sP11('const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK15)),X40) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK13)))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f81303,plain,(
  ( ! [X83,X84,X82] : (sP11('const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK15)),X83) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X82),s(fun(sK12,bool),sK13))) | ~p(s(bool,X84))) ) | ($spl4 | $spl40)),
  inference(resolution,[],[f79976,f61709])).
fof(f81321,plain,(
  ( ! [X78,X81,X79] : (sP11('const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK15)),X79) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK13)))) ) | $spl468),
  inference(cnf_transformation,[],[f81321_D])).
fof(f81321_D,plain,(
  ( ! [X78,X81,X79] : (sP11('const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK15)),X79) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK13)))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f81302,plain,(
  ( ! [X80,X78,X81,X79] : (sP11('const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK15)),X79) | ~p(s(bool,X80)) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X78),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X81)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79976,f21063])).
fof(f81296,plain,(
  ( ! [X59,X57,X58,X56,X55] : (sP11('const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK15)),X56) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK13))) | ~p(s(bool,X58)) | ~p(s(bool,X57)) | p(s(bool,X59))) ) | ($spl40 | $spl152)),
  inference(resolution,[],[f79976,f71583])).
fof(f81295,plain,(
  ( ! [X54,X52,X53,X51] : (sP11('const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK15)),X52) | ~p(s(bool,X53)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X51),s(fun(sK12,bool),sK13))) | p(s(bool,X54))) ) | ($spl40 | $spl152)),
  inference(resolution,[],[f79976,f40490])).
fof(f81292,plain,(
  ( ! [X39,X41,X42,X40] : (sP11('const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK15)),X40) | ~p(s(bool,X41)) | s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79976,f21264])).
fof(f81007,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | $spl446),
  inference(cnf_transformation,[],[f81007_D])).
fof(f81007_D,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f81315,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81279,f427])).
fof(f81279,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79976,f375])).
fof(f81314,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81278,f427])).
fof(f81278,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79976,f379])).
fof(f81277,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79976,f51027])).
fof(f80966,plain,(
  ( ! [X2,X0,X1] : (sP4('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl444),
  inference(cnf_transformation,[],[f80966_D])).
fof(f80966_D,plain,(
  ( ! [X2,X0,X1] : (sP4('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f81275,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81257,f427])).
fof(f81257,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79964,f375])).
fof(f79964,plain,(
  ( ! [X163,X161,X162,X160] : (sP11('const.bool.IN_2'(s(sK12,X160),s(fun(sK12,bool),sK15)),X161) | sP4('const.bool.IN_2'(s(sK12,X160),s(fun(sK12,bool),sK13)),X162,X163)) ) | $spl40),
  inference(resolution,[],[f72909,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f81274,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81256,f427])).
fof(f81256,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79964,f379])).
fof(f81255,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79964,f51027])).
fof(f80942,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | $spl442),
  inference(cnf_transformation,[],[f80942_D])).
fof(f80942_D,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f81227,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81209,f427])).
fof(f81209,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79963,f375])).
fof(f79963,plain,(
  ( ! [X158,X156,X159,X157] : (sP11('const.bool.IN_2'(s(sK12,X156),s(fun(sK12,bool),sK15)),X157) | sP4(X158,'const.bool.IN_2'(s(sK12,X156),s(fun(sK12,bool),sK13)),X159)) ) | $spl40),
  inference(resolution,[],[f72909,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f81226,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81208,f427])).
fof(f81208,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79963,f379])).
fof(f81207,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79963,f51027])).
fof(f80901,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl440),
  inference(cnf_transformation,[],[f80901_D])).
fof(f80901_D,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f81204,plain,(
  ( ! [X10,X8,X11,X9] : (sP2('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81188,f427])).
fof(f81188,plain,(
  ( ! [X10,X8,X11,X9] : (sP2('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79950,f375])).
fof(f79950,plain,(
  ( ! [X107,X105,X106,X104] : (sP11('const.bool.IN_2'(s(sK12,X104),s(fun(sK12,bool),sK15)),X105) | sP2('const.bool.IN_2'(s(sK12,X104),s(fun(sK12,bool),sK13)),X106,X107)) ) | $spl40),
  inference(resolution,[],[f72909,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f173])).
fof(f173,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f81203,plain,(
  ( ! [X6,X4,X7,X5] : (sP2('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81187,f427])).
fof(f81187,plain,(
  ( ! [X6,X4,X7,X5] : (sP2('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79950,f379])).
fof(f81186,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79950,f51027])).
fof(f80879,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl438),
  inference(cnf_transformation,[],[f80879_D])).
fof(f80879_D,plain,(
  ( ! [X2,X0,X1] : (sP0('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f81183,plain,(
  ( ! [X10,X8,X11,X9] : (sP0('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81165,f427])).
fof(f81165,plain,(
  ( ! [X10,X8,X11,X9] : (sP0('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79934,f375])).
fof(f79934,plain,(
  ( ! [X43,X41,X42,X40] : (sP11('const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK15)),X41) | sP0('const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK13)),X42,X43)) ) | $spl40),
  inference(resolution,[],[f72909,f275])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f165])).
fof(f165,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f81182,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81164,f427])).
fof(f81164,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79934,f379])).
fof(f81163,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79934,f51027])).
fof(f80855,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15))))) ) | $spl436),
  inference(cnf_transformation,[],[f80855_D])).
fof(f80855_D,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15))))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f81160,plain,(
  ( ! [X10,X8,X11,X9] : (sP0(X8,X9,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(subsumption_resolution,[],[f81142,f427])).
fof(f81142,plain,(
  ( ! [X10,X8,X11,X9] : (sP0(X8,X9,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13))) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK15)))) | p(s(bool,X11))) ) | $spl40),
  inference(resolution,[],[f79933,f375])).
fof(f79933,plain,(
  ( ! [X39,X37,X38,X36] : (sP11('const.bool.IN_2'(s(sK12,X36),s(fun(sK12,bool),sK15)),X37) | sP0(X38,X39,'const.bool.IN_2'(s(sK12,X36),s(fun(sK12,bool),sK13)))) ) | $spl40),
  inference(resolution,[],[f72909,f273])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f166])).
fof(f81159,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4,X5,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(subsumption_resolution,[],[f81141,f427])).
fof(f81141,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4,X5,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK13))) | p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK15)))) | p(s(bool,X7))) ) | $spl40),
  inference(resolution,[],[f79933,f379])).
fof(f81140,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15)))) | p(s(bool,X3))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79933,f51027])).
fof(f81014,plain,(
  ( ! [X57,X56,X55] : (sP11(X55,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK15))) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK13)))) ) | $spl452),
  inference(cnf_transformation,[],[f81014_D])).
fof(f81014_D,plain,(
  ( ! [X57,X56,X55] : (sP11(X55,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK15))) | ~p(s(bool,X57)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK13)))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f81121,plain,(
  ( ! [X26,X24,X23,X25] : (sP11(X23,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK15))) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79863,f13485])).
fof(f81137,plain,(
  ( ! [X10,X8,X11,X9] : (sP7('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81112,f427])).
fof(f81112,plain,(
  ( ! [X10,X8,X11,X9] : (sP7('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79863,f375])).
fof(f81136,plain,(
  ( ! [X6,X4,X7,X5] : (sP7('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81111,f427])).
fof(f81111,plain,(
  ( ! [X6,X4,X7,X5] : (sP7('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79863,f379])).
fof(f81110,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79863,f51027])).
fof(f81088,plain,(
  ( ! [X10,X8,X11,X9] : (sP7(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81062,f427])).
fof(f81062,plain,(
  ( ! [X10,X8,X11,X9] : (sP7(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79862,f375])).
fof(f79862,plain,(
  ( ! [X222,X220,X223,X221] : (sP11(X220,'const.bool.IN_2'(s(sK12,X221),s(fun(sK12,bool),sK15))) | sP7(X222,'const.bool.IN_2'(s(sK12,X221),s(fun(sK12,bool),sK13)),X223)) ) | $spl40),
  inference(resolution,[],[f72908,f329])).
fof(f81087,plain,(
  ( ! [X6,X4,X7,X5] : (sP7(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81061,f427])).
fof(f81061,plain,(
  ( ! [X6,X4,X7,X5] : (sP7(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79862,f379])).
fof(f81060,plain,(
  ( ! [X2,X0,X3,X1] : (sP7(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79862,f51027])).
fof(f81056,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81018,f427])).
fof(f81018,plain,(
  ( ! [X10,X8,X11,X9] : (sP8('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79860,f375])).
fof(f79860,plain,(
  ( ! [X212,X215,X213,X214] : (sP11(X212,'const.bool.IN_2'(s(sK12,X213),s(fun(sK12,bool),sK15))) | sP8('const.bool.IN_2'(s(sK12,X213),s(fun(sK12,bool),sK13)),X214,X215)) ) | $spl40),
  inference(resolution,[],[f72908,f326])).
fof(f81055,plain,(
  ( ! [X6,X4,X7,X5] : (sP8('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f81017,f427])).
fof(f81017,plain,(
  ( ! [X6,X4,X7,X5] : (sP8('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79860,f379])).
fof(f81016,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79860,f51027])).
fof(f81012,plain,(
  ( ! [X39,X40] : (sP11(X39,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK13)))) ) | $spl448),
  inference(cnf_transformation,[],[f81012_D])).
fof(f81012_D,plain,(
  ( ! [X39,X40] : (sP11(X39,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK13)))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f80996,plain,(
  ( ! [X83,X84,X82] : (sP11(X82,'const.bool.IN_2'(s(sK12,X83),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X83),s(fun(sK12,bool),sK13))) | ~p(s(bool,X84))) ) | ($spl4 | $spl40)),
  inference(resolution,[],[f79859,f61709])).
fof(f81015,plain,(
  ( ! [X78,X81,X79] : (sP11(X78,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK13)))) ) | $spl454),
  inference(cnf_transformation,[],[f81015_D])).
fof(f81015_D,plain,(
  ( ! [X78,X81,X79] : (sP11(X78,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,X81) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK13)))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f80995,plain,(
  ( ! [X80,X78,X81,X79] : (sP11(X78,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK15))) | ~p(s(bool,X80)) | s(bool,X81) != s(bool,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X81)) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79859,f21063])).
fof(f80989,plain,(
  ( ! [X59,X57,X58,X56,X55] : (sP11(X55,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK15))) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK13))) | ~p(s(bool,X58)) | ~p(s(bool,X57)) | p(s(bool,X59))) ) | ($spl40 | $spl152)),
  inference(resolution,[],[f79859,f71583])).
fof(f80988,plain,(
  ( ! [X54,X52,X53,X51] : (sP11(X51,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK15))) | ~p(s(bool,X53)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK13))) | p(s(bool,X54))) ) | ($spl40 | $spl152)),
  inference(resolution,[],[f79859,f40490])).
fof(f80985,plain,(
  ( ! [X39,X41,X42,X40] : (sP11(X39,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK15))) | ~p(s(bool,X41)) | s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl40)),
  inference(resolution,[],[f79859,f21264])).
fof(f81009,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80972,f427])).
fof(f80972,plain,(
  ( ! [X10,X8,X11,X9] : (sP8(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79859,f375])).
fof(f81008,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80971,f427])).
fof(f80971,plain,(
  ( ! [X6,X4,X7,X5] : (sP8(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79859,f379])).
fof(f80970,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79859,f51027])).
fof(f80968,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80949,f427])).
fof(f80949,plain,(
  ( ! [X10,X8,X11,X9] : (sP4('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79847,f375])).
fof(f79847,plain,(
  ( ! [X163,X161,X162,X160] : (sP11(X160,'const.bool.IN_2'(s(sK12,X161),s(fun(sK12,bool),sK15))) | sP4('const.bool.IN_2'(s(sK12,X161),s(fun(sK12,bool),sK13)),X162,X163)) ) | $spl40),
  inference(resolution,[],[f72908,f309])).
fof(f80967,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80948,f427])).
fof(f80948,plain,(
  ( ! [X6,X4,X7,X5] : (sP4('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79847,f379])).
fof(f80947,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79847,f51027])).
fof(f80944,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80925,f427])).
fof(f80925,plain,(
  ( ! [X10,X8,X11,X9] : (sP4(X8,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK13)),X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79846,f375])).
fof(f79846,plain,(
  ( ! [X158,X156,X159,X157] : (sP11(X156,'const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),sK15))) | sP4(X158,'const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),sK13)),X159)) ) | $spl40),
  inference(resolution,[],[f72908,f308])).
fof(f80943,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80924,f427])).
fof(f80924,plain,(
  ( ! [X6,X4,X7,X5] : (sP4(X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13)),X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79846,f379])).
fof(f80923,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK13)),X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79846,f51027])).
fof(f80903,plain,(
  ( ! [X10,X8,X11,X9] : (sP2('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80886,f427])).
fof(f80886,plain,(
  ( ! [X10,X8,X11,X9] : (sP2('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79833,f375])).
fof(f79833,plain,(
  ( ! [X107,X105,X106,X104] : (sP11(X104,'const.bool.IN_2'(s(sK12,X105),s(fun(sK12,bool),sK15))) | sP2('const.bool.IN_2'(s(sK12,X105),s(fun(sK12,bool),sK13)),X106,X107)) ) | $spl40),
  inference(resolution,[],[f72908,f292])).
fof(f80902,plain,(
  ( ! [X6,X4,X7,X5] : (sP2('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80885,f427])).
fof(f80885,plain,(
  ( ! [X6,X4,X7,X5] : (sP2('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79833,f379])).
fof(f80884,plain,(
  ( ! [X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79833,f51027])).
fof(f80881,plain,(
  ( ! [X10,X8,X11,X9] : (sP0('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80862,f427])).
fof(f80862,plain,(
  ( ! [X10,X8,X11,X9] : (sP0('const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13)),X9,X10) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79817,f375])).
fof(f79817,plain,(
  ( ! [X43,X41,X42,X40] : (sP11(X40,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK15))) | sP0('const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)),X42,X43)) ) | $spl40),
  inference(resolution,[],[f72908,f275])).
fof(f80880,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80861,f427])).
fof(f80861,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK13)),X5,X6) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79817,f379])).
fof(f80860,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)),X1,X2) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79817,f51027])).
fof(f80857,plain,(
  ( ! [X10,X8,X11,X9] : (sP0(X8,X9,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13))) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80838,f427])).
fof(f80838,plain,(
  ( ! [X10,X8,X11,X9] : (sP0(X8,X9,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13))) | p(s(bool,'F_0')) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79816,f375])).
fof(f79816,plain,(
  ( ! [X39,X37,X38,X36] : (sP11(X36,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),sK15))) | sP0(X38,X39,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),sK13)))) ) | $spl40),
  inference(resolution,[],[f72908,f273])).
fof(f80856,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4,X5,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK13))) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(subsumption_resolution,[],[f80837,f427])).
fof(f80837,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4,X5,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK13))) | p(s(bool,'F_0')) | p(s(bool,X7)) | p(s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),sK15))))) ) | $spl40),
  inference(resolution,[],[f79816,f379])).
fof(f80836,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl256)),
  inference(resolution,[],[f79816,f51027])).
fof(f75878,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl425),
  inference(cnf_transformation,[],[f75878_D])).
fof(f75878_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl425),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl425])])).
fof(f77180,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16))) ) | $spl434),
  inference(cnf_transformation,[],[f77180_D])).
fof(f77180_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f77170,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl114)),
  inference(resolution,[],[f63196,f27497])).
fof(f27497,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | sP9(X16,X17,X15) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X15)) ) | $spl114),
  inference(subsumption_resolution,[],[f27489,f326])).
fof(f27489,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl114),
  inference(resolution,[],[f27460,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f27460,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X4,X3) | s(bool,'F_0') != s(bool,X4)) ) | $spl114),
  inference(subsumption_resolution,[],[f27459,f328])).
fof(f27459,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3)) ) | $spl114),
  inference(subsumption_resolution,[],[f27391,f330])).
fof(f27391,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | sP7(X5,X4,X3)) ) | $spl114),
  inference(resolution,[],[f26923,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f186])).
fof(f63196,plain,(
  ( ! [X116,X115] : (sP6('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X115,X116)) ) | $spl1),
  inference(resolution,[],[f60499,f334])).
fof(f75881,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl428),
  inference(cnf_transformation,[],[f75881_D])).
fof(f75881_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f77169,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl106)),
  inference(resolution,[],[f63196,f27389])).
fof(f27389,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | sP9(X16,X17,X15) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15)) ) | $spl106),
  inference(subsumption_resolution,[],[f27381,f326])).
fof(f27381,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl106),
  inference(resolution,[],[f27353,f322])).
fof(f27353,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X3,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl106),
  inference(subsumption_resolution,[],[f27352,f328])).
fof(f27352,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X5,X3,X4)) ) | $spl106),
  inference(subsumption_resolution,[],[f27285,f330])).
fof(f27285,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X5,X3,X4)) ) | $spl106),
  inference(resolution,[],[f26912,f317])).
fof(f76933,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8))) ) | $spl432),
  inference(cnf_transformation,[],[f76933_D])).
fof(f76933_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f76923,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl96)),
  inference(resolution,[],[f63195,f26605])).
fof(f26605,plain,(
  ( ! [X17,X15,X16] : (~sP6(X17,X15,X16) | sP9(X16,X15,X17) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15)) ) | $spl96),
  inference(subsumption_resolution,[],[f26596,f326])).
fof(f26596,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X15,X17) | p(s(bool,X16)) | ~sP8(X16,X17,X15) | ~sP6(X17,X15,X16)) ) | $spl96),
  inference(resolution,[],[f26570,f322])).
fof(f26570,plain,(
  ( ! [X4,X5,X3] : (sP7(X3,X5,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl96),
  inference(subsumption_resolution,[],[f26569,f328])).
fof(f26569,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X3,X5,X4)) ) | $spl96),
  inference(subsumption_resolution,[],[f26502,f329])).
fof(f26502,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X3,X5,X4)) ) | $spl96),
  inference(resolution,[],[f26344,f317])).
fof(f63195,plain,(
  ( ! [X114,X113] : (sP6(X113,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X114)) ) | $spl1),
  inference(resolution,[],[f60499,f333])).
fof(f76920,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f63195,f12103])).
fof(f12103,plain,(
  ( ! [X10,X11,X9] : (~sP6(X11,X9,X10) | sP9(X10,X9,X11) | p(s(bool,X10)) | s(bool,'F_0') != s(bool,X9)) ) | $spl8),
  inference(subsumption_resolution,[],[f12099,f326])).
fof(f12099,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP9(X10,X9,X11) | p(s(bool,X10)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl8),
  inference(resolution,[],[f11867,f322])).
fof(f11867,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl8),
  inference(subsumption_resolution,[],[f11866,f329])).
fof(f11866,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl8),
  inference(subsumption_resolution,[],[f11831,f328])).
fof(f11831,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl8),
  inference(resolution,[],[f11501,f317])).
fof(f11501,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl8),
  inference(resolution,[],[f11475,f326])).
fof(f11475,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X4,X3) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f11474,f416])).
fof(f11474,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f11468,f11408])).
fof(f11468,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3) | ~sP6(X4,X3,X5)) ) | $spl8),
  inference(resolution,[],[f11409,f320])).
fof(f11409,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl8),
  inference(subsumption_resolution,[],[f11377,f328])).
fof(f11377,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl8),
  inference(resolution,[],[f5323,f317])).
fof(f75882,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl430),
  inference(cnf_transformation,[],[f75882_D])).
fof(f75882_D,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f75863,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X31))) ) | ($spl1 | $spl94)),
  inference(resolution,[],[f63190,f72212])).
fof(f63190,plain,(
  ( ! [X103,X104] : (sP8(X103,X104,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl1),
  inference(resolution,[],[f60499,f324])).
fof(f75879,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13))) ) | $spl426),
  inference(cnf_transformation,[],[f75879_D])).
fof(f75879_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f75859,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl96)),
  inference(resolution,[],[f63190,f26604])).
fof(f26604,plain,(
  ( ! [X10,X11,X9] : (~sP8(X10,X11,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X9)) ) | $spl96),
  inference(subsumption_resolution,[],[f26594,f334])).
fof(f26594,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl96),
  inference(resolution,[],[f26570,f320])).
fof(f75858,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f63190,f21264])).
fof(f75857,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f63190,f21261])).
fof(f75855,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f63190,f12101])).
fof(f12101,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X5,X3) | sP9(X4,X3,X5) | ~p(s(bool,X5)) | s(bool,'F_0') != s(bool,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f12097,f334])).
fof(f12097,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | sP9(X4,X3,X5) | ~p(s(bool,X5)) | ~sP8(X4,X5,X3) | ~sP6(X5,X3,X4)) ) | $spl8),
  inference(resolution,[],[f11867,f320])).
fof(f74626,plain,(
  ( ! [X352,X351] : (sP9(X352,X351,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X352)) | p(s(bool,X351))) ) | $spl422),
  inference(cnf_transformation,[],[f74626_D])).
fof(f74626_D,plain,(
  ( ! [X352,X351] : (sP9(X352,X351,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X352)) | p(s(bool,X351))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f74625,plain,(
  ( ! [X350,X352,X351] : (~p(s(bool,X350)) | p(s(bool,X351)) | sP9(X352,X351,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X352))) ) | ($spl42 | $spl94)),
  inference(subsumption_resolution,[],[f74441,f416])).
fof(f74441,plain,(
  ( ! [X350,X352,X351] : (s(bool,'T_0') != s(bool,X350) | ~p(s(bool,X350)) | p(s(bool,X351)) | sP9(X352,X351,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X352))) ) | ($spl42 | $spl94)),
  inference(superposition,[],[f61729,f13005])).
fof(f74624,plain,(
  ( ! [X346,X345] : (sP9(X346,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X345) | p(s(bool,X346)) | p(s(bool,X345))) ) | $spl420),
  inference(cnf_transformation,[],[f74624_D])).
fof(f74624_D,plain,(
  ( ! [X346,X345] : (sP9(X346,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X345) | p(s(bool,X346)) | p(s(bool,X345))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f74623,plain,(
  ( ! [X344,X346,X345] : (~p(s(bool,X344)) | p(s(bool,X345)) | sP9(X346,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X345) | p(s(bool,X346))) ) | ($spl42 | $spl94)),
  inference(subsumption_resolution,[],[f74439,f416])).
fof(f74439,plain,(
  ( ! [X344,X346,X345] : (s(bool,'T_0') != s(bool,X344) | ~p(s(bool,X344)) | p(s(bool,X345)) | sP9(X346,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X345) | p(s(bool,X346))) ) | ($spl42 | $spl94)),
  inference(superposition,[],[f61728,f13005])).
fof(f73581,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl411),
  inference(cnf_transformation,[],[f73581_D])).
fof(f73581_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl411),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl411])])).
fof(f73797,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X16))) ) | $spl418),
  inference(cnf_transformation,[],[f73797_D])).
fof(f73797_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X16))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f73787,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl114 | $spl402)),
  inference(resolution,[],[f73190,f27497])).
fof(f73190,plain,(
  ( ! [X116,X115] : (sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X115,X116)) ) | $spl402),
  inference(resolution,[],[f72961,f334])).
fof(f13005,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl42),
  inference(cnf_transformation,[],[f13005_D])).
fof(f13005_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f73786,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl106 | $spl402)),
  inference(resolution,[],[f73190,f27389])).
fof(f73777,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X9) | p(s(bool,X8))) ) | $spl416),
  inference(cnf_transformation,[],[f73777_D])).
fof(f73777_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X9) | p(s(bool,X8))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f73767,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl96 | $spl402)),
  inference(resolution,[],[f73189,f26605])).
fof(f73189,plain,(
  ( ! [X114,X113] : (sP6(X113,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X114)) ) | $spl402),
  inference(resolution,[],[f72961,f333])).
fof(f73764,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl8 | $spl402)),
  inference(resolution,[],[f73189,f12103])).
fof(f73583,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl414),
  inference(cnf_transformation,[],[f73583_D])).
fof(f73583_D,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f73566,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X31))) ) | ($spl94 | $spl402)),
  inference(resolution,[],[f73184,f72212])).
fof(f73184,plain,(
  ( ! [X103,X104] : (sP8(X103,X104,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl402),
  inference(resolution,[],[f72961,f324])).
fof(f73582,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13))) ) | $spl412),
  inference(cnf_transformation,[],[f73582_D])).
fof(f73582_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f73562,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl96 | $spl402)),
  inference(resolution,[],[f73184,f26604])).
fof(f73561,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl8 | $spl402)),
  inference(resolution,[],[f73184,f21264])).
fof(f73560,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl402)),
  inference(resolution,[],[f73184,f21261])).
fof(f73558,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl8 | $spl402)),
  inference(resolution,[],[f73184,f12101])).
fof(f73349,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl407),
  inference(cnf_transformation,[],[f73349_D])).
fof(f73349_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) <=> ~$spl407),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl407])])).
fof(f73351,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl409),
  inference(cnf_transformation,[],[f73351_D])).
fof(f73351_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) <=> ~$spl409),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl409])])).
fof(f73242,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl1),
  inference(resolution,[],[f60500,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X7,X8,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_INTER','thm.pred_set.IN_INTER')).
fof(f60500,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | $spl1),
  inference(resolution,[],[f468,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f73120,plain,(
  $false | ($spl40 | $spl94 | $spl403)),
  inference(subsumption_resolution,[],[f73119,f433])).
fof(f433,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.TRUTH')).
fof(f73119,plain,(
  ~p(s(bool,'T_0')) | ($spl40 | $spl94 | $spl403)),
  inference(forward_demodulation,[],[f73013,f36884])).
fof(f36884,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl94),
  inference(resolution,[],[f14077,f416])).
fof(f73013,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl40 | $spl403)),
  inference(resolution,[],[f72962,f62415])).
fof(f73012,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl404),
  inference(cnf_transformation,[],[f73012_D])).
fof(f73012_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f34991,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl100),
  inference(duplicate_literal_removal,[],[f34931])).
fof(f34931,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl100),
  inference(resolution,[],[f26353,f338])).
fof(f35101,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl112),
  inference(duplicate_literal_removal,[],[f35043])).
fof(f35043,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X12)) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl112),
  inference(resolution,[],[f26922,f338])).
fof(f72961,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl402),
  inference(cnf_transformation,[],[f72961_D])).
fof(f72961_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f73003,plain,(
  ( ! [X68] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X68))) ) | ($spl40 | $spl94 | $spl112)),
  inference(subsumption_resolution,[],[f49193,f62415])).
fof(f49193,plain,(
  ( ! [X68] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X68)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | ($spl94 | $spl112)),
  inference(duplicate_literal_removal,[],[f49162])).
fof(f49162,plain,(
  ( ! [X68] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X68)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X68))) ) | ($spl94 | $spl112)),
  inference(resolution,[],[f36851,f26922])).
fof(f73002,plain,(
  ( ! [X84] : (~p(s(bool,X84)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl94 | $spl100)),
  inference(subsumption_resolution,[],[f50438,f62415])).
fof(f50438,plain,(
  ( ! [X84] : (~p(s(bool,X84)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | ($spl94 | $spl100)),
  inference(duplicate_literal_removal,[],[f50431])).
fof(f50431,plain,(
  ( ! [X84] : (~p(s(bool,X84)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X84))) ) | ($spl94 | $spl100)),
  inference(resolution,[],[f49495,f26353])).
fof(f49495,plain,(
  ( ! [X10,X11] : (~sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X10,X11) | ~p(s(bool,X11)) | p(s(bool,X10))) ) | $spl94),
  inference(resolution,[],[f36854,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X2,X1) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f186])).
fof(f36854,plain,(
  ( ! [X99,X100] : (~sP8('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X99,X100) | p(s(bool,X100)) | ~p(s(bool,X99))) ) | $spl94),
  inference(resolution,[],[f14077,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f189])).
fof(f72962,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl403),
  inference(cnf_transformation,[],[f72962_D])).
fof(f72962_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f72995,plain,(
  ( ! [X70] : (~p(s(bool,X70)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl112)),
  inference(subsumption_resolution,[],[f72994,f397])).
fof(f72994,plain,(
  ( ! [X70] : (p(s(bool,'F_0')) | ~p(s(bool,X70)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl112)),
  inference(forward_demodulation,[],[f52654,f37913])).
fof(f37913,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl21),
  inference(trivial_inequality_removal,[],[f37911])).
fof(f37911,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl21),
  inference(superposition,[],[f6195,f442])).
fof(f52654,plain,(
  ( ! [X70] : (~p(s(bool,X70)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl21 | $spl112)),
  inference(duplicate_literal_removal,[],[f52617])).
fof(f52617,plain,(
  ( ! [X70] : (~p(s(bool,X70)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X70))) ) | ($spl21 | $spl112)),
  inference(resolution,[],[f51960,f26922])).
fof(f51960,plain,(
  ( ! [X2,X3] : (~sP9('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X3,X2) | ~p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl21),
  inference(resolution,[],[f37974,f317])).
fof(f37974,plain,(
  ( ! [X105,X106] : (~sP7(X105,X106,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl21),
  inference(resolution,[],[f37914,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f72993,plain,(
  ( ! [X87] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X87))) ) | ($spl21 | $spl100)),
  inference(subsumption_resolution,[],[f72992,f397])).
fof(f72992,plain,(
  ( ! [X87] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X87))) ) | ($spl21 | $spl100)),
  inference(forward_demodulation,[],[f52649,f37913])).
fof(f52649,plain,(
  ( ! [X87] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X87)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl21 | $spl100)),
  inference(duplicate_literal_removal,[],[f52642])).
fof(f52642,plain,(
  ( ! [X87] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X87)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X87))) ) | ($spl21 | $spl100)),
  inference(resolution,[],[f51960,f26353])).
fof(f72979,plain,(
  ( ! [X68] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X68))) ) | ($spl112 | $spl264)),
  inference(subsumption_resolution,[],[f72978,f403])).
fof(f72978,plain,(
  ( ! [X68] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X68))) ) | ($spl112 | $spl264)),
  inference(forward_demodulation,[],[f59671,f53360])).
fof(f53360,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl264),
  inference(resolution,[],[f53275,f416])).
fof(f59671,plain,(
  ( ! [X68] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X68)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl112 | $spl264)),
  inference(duplicate_literal_removal,[],[f59646])).
fof(f59646,plain,(
  ( ! [X68] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X68)) | p(s(bool,X68)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl112 | $spl264)),
  inference(resolution,[],[f57924,f26922])).
fof(f57924,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl264),
  inference(resolution,[],[f53333,f317])).
fof(f53333,plain,(
  ( ! [X105,X106] : (~sP7('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X105,X106) | p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl264),
  inference(resolution,[],[f53275,f327])).
fof(f72977,plain,(
  ( ! [X69] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X69))) ) | ($spl100 | $spl264)),
  inference(subsumption_resolution,[],[f72976,f403])).
fof(f72976,plain,(
  ( ! [X69] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X69))) ) | ($spl100 | $spl264)),
  inference(forward_demodulation,[],[f59853,f53360])).
fof(f59853,plain,(
  ( ! [X69] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X69)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl100 | $spl264)),
  inference(duplicate_literal_removal,[],[f59823])).
fof(f59823,plain,(
  ( ! [X69] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X69)) | p(s(bool,X69)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | ($spl100 | $spl264)),
  inference(resolution,[],[f57975,f26353])).
fof(f57975,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl264),
  inference(resolution,[],[f53334,f317])).
fof(f53334,plain,(
  ( ! [X107,X108] : (~sP7(X107,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X108) | p(s(bool,X108)) | ~p(s(bool,X107))) ) | $spl264),
  inference(resolution,[],[f53275,f327])).
fof(f68444,plain,(
  ( ! [X60] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X60))) ) | ($spl40 | $spl112)),
  inference(subsumption_resolution,[],[f68387,f436])).
fof(f68387,plain,(
  ( ! [X60] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X60)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl112)),
  inference(duplicate_literal_removal,[],[f68366])).
fof(f68366,plain,(
  ( ! [X60] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X60)) | p(s(bool,X60)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl112)),
  inference(resolution,[],[f67930,f26922])).
fof(f67930,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl40),
  inference(resolution,[],[f62493,f317])).
fof(f62493,plain,(
  ( ! [X208,X207] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X207,X208) | p(s(bool,X208)) | ~p(s(bool,X207))) ) | $spl40),
  inference(subsumption_resolution,[],[f62275,f403])).
fof(f62275,plain,(
  ( ! [X208,X207] : (~p(s(bool,'T_0')) | ~p(s(bool,X207)) | p(s(bool,X208)) | ~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X207,X208)) ) | $spl40),
  inference(superposition,[],[f327,f13003])).
fof(f68647,plain,(
  ( ! [X62] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X62))) ) | ($spl40 | $spl100)),
  inference(subsumption_resolution,[],[f68596,f436])).
fof(f68596,plain,(
  ( ! [X62] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X62)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl100)),
  inference(duplicate_literal_removal,[],[f68570])).
fof(f68570,plain,(
  ( ! [X62] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X62)) | p(s(bool,X62)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl40 | $spl100)),
  inference(resolution,[],[f67974,f26353])).
fof(f67974,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl40),
  inference(resolution,[],[f62494,f317])).
fof(f62494,plain,(
  ( ! [X210,X209] : (~sP7(X209,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X210) | p(s(bool,X210)) | ~p(s(bool,X209))) ) | $spl40),
  inference(subsumption_resolution,[],[f62276,f403])).
fof(f62276,plain,(
  ( ! [X210,X209] : (~p(s(bool,'T_0')) | ~p(s(bool,X209)) | p(s(bool,X210)) | ~sP7(X209,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X210)) ) | $spl40),
  inference(superposition,[],[f327,f13003])).
fof(f72960,plain,(
  ( ! [X94,X93] : (sP9(X94,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X93) | p(s(bool,X94)) | p(s(bool,X93))) ) | $spl400),
  inference(cnf_transformation,[],[f72960_D])).
fof(f72960_D,plain,(
  ( ! [X94,X93] : (sP9(X94,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X93) | p(s(bool,X94)) | p(s(bool,X93))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f71715,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,X92)) | p(s(bool,X93)) | sP9(X94,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X93) | p(s(bool,X94))) ) | ($spl40 | $spl94)),
  inference(subsumption_resolution,[],[f71668,f416])).
fof(f71668,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') != s(bool,X92) | ~p(s(bool,X92)) | p(s(bool,X93)) | sP9(X94,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X93) | p(s(bool,X94))) ) | ($spl40 | $spl94)),
  inference(superposition,[],[f61728,f13003])).
fof(f72959,plain,(
  ( ! [X94,X93] : (sP9(X94,X93,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X94)) | p(s(bool,X93))) ) | $spl398),
  inference(cnf_transformation,[],[f72959_D])).
fof(f72959_D,plain,(
  ( ! [X94,X93] : (sP9(X94,X93,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X94)) | p(s(bool,X93))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f71806,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,X92)) | p(s(bool,X93)) | sP9(X94,X93,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X94))) ) | ($spl40 | $spl94)),
  inference(subsumption_resolution,[],[f71759,f416])).
fof(f71759,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') != s(bool,X92) | ~p(s(bool,X92)) | p(s(bool,X93)) | sP9(X94,X93,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X94))) ) | ($spl40 | $spl94)),
  inference(superposition,[],[f61729,f13003])).
fof(f72950,plain,(
  $false | ($spl40 | $spl43 | $spl94)),
  inference(subsumption_resolution,[],[f72949,f433])).
fof(f72949,plain,(
  ~p(s(bool,'T_0')) | ($spl40 | $spl43 | $spl94)),
  inference(forward_demodulation,[],[f72948,f36884])).
fof(f72948,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl40 | $spl43)),
  inference(subsumption_resolution,[],[f72928,f425])).
fof(f425,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) => ~p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_F')).
fof(f72928,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl40 | $spl43)),
  inference(superposition,[],[f62415,f15564])).
fof(f15564,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl43),
  inference(trivial_inequality_removal,[],[f15562])).
fof(f15562,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl43),
  inference(superposition,[],[f13006,f442])).
fof(f72015,plain,(
  ( ! [X99,X100] : (~p(s(bool,X99)) | p(s(bool,X100)) | sP9(X100,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X99)) ) | $spl396),
  inference(cnf_transformation,[],[f72015_D])).
fof(f72015_D,plain,(
  ( ! [X99,X100] : (~p(s(bool,X99)) | p(s(bool,X100)) | sP9(X100,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X99)) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f72014,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X98) | ~p(s(bool,X99)) | sP9(X100,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X99) | p(s(bool,X100))) ) | ($spl1 | $spl94)),
  inference(subsumption_resolution,[],[f71968,f411])).
fof(f71968,plain,(
  ( ! [X99,X100,X98] : (s(bool,'F_0') != s(bool,X98) | ~p(s(bool,X99)) | p(s(bool,X98)) | sP9(X100,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X99) | p(s(bool,X100))) ) | ($spl1 | $spl94)),
  inference(superposition,[],[f61752,f60578])).
fof(f60578,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | $spl1),
  inference(resolution,[],[f468,f412])).
fof(f4225,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | $spl13),
  inference(cnf_transformation,[],[f4225_D])).
fof(f4225_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f62116,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X13) | p(s(bool,X12))) ) | $spl394),
  inference(cnf_transformation,[],[f62116_D])).
fof(f62116_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X13) | p(s(bool,X12))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f62106,plain,(
  ( ! [X17,X16] : (sP9(X16,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X17) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl246)),
  inference(resolution,[],[f750,f45868])).
fof(f45868,plain,(
  ( ! [X6,X8,X7] : (~sP7(X6,X8,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl246),
  inference(subsumption_resolution,[],[f45848,f326])).
fof(f45848,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl246),
  inference(resolution,[],[f45831,f322])).
fof(f45831,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,'T_0') != s(bool,X1)) ) | $spl246),
  inference(subsumption_resolution,[],[f45830,f332])).
fof(f45830,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | sP6(X2,X1,X0)) ) | $spl246),
  inference(subsumption_resolution,[],[f45759,f334])).
fof(f45759,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | sP6(X2,X1,X0)) ) | $spl246),
  inference(resolution,[],[f45677,f316])).
fof(f750,plain,(
  ( ! [X107,X106] : (sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X106,X107)) ) | $spl1),
  inference(resolution,[],[f468,f330])).
fof(f61839,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | $spl388),
  inference(cnf_transformation,[],[f61839_D])).
fof(f61839_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f62104,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X13) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl116)),
  inference(resolution,[],[f750,f29496])).
fof(f29496,plain,(
  ( ! [X6,X8,X7] : (~sP7(X6,X8,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) ) | $spl116),
  inference(subsumption_resolution,[],[f29475,f326])).
fof(f29475,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl116),
  inference(resolution,[],[f29460,f322])).
fof(f29460,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl116),
  inference(subsumption_resolution,[],[f29459,f332])).
fof(f29459,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X2,X0,X1)) ) | $spl116),
  inference(subsumption_resolution,[],[f29384,f334])).
fof(f29384,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X2,X0,X1)) ) | $spl116),
  inference(resolution,[],[f28512,f316])).
fof(f62115,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | $spl392),
  inference(cnf_transformation,[],[f62115_D])).
fof(f62115_D,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f62103,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f750,f13485])).
fof(f61941,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X8))) ) | $spl390),
  inference(cnf_transformation,[],[f61941_D])).
fof(f61941_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X8))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f61930,plain,(
  ( ! [X14,X15] : (sP9(X14,X15,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl130)),
  inference(resolution,[],[f749,f30672])).
fof(f30672,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X6,X7) | sP9(X7,X8,X6) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) ) | $spl130),
  inference(subsumption_resolution,[],[f30649,f326])).
fof(f30649,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X7)) | ~sP8(X7,X6,X8) | ~sP7(X8,X6,X7)) ) | $spl130),
  inference(resolution,[],[f30636,f322])).
fof(f30636,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X2,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl130),
  inference(subsumption_resolution,[],[f30635,f332])).
fof(f30635,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X0,X2,X1)) ) | $spl130),
  inference(subsumption_resolution,[],[f30561,f333])).
fof(f30561,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X0,X2,X1)) ) | $spl130),
  inference(resolution,[],[f30056,f316])).
fof(f749,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X105)) ) | $spl1),
  inference(resolution,[],[f468,f329])).
fof(f61927,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f749,f12095])).
fof(f12095,plain,(
  ( ! [X10,X11,X9] : (~sP7(X11,X9,X10) | sP9(X10,X11,X9) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl8),
  inference(subsumption_resolution,[],[f12089,f326])).
fof(f12089,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP9(X10,X11,X9) | p(s(bool,X10)) | ~sP8(X10,X9,X11) | ~sP7(X11,X9,X10)) ) | $spl8),
  inference(resolution,[],[f11558,f322])).
fof(f61828,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | ~p(s(bool,X33))) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f746,f61709])).
fof(f746,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X99)) ) | $spl1),
  inference(resolution,[],[f468,f325])).
fof(f11886,plain,(
  ( ! [X317] : (s(bool,'T_0') != s(bool,X317) | s(bool,X317) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | $spl26),
  inference(cnf_transformation,[],[f11886_D])).
fof(f11886_D,plain,(
  ( ! [X317] : (s(bool,'T_0') != s(bool,X317) | s(bool,X317) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f61827,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | s(bool,X32) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | s(bool,'T_0') != s(bool,X32)) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f746,f21063])).
fof(f4226,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | $spl14),
  inference(cnf_transformation,[],[f4226_D])).
fof(f4226_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f61826,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,X29) | sP9(X30,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | s(bool,X29) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f746,f20975])).
fof(f20975,plain,(
  ( ! [X10,X11,X9] : (~sP8(X11,X10,'F_0') | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | s(bool,X9) != s(bool,X10)) ) | $spl8),
  inference(subsumption_resolution,[],[f20974,f427])).
fof(f20974,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | p(s(bool,'F_0')) | ~sP8(X11,X10,'F_0')) ) | $spl8),
  inference(subsumption_resolution,[],[f20964,f330])).
fof(f20964,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | p(s(bool,'F_0')) | ~sP8(X11,X10,'F_0') | ~sP7('F_0',X10,X11)) ) | $spl8),
  inference(resolution,[],[f20950,f321])).
fof(f61822,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X25))) ) | ($spl1 | $spl152)),
  inference(resolution,[],[f746,f40490])).
fof(f61837,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X11))) ) | $spl386),
  inference(cnf_transformation,[],[f61837_D])).
fof(f61837_D,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X11))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f61821,plain,(
  ( ! [X23,X22] : (sP9(X22,X23,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X23)) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl130)),
  inference(resolution,[],[f746,f30671])).
fof(f30671,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | sP9(X4,X5,X3) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3)) ) | $spl130),
  inference(subsumption_resolution,[],[f30648,f330])).
fof(f30648,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | sP9(X4,X5,X3) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl130),
  inference(resolution,[],[f30636,f321])).
fof(f61819,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f746,f21264])).
fof(f61815,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | p(s(bool,X11)) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f746,f12094])).
fof(f12094,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X6,X8) | sP9(X7,X8,X6) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(subsumption_resolution,[],[f12088,f330])).
fof(f12088,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X8)) | ~sP8(X7,X6,X8) | ~sP7(X8,X6,X7)) ) | $spl8),
  inference(resolution,[],[f11558,f321])).
fof(f472,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl5),
  inference(cnf_transformation,[],[f472_D])).
fof(f472_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f61167,plain,(
  ( ! [X46] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X46))) ) | ($spl21 | $spl62)),
  inference(resolution,[],[f13829,f38012])).
fof(f38012,plain,(
  ( ! [X6,X7] : (~sP9(X7,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl21),
  inference(superposition,[],[f37914,f337])).
fof(f471,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl4),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f61166,plain,(
  ( ! [X45] : (p(s(bool,X45)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl62)),
  inference(resolution,[],[f13829,f52019])).
fof(f52019,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl21),
  inference(resolution,[],[f37978,f316])).
fof(f37978,plain,(
  ( ! [X114,X113] : (~sP6(X113,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl21),
  inference(resolution,[],[f37914,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f194])).
fof(f61165,plain,(
  ( ! [X44] : (p(s(bool,X44)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl264)),
  inference(resolution,[],[f13829,f53328])).
fof(f53328,plain,(
  ( ! [X95,X96] : (~sP9(X95,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X96) | p(s(bool,X96)) | ~p(s(bool,X95))) ) | $spl264),
  inference(resolution,[],[f53275,f319])).
fof(f61164,plain,(
  ( ! [X43] : (~p(s(bool,X43)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl264)),
  inference(resolution,[],[f13829,f57924])).
fof(f61163,plain,(
  ( ! [X42] : (p(s(bool,X42)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl94)),
  inference(resolution,[],[f13829,f36852])).
fof(f36852,plain,(
  ( ! [X95,X96] : (~sP9(X95,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X96) | p(s(bool,X96)) | ~p(s(bool,X95))) ) | $spl94),
  inference(resolution,[],[f14077,f319])).
fof(f61162,plain,(
  ( ! [X41] : (~p(s(bool,X41)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl94)),
  inference(resolution,[],[f13829,f49551])).
fof(f49551,plain,(
  ( ! [X2,X3] : (~sP9(X2,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X3) | ~p(s(bool,X3)) | p(s(bool,X2))) ) | $spl94),
  inference(resolution,[],[f36857,f317])).
fof(f36857,plain,(
  ( ! [X105,X106] : (~sP7('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X105,X106) | p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl94),
  inference(resolution,[],[f14077,f327])).
fof(f61161,plain,(
  ( ! [X40] : (p(s(bool,X40)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl43 | $spl62)),
  inference(resolution,[],[f13829,f34841])).
fof(f34841,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl43),
  inference(resolution,[],[f16539,f316])).
fof(f16539,plain,(
  ( ! [X225,X224] : (~sP6(X224,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X225) | p(s(bool,X225)) | p(s(bool,X224))) ) | $spl43),
  inference(subsumption_resolution,[],[f16335,f427])).
fof(f16335,plain,(
  ( ! [X225,X224] : (p(s(bool,'F_0')) | p(s(bool,X224)) | p(s(bool,X225)) | ~sP6(X224,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X225)) ) | $spl43),
  inference(superposition,[],[f331,f15564])).
fof(f11902,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37)) ) | $spl32),
  inference(cnf_transformation,[],[f11902_D])).
fof(f11902_D,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f61160,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X38) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,X38) != s(bool,X39)) ) | ($spl8 | $spl62)),
  inference(resolution,[],[f13829,f12166])).
fof(f61159,plain,(
  ( ! [X37] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X37))) ) | ($spl2 | $spl62)),
  inference(resolution,[],[f13829,f1663])).
fof(f1663,plain,(
  ( ! [X8,X9] : (~sP9(X9,'F_0',X8) | p(s(bool,X9)) | p(s(bool,X8))) ) | $spl2),
  inference(resolution,[],[f1658,f316])).
fof(f1658,plain,(
  ( ! [X2,X3] : (~sP6(X2,'F_0',X3) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl2),
  inference(subsumption_resolution,[],[f1652,f412])).
fof(f1652,plain,(
  ( ! [X2,X3] : (~sP6(X2,'F_0',X3) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f1640,f338])).
fof(f1640,plain,(
  ( ! [X2,X3] : (sP9(X2,'F_0',X3) | ~sP6(X3,'F_0',X2) | p(s(bool,X3))) ) | $spl2),
  inference(resolution,[],[f1637,f325])).
fof(f1637,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | sP9(X4,'F_0',X5) | ~sP6(X5,'F_0',X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f1635,f427])).
fof(f1635,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl2),
  inference(resolution,[],[f1632,f321])).
fof(f1632,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7)) ) | $spl2),
  inference(subsumption_resolution,[],[f1631,f329])).
fof(f1631,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP7('F_0',X6,X7)) ) | $spl2),
  inference(subsumption_resolution,[],[f1627,f328])).
fof(f1627,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl2),
  inference(resolution,[],[f1622,f317])).
fof(f1622,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f1617,f326])).
fof(f1617,plain,(
  ( ! [X2,X3] : (~sP8(X3,X2,'F_0') | ~p(s(bool,X2)) | sP9(X3,'F_0',X2)) ) | $spl2),
  inference(subsumption_resolution,[],[f1616,f416])).
fof(f1616,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl2),
  inference(subsumption_resolution,[],[f1613,f1236])).
fof(f1236,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f1231,f332])).
fof(f1231,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl2),
  inference(resolution,[],[f980,f316])).
fof(f980,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl2),
  inference(superposition,[],[f973,f335])).
fof(f973,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(subsumption_resolution,[],[f964,f469])).
fof(f964,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | $spl2),
  inference(superposition,[],[f550,f416])).
fof(f550,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl2),
  inference(resolution,[],[f469,f411])).
fof(f1613,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl2),
  inference(resolution,[],[f1237,f320])).
fof(f1237,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f1232,f328])).
fof(f1232,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl2),
  inference(resolution,[],[f980,f317])).
fof(f61158,plain,(
  ( ! [X36] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X36))) ) | ($spl4 | $spl62)),
  inference(resolution,[],[f13829,f5541])).
fof(f5541,plain,(
  ( ! [X6,X7] : (~sP9(X7,'F_0',X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f5534,f318])).
fof(f5534,plain,(
  ( ! [X10,X11] : (~sP8(X10,X11,'F_0') | ~p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl4),
  inference(subsumption_resolution,[],[f5533,f334])).
fof(f5533,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP8(X10,X11,'F_0') | ~sP6(X11,'F_0',X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f5532,f328])).
fof(f5532,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP8(X10,X11,'F_0') | ~sP7('F_0',X11,X10) | ~sP6(X11,'F_0',X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f5522,f416])).
fof(f5522,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~sP8(X10,X11,'F_0') | ~sP7('F_0',X11,X10) | ~sP6(X11,'F_0',X10)) ) | $spl4),
  inference(resolution,[],[f4665,f320])).
fof(f4665,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl4),
  inference(superposition,[],[f1821,f337])).
fof(f61157,plain,(
  ( ! [X35] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X35))) ) | ($spl62 | $spl138)),
  inference(resolution,[],[f13829,f31178])).
fof(f61156,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl62)),
  inference(resolution,[],[f13829,f37968])).
fof(f37968,plain,(
  ( ! [X94,X93] : (~sP9(X93,X94,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X94)) | ~p(s(bool,X93))) ) | $spl21),
  inference(resolution,[],[f37914,f319])).
fof(f61155,plain,(
  ( ! [X33] : (p(s(bool,X33)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl21 | $spl62)),
  inference(resolution,[],[f13829,f52057])).
fof(f52057,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl21),
  inference(resolution,[],[f37979,f316])).
fof(f37979,plain,(
  ( ! [X116,X115] : (~sP6('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X115,X116) | p(s(bool,X116)) | p(s(bool,X115))) ) | $spl21),
  inference(resolution,[],[f37914,f331])).
fof(f61154,plain,(
  ( ! [X32] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X32))) ) | ($spl62 | $spl264)),
  inference(resolution,[],[f13829,f53373])).
fof(f53373,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X5)) | p(s(bool,X4))) ) | $spl264),
  inference(superposition,[],[f53275,f337])).
fof(f61153,plain,(
  ( ! [X31] : (~p(s(bool,X31)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl264)),
  inference(resolution,[],[f13829,f57975])).
fof(f61152,plain,(
  ( ! [X30] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X30))) ) | ($spl62 | $spl94)),
  inference(resolution,[],[f13829,f36895])).
fof(f36895,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X5)) | p(s(bool,X4))) ) | $spl94),
  inference(superposition,[],[f14077,f337])).
fof(f61151,plain,(
  ( ! [X29] : (~p(s(bool,X29)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl94)),
  inference(resolution,[],[f13829,f49600])).
fof(f49600,plain,(
  ( ! [X2,X3] : (~sP9(X2,X3,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X3)) | p(s(bool,X2))) ) | $spl94),
  inference(resolution,[],[f36858,f317])).
fof(f36858,plain,(
  ( ! [X107,X108] : (~sP7(X107,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108) | p(s(bool,X108)) | ~p(s(bool,X107))) ) | $spl94),
  inference(resolution,[],[f14077,f327])).
fof(f61150,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl43 | $spl62)),
  inference(resolution,[],[f13829,f16531])).
fof(f16531,plain,(
  ( ! [X187,X186] : (~sP9(X186,X187,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X187)) | ~p(s(bool,X186))) ) | $spl43),
  inference(subsumption_resolution,[],[f16316,f397])).
fof(f16316,plain,(
  ( ! [X187,X186] : (p(s(bool,'F_0')) | ~p(s(bool,X186)) | ~p(s(bool,X187)) | ~sP9(X186,X187,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl43),
  inference(superposition,[],[f319,f15564])).
fof(f61149,plain,(
  ( ! [X27] : (p(s(bool,X27)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl43 | $spl62)),
  inference(resolution,[],[f13829,f34884])).
fof(f34884,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl43),
  inference(resolution,[],[f16540,f316])).
fof(f16540,plain,(
  ( ! [X227,X226] : (~sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X226,X227) | p(s(bool,X227)) | p(s(bool,X226))) ) | $spl43),
  inference(subsumption_resolution,[],[f16336,f427])).
fof(f16336,plain,(
  ( ! [X227,X226] : (p(s(bool,'F_0')) | p(s(bool,X226)) | p(s(bool,X227)) | ~sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X226,X227)) ) | $spl43),
  inference(superposition,[],[f331,f15564])).
fof(f61148,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,X25) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,X25) != s(bool,X26)) ) | ($spl8 | $spl62)),
  inference(resolution,[],[f13829,f12165])).
fof(f12165,plain,(
  ( ! [X14,X12,X13] : (~sP9(X14,X12,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13)) ) | $spl8),
  inference(superposition,[],[f1003,f337])).
fof(f61147,plain,(
  ( ! [X24] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X24))) ) | ($spl62 | $spl148)),
  inference(resolution,[],[f13829,f31533])).
fof(f61146,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') = s(bool,X23)) ) | ($spl62 | $spl152)),
  inference(resolution,[],[f13829,f32078])).
fof(f61145,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | ($spl62 | $spl144)),
  inference(resolution,[],[f13829,f31526])).
fof(f61144,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,X18)) | s(bool,'F_0') = s(bool,X19)) ) | ($spl62 | $spl142)),
  inference(resolution,[],[f13829,f31454])).
fof(f31454,plain,(
  ( ! [X50,X48,X49] : (~sP9(X49,X48,X50) | ~p(s(bool,X49)) | p(s(bool,X50)) | s(bool,'F_0') = s(bool,X48)) ) | $spl142),
  inference(trivial_inequality_removal,[],[f31432])).
fof(f31432,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,'T_0') | ~sP9(X49,X48,X50) | ~p(s(bool,X49)) | p(s(bool,X50)) | s(bool,'F_0') = s(bool,X48)) ) | $spl142),
  inference(superposition,[],[f31183,f442])).
fof(f61143,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') != s(bool,X17)) ) | ($spl8 | $spl62)),
  inference(resolution,[],[f13829,f21285])).
fof(f21285,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,X9,X10) | ~p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f21261,f318])).
fof(f61142,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') != s(bool,X15)) ) | ($spl8 | $spl62)),
  inference(resolution,[],[f13829,f12180])).
fof(f12180,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,X9,X10) | s(bool,'T_0') != s(bool,X9) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,X10)) ) | $spl8),
  inference(superposition,[],[f1003,f337])).
fof(f61141,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,'F_0') != s(bool,X13)) ) | ($spl8 | $spl62)),
  inference(resolution,[],[f13829,f12179])).
fof(f58912,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4)) ) | $spl290),
  inference(cnf_transformation,[],[f58912_D])).
fof(f58912_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f61140,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | s(bool,X10) != s(bool,X11)) ) | $spl62),
  inference(resolution,[],[f13829,f338])).
fof(f58793,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5)) ) | $spl288),
  inference(cnf_transformation,[],[f58793_D])).
fof(f58793_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f61139,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | $spl62),
  inference(resolution,[],[f13829,f337])).
fof(f61138,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | $spl62),
  inference(resolution,[],[f13829,f319])).
fof(f1772,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl11),
  inference(cnf_transformation,[],[f1772_D])).
fof(f1772_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f61134,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X16))) ) | $spl384),
  inference(cnf_transformation,[],[f61134_D])).
fof(f61134_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X16))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f61124,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl114)),
  inference(resolution,[],[f638,f27497])).
fof(f638,plain,(
  ( ! [X114,X113] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X113,X114)) ) | $spl4),
  inference(resolution,[],[f471,f334])).
fof(f13003,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl40),
  inference(cnf_transformation,[],[f13003_D])).
fof(f13003_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f61123,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl106)),
  inference(resolution,[],[f638,f27389])).
fof(f13839,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1) | p(s(bool,X0))) ) | $spl66),
  inference(cnf_transformation,[],[f13839_D])).
fof(f13839_D,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X1) | p(s(bool,X0))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f61104,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl96)),
  inference(resolution,[],[f637,f26605])).
fof(f637,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X112)) ) | $spl4),
  inference(resolution,[],[f471,f333])).
fof(f61101,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f637,f12103])).
fof(f60927,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13))) ) | $spl348),
  inference(cnf_transformation,[],[f60927_D])).
fof(f60927_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f60912,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl96)),
  inference(resolution,[],[f632,f26604])).
fof(f632,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl4),
  inference(resolution,[],[f471,f324])).
fof(f60911,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f632,f21264])).
fof(f60910,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f632,f21261])).
fof(f60908,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f632,f12101])).
fof(f11888,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f11888_D])).
fof(f11888_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f60496,plain,(
  ( ! [X17,X16] : (sP9(X16,'F_0',X17) | p(s(bool,X16))) ) | $spl344),
  inference(cnf_transformation,[],[f60496_D])).
fof(f60496_D,plain,(
  ( ! [X17,X16] : (sP9(X16,'F_0',X17) | p(s(bool,X16))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f60486,plain,(
  ( ! [X17,X16] : (sP9(X16,'F_0',X17) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl4 | $spl106 | $spl246)),
  inference(resolution,[],[f60475,f45868])).
fof(f60475,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7)) ) | ($spl4 | $spl106)),
  inference(subsumption_resolution,[],[f5365,f27353])).
fof(f5365,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f5354,f328])).
fof(f5354,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl4),
  inference(resolution,[],[f4663,f317])).
fof(f4663,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl4),
  inference(superposition,[],[f1821,f335])).
fof(f36521,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl179),
  inference(cnf_transformation,[],[f36521_D])).
fof(f36521_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f39241,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl108)),
  inference(duplicate_literal_removal,[],[f39228])).
fof(f39228,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl108)),
  inference(resolution,[],[f36585,f338])).
fof(f36585,plain,(
  ( ! [X22] : (sP9(X22,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),'F_0') | p(s(bool,X22))) ) | ($spl2 | $spl108)),
  inference(resolution,[],[f529,f27096])).
fof(f27096,plain,(
  ( ! [X10,X11] : (~sP6('F_0',X11,X10) | p(s(bool,X10)) | sP9(X10,X11,'F_0')) ) | $spl108),
  inference(subsumption_resolution,[],[f27088,f326])).
fof(f27088,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'F_0') | p(s(bool,X10)) | ~sP8(X10,'F_0',X11) | ~sP6('F_0',X11,X10)) ) | $spl108),
  inference(resolution,[],[f27061,f322])).
fof(f27061,plain,(
  ( ! [X8,X7] : (sP7(X8,'F_0',X7)) ) | $spl108),
  inference(subsumption_resolution,[],[f27060,f328])).
fof(f27060,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | sP7(X8,'F_0',X7)) ) | $spl108),
  inference(subsumption_resolution,[],[f27011,f330])).
fof(f27011,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | ~p(s(bool,X8)) | sP7(X8,'F_0',X7)) ) | $spl108),
  inference(resolution,[],[f26914,f317])).
fof(f529,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X112)) ) | $spl2),
  inference(resolution,[],[f469,f333])).
fof(f36522,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13))) ) | $spl180),
  inference(cnf_transformation,[],[f36522_D])).
fof(f36522_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f41499,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9(X9,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X8) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f36516,f11867])).
fof(f36592,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8))) ) | $spl214),
  inference(cnf_transformation,[],[f36592_D])).
fof(f36592_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f42127,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | sP9(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X9) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f36591,f11867])).
fof(f36591,plain,(
  ( ! [X4,X5] : (~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5,X4) | p(s(bool,X4)) | sP9(X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5)) ) | $spl2),
  inference(subsumption_resolution,[],[f36575,f326])).
fof(f36575,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5) | p(s(bool,X4)) | ~sP8(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~sP7('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5,X4)) ) | $spl2),
  inference(resolution,[],[f529,f322])).
fof(f60440,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl338),
  inference(cnf_transformation,[],[f60440_D])).
fof(f60440_D,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,X8) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f42192,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f42150])).
fof(f42150,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X9))) ) | $spl2),
  inference(resolution,[],[f41496,f337])).
fof(f42195,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f42191,f416])).
fof(f42191,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X12) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f42152])).
fof(f42152,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f41496,f12179])).
fof(f42189,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f42154])).
fof(f42154,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f41496,f21285])).
fof(f42389,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f42356])).
fof(f42356,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(resolution,[],[f42125,f338])).
fof(f42125,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f36591,f329])).
fof(f42531,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f42494])).
fof(f42494,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(resolution,[],[f42475,f338])).
fof(f42475,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f36611,f326])).
fof(f36611,plain,(
  ( ! [X2,X3] : (~sP8(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X3) | p(s(bool,X3)) | sP9(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(subsumption_resolution,[],[f36595,f330])).
fof(f36595,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X3)) | ~sP8(X2,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X3) | ~sP7(X3,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X2)) ) | $spl2),
  inference(resolution,[],[f530,f321])).
fof(f530,plain,(
  ( ! [X114,X113] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X113,X114)) ) | $spl2),
  inference(resolution,[],[f469,f334])).
fof(f36614,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X16))) ) | $spl220),
  inference(cnf_transformation,[],[f36614_D])).
fof(f36614_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X16))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f42712,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | sP9(X12,X13,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl114)),
  inference(resolution,[],[f36612,f27460])).
fof(f36612,plain,(
  ( ! [X4,X5] : (~sP7(X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X4) | p(s(bool,X4)) | sP9(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(subsumption_resolution,[],[f36596,f326])).
fof(f36596,plain,(
  ( ! [X4,X5] : (sP9(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X4)) | ~sP8(X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X5) | ~sP7(X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X4)) ) | $spl2),
  inference(resolution,[],[f530,f322])).
fof(f47145,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X47)) | sP9(X47,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X48)) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f26570,f36591])).
fof(f47146,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X49)) | sP9(X50,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X49)) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f26570,f36516])).
fof(f47325,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X15) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f26604,f524])).
fof(f47359,plain,(
  ( ! [X47,X46] : (sP9(X46,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X47) | p(s(bool,X46)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f26605,f529])).
fof(f47698,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X39)) | sP9(X39,X40,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl106)),
  inference(resolution,[],[f27353,f36612])).
fof(f47936,plain,(
  ( ! [X59,X60] : (sP9(X59,X60,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X59)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl106)),
  inference(resolution,[],[f27389,f530])).
fof(f13829,plain,(
  ( ! [X10,X11] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X10,X11)) ) | $spl62),
  inference(cnf_transformation,[],[f13829_D])).
fof(f13829_D,plain,(
  ( ! [X10,X11] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X10,X11)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f59511,plain,(
  ( ! [X4,X5] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X5,X4) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | $spl88),
  inference(subsumption_resolution,[],[f59490,f332])).
fof(f59490,plain,(
  ( ! [X4,X5] : (~sP6(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X5,X4) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | $spl88),
  inference(resolution,[],[f14053,f328])).
fof(f12210,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24)) ) | $spl8),
  inference(subsumption_resolution,[],[f12189,f1003])).
fof(f12189,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl8),
  inference(superposition,[],[f1003,f442])).
fof(f12203,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f12191])).
fof(f12191,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,X29) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(superposition,[],[f1003,f442])).
fof(f12202,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f12192])).
fof(f12192,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,X30) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl8),
  inference(superposition,[],[f1003,f414])).
fof(f414,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.EQ_CLAUSES|split|2')).
fof(f12201,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f12193])).
fof(f12193,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,X31) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl8),
  inference(superposition,[],[f1003,f412])).
fof(f28517,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X72))) ) | $spl122),
  inference(cnf_transformation,[],[f28517_D])).
fof(f28517_D,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X72))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f17417,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37))) ) | $spl24),
  inference(subsumption_resolution,[],[f17388,f11865])).
fof(f17388,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | s(bool,'F_0') = s(bool,X36)) ) | $spl24),
  inference(superposition,[],[f11865,f442])).
fof(f26915,plain,(
  ( ! [X80,X81] : (s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X80)) ) | $spl110),
  inference(cnf_transformation,[],[f26915_D])).
fof(f26915_D,plain,(
  ( ! [X80,X81] : (s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X80)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f60349,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl8),
  inference(subsumption_resolution,[],[f20154,f12160])).
fof(f20154,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl8),
  inference(superposition,[],[f12160,f442])).
fof(f60348,plain,(
  ( ! [X58] : (s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X58)) ) | $spl324),
  inference(cnf_transformation,[],[f60348_D])).
fof(f60348_D,plain,(
  ( ! [X58] : (s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X58)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f20167,plain,(
  ( ! [X58] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl8 | $spl43)),
  inference(superposition,[],[f12160,f15564])).
fof(f60346,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | sP9(X7,X8,'F_0')) ) | $spl322),
  inference(cnf_transformation,[],[f60346_D])).
fof(f60346_D,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | sP9(X7,X8,'F_0')) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f27084,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'F_0') | p(s(bool,X2)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl8 | $spl108)),
  inference(resolution,[],[f27061,f12095])).
fof(f27102,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | sP9(X7,X8,'F_0') | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl8 | $spl108)),
  inference(resolution,[],[f27096,f11558])).
fof(f60344,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | sP9(X8,'T_0',X9)) ) | $spl320),
  inference(cnf_transformation,[],[f60344_D])).
fof(f60344_D,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | sP9(X8,'T_0',X9)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f28894,plain,(
  ( ! [X8,X9] : (sP9(X8,'T_0',X9) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl8 | $spl118)),
  inference(resolution,[],[f28877,f12103])).
fof(f28877,plain,(
  ( ! [X0,X1] : (sP6(X1,'T_0',X0)) ) | $spl118),
  inference(subsumption_resolution,[],[f28876,f332])).
fof(f28876,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP6(X1,'T_0',X0)) ) | $spl118),
  inference(subsumption_resolution,[],[f28823,f334])).
fof(f28823,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP6(X1,'T_0',X0)) ) | $spl118),
  inference(resolution,[],[f28514,f316])).
fof(f29016,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl108 | $spl118)),
  inference(duplicate_literal_removal,[],[f28993])).
fof(f28993,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl108 | $spl118)),
  inference(resolution,[],[f28904,f338])).
fof(f28904,plain,(
  ( ! [X24] : (sP9(X24,'T_0','F_0') | p(s(bool,X24))) ) | ($spl108 | $spl118)),
  inference(resolution,[],[f28877,f27096])).
fof(f29036,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | sP9(X8,'T_0',X9) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl8 | $spl118)),
  inference(resolution,[],[f28909,f11867])).
fof(f28909,plain,(
  ( ! [X4,X5] : (~sP7('T_0',X5,X4) | p(s(bool,X4)) | sP9(X4,'T_0',X5)) ) | $spl118),
  inference(subsumption_resolution,[],[f28892,f326])).
fof(f28892,plain,(
  ( ! [X4,X5] : (sP9(X4,'T_0',X5) | p(s(bool,X4)) | ~sP8(X4,X5,'T_0') | ~sP7('T_0',X5,X4)) ) | $spl118),
  inference(resolution,[],[f28877,f322])).
fof(f60343,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | sP9(X16,X17,'T_0')) ) | $spl318),
  inference(cnf_transformation,[],[f60343_D])).
fof(f60343_D,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | sP9(X16,X17,'T_0')) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f30162,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'T_0') | p(s(bool,X18)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl114 | $spl132)),
  inference(resolution,[],[f30142,f27497])).
fof(f30142,plain,(
  ( ! [X0,X1] : (sP6('T_0',X1,X0)) ) | $spl132),
  inference(subsumption_resolution,[],[f30141,f332])).
fof(f30141,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP6('T_0',X1,X0)) ) | $spl132),
  inference(subsumption_resolution,[],[f30088,f333])).
fof(f30088,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP6('T_0',X1,X0)) ) | $spl132),
  inference(resolution,[],[f30058,f316])).
fof(f30242,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | sP9(X16,X17,'T_0') | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl114 | $spl132)),
  inference(resolution,[],[f30174,f27460])).
fof(f30174,plain,(
  ( ! [X4,X5] : (~sP7(X5,'T_0',X4) | p(s(bool,X4)) | sP9(X4,X5,'T_0')) ) | $spl132),
  inference(subsumption_resolution,[],[f30155,f326])).
fof(f30155,plain,(
  ( ! [X4,X5] : (sP9(X4,X5,'T_0') | p(s(bool,X4)) | ~sP8(X4,'T_0',X5) | ~sP7(X5,'T_0',X4)) ) | $spl132),
  inference(resolution,[],[f30142,f322])).
fof(f60337,plain,(
  ( ! [X39,X41,X40] : (~sP9(X41,X40,X39) | p(s(bool,X39)) | ~p(s(bool,X41))) ) | $spl314),
  inference(cnf_transformation,[],[f60337_D])).
fof(f60337_D,plain,(
  ( ! [X39,X41,X40] : (~sP9(X41,X40,X39) | p(s(bool,X39)) | ~p(s(bool,X41))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f60342,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X37,X36,X38) | ~p(s(bool,X37)) | p(s(bool,X38))) ) | $spl142),
  inference(subsumption_resolution,[],[f31428,f319])).
fof(f31428,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X37,X36,X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | p(s(bool,X36))) ) | $spl142),
  inference(superposition,[],[f31183,f412])).
fof(f60341,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X40,X39,X41) | ~p(s(bool,X40)) | p(s(bool,X41))) ) | $spl142),
  inference(subsumption_resolution,[],[f31429,f319])).
fof(f31429,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X40,X39,X41) | ~p(s(bool,X40)) | p(s(bool,X41)) | p(s(bool,X39))) ) | $spl142),
  inference(superposition,[],[f31183,f414])).
fof(f31463,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X52,X51,X53) | ~p(s(bool,X52)) | p(s(bool,X53))) ) | $spl142),
  inference(subsumption_resolution,[],[f31433,f31183])).
fof(f31433,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X52,X51,X53) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl142),
  inference(superposition,[],[f31183,f442])).
fof(f60340,plain,(
  ( ! [X85,X84] : (~sP9(X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X85) | p(s(bool,X85))) ) | $spl316),
  inference(cnf_transformation,[],[f60340_D])).
fof(f60340_D,plain,(
  ( ! [X85,X84] : (~sP9(X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X85) | p(s(bool,X85))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f60339,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X85) | p(s(bool,X85))) ) | ($spl43 | $spl142)),
  inference(subsumption_resolution,[],[f31446,f34841])).
fof(f31446,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X85) | ~p(s(bool,X84)) | p(s(bool,X85))) ) | ($spl43 | $spl142)),
  inference(superposition,[],[f31183,f15564])).
fof(f60338,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X38,X37,X36) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl152),
  inference(subsumption_resolution,[],[f32052,f319])).
fof(f32052,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X37)) | ~sP9(X38,X37,X36) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl152),
  inference(superposition,[],[f31538,f412])).
fof(f60336,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X41,X40,X39) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl152),
  inference(subsumption_resolution,[],[f32053,f319])).
fof(f32053,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X40)) | ~sP9(X41,X40,X39) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl152),
  inference(superposition,[],[f31538,f414])).
fof(f60335,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | ~p(s(bool,X53)) | ~sP9(X53,X52,X51)) ) | $spl312),
  inference(cnf_transformation,[],[f60335_D])).
fof(f60335_D,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | ~p(s(bool,X53)) | ~sP9(X53,X52,X51)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f32087,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X52)) | ~sP9(X53,X52,X51) | ~p(s(bool,X53))) ) | $spl152),
  inference(subsumption_resolution,[],[f32057,f31538])).
fof(f32057,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X52)) | ~sP9(X53,X52,X51) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl152),
  inference(superposition,[],[f31538,f442])).
fof(f60334,plain,(
  ( ! [X85,X84] : (p(s(bool,X84)) | ~sP9(X85,X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl310),
  inference(cnf_transformation,[],[f60334_D])).
fof(f60334_D,plain,(
  ( ! [X85,X84] : (p(s(bool,X84)) | ~sP9(X85,X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f60333,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X84)) | ~sP9(X85,X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl43 | $spl152)),
  inference(subsumption_resolution,[],[f32070,f34884])).
fof(f32070,plain,(
  ( ! [X85,X84] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X84)) | ~sP9(X85,X84,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X85))) ) | ($spl43 | $spl152)),
  inference(superposition,[],[f31538,f15564])).
fof(f26907,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) | $spl104),
  inference(cnf_transformation,[],[f26907_D])).
fof(f26907_D,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f36260,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | p(s(bool,X30))) ) | $spl8),
  inference(subsumption_resolution,[],[f20150,f12161])).
fof(f12161,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f1003,f414])).
fof(f20150,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X30))) ) | $spl8),
  inference(superposition,[],[f12160,f414])).
fof(f36261,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | p(s(bool,X28))) ) | $spl8),
  inference(subsumption_resolution,[],[f20149,f12161])).
fof(f20149,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f12160,f412])).
fof(f60332,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl24),
  inference(subsumption_resolution,[],[f17387,f391])).
fof(f17387,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | ~p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl24),
  inference(superposition,[],[f11865,f418])).
fof(f60331,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl24),
  inference(subsumption_resolution,[],[f17386,f391])).
fof(f17386,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl24),
  inference(superposition,[],[f11865,f416])).
fof(f36280,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | p(s(bool,X39))) ) | $spl8),
  inference(subsumption_resolution,[],[f16114,f12161])).
fof(f16114,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,X38)) ) | $spl8),
  inference(superposition,[],[f12161,f442])).
fof(f60330,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31)) | p(s(bool,X30))) ) | $spl8),
  inference(subsumption_resolution,[],[f16110,f388])).
fof(f388,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f125])).
fof(f16110,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | p(s(bool,X30))) ) | $spl8),
  inference(superposition,[],[f12161,f414])).
fof(f60329,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl8),
  inference(subsumption_resolution,[],[f16109,f388])).
fof(f16109,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f12161,f412])).
fof(f36218,plain,(
  ( ! [X61,X60] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X61))) ) | $spl160),
  inference(cnf_transformation,[],[f36218_D])).
fof(f36218_D,plain,(
  ( ! [X61,X60] : (s(bool,X60) != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X61))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f37141,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl102),
  inference(duplicate_literal_removal,[],[f37097])).
fof(f37097,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl102),
  inference(resolution,[],[f26906,f338])).
fof(f37227,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13)) ) | $spl144),
  inference(duplicate_literal_removal,[],[f37160])).
fof(f37160,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,X14))) ) | $spl144),
  inference(resolution,[],[f31526,f336])).
fof(f60312,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl308),
  inference(cnf_transformation,[],[f60312_D])).
fof(f60312_D,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f38486,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl18 | $spl94)),
  inference(superposition,[],[f4610,f36884])).
fof(f60302,plain,(
  ( ! [X322] : (s(bool,X322) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X322)) ) | $spl306),
  inference(cnf_transformation,[],[f60302_D])).
fof(f60302_D,plain,(
  ( ! [X322] : (s(bool,X322) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X322)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f38893,plain,(
  ( ! [X322] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X322) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X322)) ) | ($spl8 | $spl21)),
  inference(superposition,[],[f12160,f37913])).
fof(f60300,plain,(
  ( ! [X331,X330] : (~sP9(X330,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X331) | ~p(s(bool,X330))) ) | $spl304),
  inference(cnf_transformation,[],[f60300_D])).
fof(f60300_D,plain,(
  ( ! [X331,X330] : (~sP9(X330,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X331) | ~p(s(bool,X330))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f60299,plain,(
  ( ! [X331,X330] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X330,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X331) | ~p(s(bool,X330))) ) | ($spl21 | $spl142)),
  inference(subsumption_resolution,[],[f38899,f38012])).
fof(f38899,plain,(
  ( ! [X331,X330] : (s(bool,'T_0') != s(bool,'F_0') | ~sP9(X330,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X331) | ~p(s(bool,X330)) | p(s(bool,X331))) ) | ($spl21 | $spl142)),
  inference(superposition,[],[f31183,f37913])).
fof(f60298,plain,(
  ( ! [X333,X332] : (p(s(bool,X332)) | ~sP9(X333,X332,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl302),
  inference(cnf_transformation,[],[f60298_D])).
fof(f60298_D,plain,(
  ( ! [X333,X332] : (p(s(bool,X332)) | ~sP9(X333,X332,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f60297,plain,(
  ( ! [X333,X332] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X332)) | ~sP9(X333,X332,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl21 | $spl152)),
  inference(subsumption_resolution,[],[f38900,f52057])).
fof(f38900,plain,(
  ( ! [X333,X332] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X332)) | ~sP9(X333,X332,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X333))) ) | ($spl21 | $spl152)),
  inference(superposition,[],[f31538,f37913])).
fof(f36157,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X13))) ) | $spl154),
  inference(cnf_transformation,[],[f36157_D])).
fof(f36157_D,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X13))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f40395,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X14) | s(bool,X13) != s(bool,X14)) ) | $spl142),
  inference(duplicate_literal_removal,[],[f40340])).
fof(f40340,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X14) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X12))) ) | $spl142),
  inference(resolution,[],[f31454,f336])).
fof(f40487,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,X14) | s(bool,X12) != s(bool,X14)) ) | $spl152),
  inference(duplicate_literal_removal,[],[f40431])).
fof(f40431,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,X14) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X13))) ) | $spl152),
  inference(resolution,[],[f32078,f336])).
fof(f41378,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X8))) ) | ($spl94 | $spl108)),
  inference(forward_demodulation,[],[f41377,f36884])).
fof(f41377,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl108)),
  inference(duplicate_literal_removal,[],[f41361])).
fof(f41361,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl108)),
  inference(resolution,[],[f37043,f338])).
fof(f37043,plain,(
  ( ! [X22] : (sP9(X22,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | p(s(bool,X22))) ) | ($spl94 | $spl108)),
  inference(resolution,[],[f36861,f27096])).
fof(f36861,plain,(
  ( ! [X114,X113] : (sP6(X113,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X114)) ) | $spl94),
  inference(resolution,[],[f14077,f333])).
fof(f29960,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | ($spl24 | $spl128)),
  inference(subsumption_resolution,[],[f29934,f11865])).
fof(f29934,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | ~p(s(bool,X34))) ) | $spl128),
  inference(superposition,[],[f29848,f418])).
fof(f29959,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | ~p(s(bool,X32))) ) | ($spl24 | $spl128)),
  inference(subsumption_resolution,[],[f29933,f11865])).
fof(f29933,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X32))) ) | $spl128),
  inference(superposition,[],[f29848,f416])).
fof(f28515,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | s(bool,'F_0') = s(bool,X68)) ) | $spl120),
  inference(cnf_transformation,[],[f28515_D])).
fof(f28515_D,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | s(bool,'F_0') = s(bool,X68)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f29961,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl128),
  inference(subsumption_resolution,[],[f29935,f29848])).
fof(f29935,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl128),
  inference(superposition,[],[f29848,f442])).
fof(f45819,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X62))) ) | ($spl2 | $spl246)),
  inference(duplicate_literal_removal,[],[f45783])).
fof(f45783,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X62)) | p(s(bool,X61)) | p(s(bool,X62))) ) | ($spl2 | $spl246)),
  inference(resolution,[],[f45677,f1663])).
fof(f45857,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30)) | p(s(bool,X31))) ) | ($spl2 | $spl246)),
  inference(resolution,[],[f45831,f1658])).
fof(f45936,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl116),
  inference(duplicate_literal_removal,[],[f45874])).
fof(f45874,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl116),
  inference(resolution,[],[f28512,f338])).
fof(f60248,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl128),
  inference(subsumption_resolution,[],[f46112,f412])).
fof(f46112,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl128),
  inference(superposition,[],[f29848,f388])).
fof(f60247,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl128),
  inference(subsumption_resolution,[],[f46113,f412])).
fof(f46113,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl128),
  inference(superposition,[],[f29848,f388])).
fof(f46142,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | ~p(s(bool,X32))) ) | ($spl24 | $spl128)),
  inference(subsumption_resolution,[],[f46118,f11865])).
fof(f46118,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X32))) ) | $spl128),
  inference(superposition,[],[f29848,f416])).
fof(f46143,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | ($spl24 | $spl128)),
  inference(subsumption_resolution,[],[f46119,f11865])).
fof(f46119,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | ~p(s(bool,X34))) ) | $spl128),
  inference(superposition,[],[f29848,f418])).
fof(f46144,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl128),
  inference(subsumption_resolution,[],[f46120,f29848])).
fof(f46120,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl128),
  inference(superposition,[],[f29848,f442])).
fof(f60246,plain,(
  ( ! [X53] : (s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'F_0') = s(bool,X53)) ) | $spl300),
  inference(cnf_transformation,[],[f60246_D])).
fof(f60246_D,plain,(
  ( ! [X53] : (s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'F_0') = s(bool,X53)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f46129,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'F_0') = s(bool,X53)) ) | ($spl94 | $spl128)),
  inference(superposition,[],[f29848,f36884])).
fof(f46480,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl130),
  inference(duplicate_literal_removal,[],[f46417])).
fof(f46417,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl130),
  inference(resolution,[],[f30056,f338])).
fof(f36193,plain,(
  ( ! [X17,X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17) | ~p(s(bool,X17))) ) | $spl158),
  inference(cnf_transformation,[],[f36193_D])).
fof(f36193_D,plain,(
  ( ! [X17,X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17) | ~p(s(bool,X17))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f47131,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) != s(bool,X16)) ) | ($spl8 | $spl96)),
  inference(resolution,[],[f26570,f13485])).
fof(f47221,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f47154])).
fof(f47154,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl96),
  inference(resolution,[],[f26344,f338])).
fof(f36169,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,X15) != s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | $spl156),
  inference(cnf_transformation,[],[f36169_D])).
fof(f36169_D,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,X15) != s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f47434,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl96 | $spl130)),
  inference(duplicate_literal_removal,[],[f47375])).
fof(f47375,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl96 | $spl130)),
  inference(resolution,[],[f47133,f338])).
fof(f47133,plain,(
  ( ! [X23,X21,X22] : (sP9(X22,X21,X23) | s(bool,'T_0') = s(bool,X21) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,X23)) ) | ($spl96 | $spl130)),
  inference(resolution,[],[f26570,f30672])).
fof(f36405,plain,(
  ( ! [X47,X48] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl2 | $spl8 | $spl106)),
  inference(duplicate_literal_removal,[],[f36376])).
fof(f36376,plain,(
  ( ! [X47,X48] : (p(s(bool,X47)) | p(s(bool,X48)) | p(s(bool,X47)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X48))) ) | ($spl2 | $spl8 | $spl106)),
  inference(resolution,[],[f1663,f28042])).
fof(f28042,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8))) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f28039,f329])).
fof(f28039,plain,(
  ( ! [X6,X8,X7] : (~sP7(X6,X8,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | ($spl8 | $spl106)),
  inference(subsumption_resolution,[],[f28018,f326])).
fof(f28018,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f27884,f322])).
fof(f27884,plain,(
  ( ! [X2,X0,X1] : (sP6(X1,X2,X0) | s(bool,'T_0') != s(bool,X2)) ) | ($spl8 | $spl106)),
  inference(subsumption_resolution,[],[f27883,f332])).
fof(f27883,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X2) | sP6(X1,X2,X0)) ) | ($spl8 | $spl106)),
  inference(subsumption_resolution,[],[f27820,f11558])).
fof(f27820,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') != s(bool,X2) | sP6(X1,X2,X0)) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f27663,f316])).
fof(f27663,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X1)) ) | ($spl8 | $spl106)),
  inference(subsumption_resolution,[],[f27639,f442])).
fof(f27639,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X1)) ) | ($spl8 | $spl106)),
  inference(resolution,[],[f27389,f21149])).
fof(f21149,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(subsumption_resolution,[],[f21121,f332])).
fof(f21121,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl8),
  inference(resolution,[],[f12178,f316])).
fof(f12178,plain,(
  ( ! [X4,X5,X3] : (sP9(X5,X3,X4) | s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl8),
  inference(superposition,[],[f1003,f335])).
fof(f36340,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl2 | $spl8 | $spl106)),
  inference(resolution,[],[f1658,f27884])).
fof(f47889,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl106),
  inference(duplicate_literal_removal,[],[f47824])).
fof(f47824,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl106),
  inference(resolution,[],[f26912,f338])).
fof(f60223,plain,(
  ( ! [X74,X73] : (p(s(bool,X74)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X73))) ) | ($spl2 | $spl106 | $spl246)),
  inference(subsumption_resolution,[],[f48001,f415])).
fof(f415,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f230])).
fof(f48001,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,X73) | p(s(bool,X74)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X73))) ) | ($spl2 | $spl106 | $spl246)),
  inference(duplicate_literal_removal,[],[f47970])).
fof(f47970,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,X73) | p(s(bool,X74)) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X74)) | p(s(bool,X73))) ) | ($spl2 | $spl106 | $spl246)),
  inference(resolution,[],[f47691,f1663])).
fof(f47691,plain,(
  ( ! [X26,X24,X25] : (sP9(X25,X26,X24) | s(bool,'T_0') = s(bool,X24) | p(s(bool,X25)) | s(bool,'T_0') != s(bool,X26)) ) | ($spl106 | $spl246)),
  inference(resolution,[],[f27353,f45868])).
fof(f48092,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl106 | $spl116)),
  inference(duplicate_literal_removal,[],[f48032])).
fof(f48032,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl106 | $spl116)),
  inference(resolution,[],[f47689,f338])).
fof(f47689,plain,(
  ( ! [X19,X20,X18] : (sP9(X19,X20,X18) | s(bool,'T_0') = s(bool,X18) | p(s(bool,X19)) | s(bool,'F_0') = s(bool,X20)) ) | ($spl106 | $spl116)),
  inference(resolution,[],[f27353,f29496])).
fof(f49574,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X10)) | ~p(s(bool,X11))) ) | ($spl8 | $spl94)),
  inference(forward_demodulation,[],[f49555,f36884])).
fof(f49555,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36857,f11867])).
fof(f49623,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13)) ) | ($spl8 | $spl94)),
  inference(forward_demodulation,[],[f49622,f36884])).
fof(f49622,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X13)) ) | ($spl8 | $spl94)),
  inference(subsumption_resolution,[],[f49605,f415])).
fof(f49605,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X13)) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36858,f21150])).
fof(f21150,plain,(
  ( ! [X10,X11,X9] : (sP7(X9,X10,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl8),
  inference(subsumption_resolution,[],[f21122,f328])).
fof(f21122,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X9,X10,X11)) ) | $spl8),
  inference(resolution,[],[f12178,f317])).
fof(f49624,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl94 | $spl114)),
  inference(forward_demodulation,[],[f49608,f36884])).
fof(f49608,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114)),
  inference(resolution,[],[f36858,f27460])).
fof(f50655,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl8 | $spl94)),
  inference(forward_demodulation,[],[f50654,f36884])).
fof(f50654,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(subsumption_resolution,[],[f50641,f416])).
fof(f50641,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(duplicate_literal_removal,[],[f50561])).
fof(f50561,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X14))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f49551,f12177])).
fof(f12177,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f1003,f335])).
fof(f50659,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X27)) | p(s(bool,X28))) ) | ($spl8 | $spl94)),
  inference(forward_demodulation,[],[f50635,f36884])).
fof(f50635,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(duplicate_literal_removal,[],[f50568])).
fof(f50568,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X27)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X28))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f49551,f19666])).
fof(f19666,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl8),
  inference(resolution,[],[f12101,f326])).
fof(f50663,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X45)) | p(s(bool,X46))) ) | ($spl8 | $spl94 | $spl130)),
  inference(forward_demodulation,[],[f50662,f36884])).
fof(f50662,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl130)),
  inference(subsumption_resolution,[],[f50626,f411])).
fof(f50626,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl130)),
  inference(duplicate_literal_removal,[],[f50577])).
fof(f50577,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl130)),
  inference(resolution,[],[f49551,f46520])).
fof(f46520,plain,(
  ( ! [X14,X12,X13] : (sP9(X12,X13,X14) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') != s(bool,X13)) ) | ($spl8 | $spl130)),
  inference(resolution,[],[f30672,f11867])).
fof(f50665,plain,(
  ( ! [X55] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X55))) ) | ($spl8 | $spl94 | $spl132)),
  inference(forward_demodulation,[],[f50664,f36884])).
fof(f50664,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl132)),
  inference(subsumption_resolution,[],[f50621,f399])).
fof(f399,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_CLAUSES|split|1')).
fof(f50621,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl132)),
  inference(duplicate_literal_removal,[],[f50582])).
fof(f50582,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94 | $spl132)),
  inference(resolution,[],[f49551,f30238])).
fof(f30238,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'T_0') | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X9)) ) | ($spl8 | $spl132)),
  inference(resolution,[],[f30174,f11867])).
fof(f50810,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X14)) | p(s(bool,X15))) ) | ($spl8 | $spl94)),
  inference(forward_demodulation,[],[f50809,f36884])).
fof(f50809,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(subsumption_resolution,[],[f50799,f416])).
fof(f50799,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') != s(bool,X14) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(duplicate_literal_removal,[],[f50713])).
fof(f50713,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') != s(bool,X14) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X15))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f49600,f12178])).
fof(f50812,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34)) | p(s(bool,X35))) ) | ($spl94 | $spl114)),
  inference(forward_demodulation,[],[f50790,f36884])).
fof(f50790,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114)),
  inference(duplicate_literal_removal,[],[f50723])).
fof(f50723,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | p(s(bool,X35)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X34))) ) | ($spl94 | $spl114)),
  inference(resolution,[],[f49600,f26923])).
fof(f50818,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X42)) | p(s(bool,X43))) ) | ($spl94 | $spl114 | $spl116)),
  inference(forward_demodulation,[],[f50817,f36884])).
fof(f50817,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl116)),
  inference(subsumption_resolution,[],[f50786,f411])).
fof(f50786,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl116)),
  inference(duplicate_literal_removal,[],[f50727])).
fof(f50727,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') = s(bool,X42) | p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl116)),
  inference(resolution,[],[f49600,f45743])).
fof(f45743,plain,(
  ( ! [X23,X21,X22] : (sP9(X22,X21,X23) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23)) ) | ($spl114 | $spl116)),
  inference(resolution,[],[f29460,f27497])).
fof(f50820,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X54))) ) | ($spl94 | $spl114 | $spl118)),
  inference(forward_demodulation,[],[f50819,f36884])).
fof(f50819,plain,(
  ( ! [X54] : (p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl118)),
  inference(subsumption_resolution,[],[f50780,f399])).
fof(f50780,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl118)),
  inference(duplicate_literal_removal,[],[f50733])).
fof(f50733,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114 | $spl118)),
  inference(resolution,[],[f49600,f28899])).
fof(f28899,plain,(
  ( ! [X19,X18] : (sP9(X18,'T_0',X19) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19)) ) | ($spl114 | $spl118)),
  inference(resolution,[],[f28877,f27497])).
fof(f60180,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl298),
  inference(cnf_transformation,[],[f60180_D])).
fof(f60180_D,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f54186,plain,(
  ( ! [X320] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X320) | s(bool,X320) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl18 | $spl264)),
  inference(superposition,[],[f4610,f53360])).
fof(f60176,plain,(
  ( ! [X334] : (s(bool,X334) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X334)) ) | $spl296),
  inference(cnf_transformation,[],[f60176_D])).
fof(f60176_D,plain,(
  ( ! [X334] : (s(bool,X334) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X334)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f54197,plain,(
  ( ! [X334] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X334) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'F_0') = s(bool,X334)) ) | ($spl128 | $spl264)),
  inference(superposition,[],[f29848,f53360])).
fof(f55380,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X8))) ) | ($spl108 | $spl264)),
  inference(forward_demodulation,[],[f55379,f53360])).
fof(f55379,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl108 | $spl264)),
  inference(duplicate_literal_removal,[],[f55350])).
fof(f55350,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl108 | $spl264)),
  inference(resolution,[],[f53629,f338])).
fof(f53629,plain,(
  ( ! [X24] : (sP9(X24,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),'F_0') | p(s(bool,X24))) ) | ($spl108 | $spl264)),
  inference(resolution,[],[f53337,f27096])).
fof(f53337,plain,(
  ( ! [X114,X113] : (sP6(X113,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X114)) ) | $spl264),
  inference(resolution,[],[f53275,f333])).
fof(f57948,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X8)) | ~p(s(bool,X9))) ) | ($spl8 | $spl264)),
  inference(forward_demodulation,[],[f57928,f53360])).
fof(f57928,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53333,f11867])).
fof(f57999,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | ($spl8 | $spl264)),
  inference(forward_demodulation,[],[f57998,f53360])).
fof(f57998,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,X11)) ) | ($spl8 | $spl264)),
  inference(subsumption_resolution,[],[f57980,f415])).
fof(f57980,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,X11)) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53334,f21150])).
fof(f58000,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl114 | $spl264)),
  inference(forward_demodulation,[],[f57983,f53360])).
fof(f57983,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl264)),
  inference(resolution,[],[f53334,f27460])).
fof(f59715,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X13)) | p(s(bool,X14))) ) | ($spl8 | $spl264)),
  inference(forward_demodulation,[],[f59714,f53360])).
fof(f59714,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(subsumption_resolution,[],[f59701,f416])).
fof(f59701,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(duplicate_literal_removal,[],[f59613])).
fof(f59613,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X14))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f57924,f12177])).
fof(f59719,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X27)) | p(s(bool,X28))) ) | ($spl8 | $spl264)),
  inference(forward_demodulation,[],[f59695,f53360])).
fof(f59695,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(duplicate_literal_removal,[],[f59620])).
fof(f59620,plain,(
  ( ! [X28,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | ~p(s(bool,X27)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X28))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f57924,f19666])).
fof(f59723,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X45)) | p(s(bool,X46))) ) | ($spl8 | $spl130 | $spl264)),
  inference(forward_demodulation,[],[f59722,f53360])).
fof(f59722,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl130 | $spl264)),
  inference(subsumption_resolution,[],[f59686,f411])).
fof(f59686,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl130 | $spl264)),
  inference(duplicate_literal_removal,[],[f59629])).
fof(f59629,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl130 | $spl264)),
  inference(resolution,[],[f57924,f46520])).
fof(f59725,plain,(
  ( ! [X55] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X55))) ) | ($spl8 | $spl132 | $spl264)),
  inference(forward_demodulation,[],[f59724,f53360])).
fof(f59724,plain,(
  ( ! [X55] : (p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl132 | $spl264)),
  inference(subsumption_resolution,[],[f59681,f399])).
fof(f59681,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl132 | $spl264)),
  inference(duplicate_literal_removal,[],[f59634])).
fof(f59634,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X55)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl132 | $spl264)),
  inference(resolution,[],[f57924,f30238])).
fof(f59894,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X14)) | p(s(bool,X15))) ) | ($spl8 | $spl264)),
  inference(forward_demodulation,[],[f59893,f53360])).
fof(f59893,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(subsumption_resolution,[],[f59883,f416])).
fof(f59883,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') != s(bool,X14) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(duplicate_literal_removal,[],[f59789])).
fof(f59789,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') != s(bool,X14) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X15))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f57975,f12178])).
fof(f59896,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34)) | p(s(bool,X35))) ) | ($spl114 | $spl264)),
  inference(forward_demodulation,[],[f59874,f53360])).
fof(f59874,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl264)),
  inference(duplicate_literal_removal,[],[f59799])).
fof(f59799,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | p(s(bool,X35)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X34))) ) | ($spl114 | $spl264)),
  inference(resolution,[],[f57975,f26923])).
fof(f59902,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X42)) | p(s(bool,X43))) ) | ($spl114 | $spl116 | $spl264)),
  inference(forward_demodulation,[],[f59901,f53360])).
fof(f59901,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl116 | $spl264)),
  inference(subsumption_resolution,[],[f59870,f411])).
fof(f59870,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl116 | $spl264)),
  inference(duplicate_literal_removal,[],[f59803])).
fof(f59803,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | s(bool,'F_0') = s(bool,X42) | p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl116 | $spl264)),
  inference(resolution,[],[f57975,f45743])).
fof(f59904,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X54))) ) | ($spl114 | $spl118 | $spl264)),
  inference(forward_demodulation,[],[f59903,f53360])).
fof(f59903,plain,(
  ( ! [X54] : (p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl118 | $spl264)),
  inference(subsumption_resolution,[],[f59864,f399])).
fof(f59864,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl118 | $spl264)),
  inference(duplicate_literal_removal,[],[f59809])).
fof(f59809,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | p(s(bool,X54)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl118 | $spl264)),
  inference(resolution,[],[f57975,f28899])).
fof(f60119,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))))) ) | $spl294),
  inference(cnf_transformation,[],[f60119_D])).
fof(f60119_D,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f60116,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.bool.IN_2'(s(sK12,X34),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X34),s(fun(sK12,bool),sK13)))) | p(s(bool,X35))) ) | ($spl0 | $spl264)),
  inference(resolution,[],[f43014,f53287])).
fof(f53287,plain,(
  ( ! [X14,X13] : (~sP0('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X13,X14) | ~p(s(bool,X14)) | p(s(bool,X13))) ) | $spl264),
  inference(resolution,[],[f53275,f272])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f166])).
fof(f43014,plain,(
  ( ! [X6,X8,X7] : (sP0(X6,X7,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X8),s(fun(sK12,bool),sK15))))) ) | $spl0),
  inference(resolution,[],[f18836,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f18836,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(sK12,X25),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP0(X26,X27,'const.bool.IN_2'(s(sK12,X25),s(fun(sK12,bool),sK13)))) ) | $spl0),
  inference(resolution,[],[f4804,f273])).
fof(f4804,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f467,f434])).
fof(f60115,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.bool.IN_2'(s(sK12,X32),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X32),s(fun(sK12,bool),sK13)))) | p(s(bool,X33))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f43014,f36811])).
fof(f36811,plain,(
  ( ! [X14,X13] : (~sP0('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X13,X14) | ~p(s(bool,X14)) | p(s(bool,X13))) ) | $spl94),
  inference(resolution,[],[f14077,f272])).
fof(f60114,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),sK13)))) | p(s(bool,X31))) ) | ($spl0 | $spl2 | $spl41)),
  inference(resolution,[],[f43014,f36622])).
fof(f36622,plain,(
  ( ! [X12,X11] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X11,X12) | ~p(s(bool,X12)) | p(s(bool,X11))) ) | ($spl2 | $spl41)),
  inference(resolution,[],[f36441,f272])).
fof(f36441,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | ($spl2 | $spl41)),
  inference(resolution,[],[f541,f14447])).
fof(f14447,plain,(
  ( ! [X127] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X127) | p(s(bool,X127))) ) | $spl41),
  inference(subsumption_resolution,[],[f14410,f427])).
fof(f14410,plain,(
  ( ! [X127] : (p(s(bool,'F_0')) | p(s(bool,X127)) | ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X127)) ) | $spl41),
  inference(resolution,[],[f14336,f375])).
fof(f14336,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl41),
  inference(trivial_inequality_removal,[],[f14332])).
fof(f14332,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | $spl41),
  inference(superposition,[],[f13004,f418])).
fof(f541,plain,(
  ( ! [X127] : (sP11(X127,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl2),
  inference(resolution,[],[f469,f376])).
fof(f60113,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),sK13)))) | p(s(bool,X29))) ) | $spl0),
  inference(resolution,[],[f43014,f4810])).
fof(f4810,plain,(
  ( ! [X12,X11] : (~sP0('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X11,X12) | ~p(s(bool,X12)) | p(s(bool,X11))) ) | $spl0),
  inference(resolution,[],[f467,f272])).
fof(f60098,plain,(
  ( ! [X22] : (p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl292),
  inference(cnf_transformation,[],[f60098_D])).
fof(f60098_D,plain,(
  ( ! [X22] : (p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f60111,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X24)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X25),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f43014,f18835])).
fof(f18835,plain,(
  ( ! [X24,X23,X22] : (~sP0(X23,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),sK13)),X24) | ~p(s(bool,X23)) | ~p(s(bool,X24)) | p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f272])).
fof(f60110,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X22))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f43014,f37928])).
fof(f37928,plain,(
  ( ! [X14,X13] : (~sP0(X13,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X14) | ~p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl21),
  inference(resolution,[],[f37914,f272])).
fof(f60109,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.IN_2'(s(sK12,X19),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X19),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X20))) ) | ($spl0 | $spl43)),
  inference(resolution,[],[f43014,f16478])).
fof(f16478,plain,(
  ( ! [X26,X27] : (~sP0(X26,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X27) | ~p(s(bool,X27)) | ~p(s(bool,X26))) ) | $spl43),
  inference(subsumption_resolution,[],[f16236,f397])).
fof(f16236,plain,(
  ( ! [X26,X27] : (p(s(bool,'F_0')) | ~p(s(bool,X26)) | ~p(s(bool,X27)) | ~sP0(X26,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X27)) ) | $spl43),
  inference(superposition,[],[f272,f15564])).
fof(f60108,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X18))) ) | ($spl0 | $spl41)),
  inference(resolution,[],[f43014,f14347])).
fof(f14347,plain,(
  ( ! [X10,X11] : (~sP0(X10,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X11) | ~p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl41),
  inference(resolution,[],[f14336,f272])).
fof(f60103,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))))) ) | $spl0),
  inference(resolution,[],[f43014,f272])).
fof(f60118,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl43 | $spl94)),
  inference(subsumption_resolution,[],[f60117,f397])).
fof(f60117,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl43 | $spl94)),
  inference(forward_demodulation,[],[f60102,f15564])).
fof(f60102,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f43014,f36812])).
fof(f36812,plain,(
  ( ! [X15,X16] : (~sP0(X15,X16,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl94),
  inference(resolution,[],[f14077,f272])).
fof(f53276,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))))) ) | $spl266),
  inference(cnf_transformation,[],[f53276_D])).
fof(f53276_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f60089,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK14))))) ) | $spl0),
  inference(resolution,[],[f18835,f43013])).
fof(f43013,plain,(
  ( ! [X4,X5,X3] : (sP0(X3,X4,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),sK14))))) ) | $spl0),
  inference(resolution,[],[f18836,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f48457,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | $spl252),
  inference(cnf_transformation,[],[f48457_D])).
fof(f48457_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f60081,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18835,f36813])).
fof(f36813,plain,(
  ( ! [X17,X18] : (sP0(X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X18)) ) | $spl94),
  inference(resolution,[],[f14077,f274])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f166])).
fof(f58909,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f58832])).
fof(f58832,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X3) != s(bool,X4)) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18888,f36863])).
fof(f18888,plain,(
  ( ! [X182,X183,X181] : (~sP9('const.bool.IN_2'(s(sK12,X181),s(fun(sK12,bool),sK13)),X183,X182) | s(bool,X182) != s(bool,X183) | p(s(bool,'const.bool.IN_2'(s(sK12,X181),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f338])).
fof(f58911,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl8 | $spl94)),
  inference(subsumption_resolution,[],[f58831,f388])).
fof(f58831,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2)) | p(s(bool,X1))) ) | ($spl0 | $spl8 | $spl94)),
  inference(resolution,[],[f18888,f37016])).
fof(f37016,plain,(
  ( ! [X10,X11] : (sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X10,X11) | p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36860,f13414])).
fof(f13414,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | sP9(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5))) ) | $spl8),
  inference(resolution,[],[f12102,f325])).
fof(f12102,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | sP9(X7,X6,X8) | ~sP6(X8,X6,X7)) ) | $spl8),
  inference(subsumption_resolution,[],[f12098,f412])).
fof(f12098,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl8),
  inference(resolution,[],[f11867,f321])).
fof(f36860,plain,(
  ( ! [X111,X112] : (sP6(X111,X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl94),
  inference(resolution,[],[f14077,f332])).
fof(f58910,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0))) ) | ($spl0 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f58830,f412])).
fof(f58830,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f18888,f12807])).
fof(f12807,plain,(
  ( ! [X9] : (sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0',X9) | p(s(bool,X9))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12268,f12578])).
fof(f12578,plain,(
  ( ! [X2,X3] : (~sP6(X3,'F_0',X2) | sP9(X2,'F_0',X3) | p(s(bool,X3))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f12575,f325])).
fof(f12575,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | sP9(X4,'F_0',X5) | ~sP6(X5,'F_0',X4)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12572,f427])).
fof(f12572,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f12564,f321])).
fof(f12564,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12563,f329])).
fof(f12563,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP7('F_0',X6,X7)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12550,f328])).
fof(f12550,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f12536,f317])).
fof(f12536,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f12517,f326])).
fof(f12517,plain,(
  ( ! [X2,X3] : (~sP8(X3,X2,'F_0') | ~p(s(bool,X2)) | sP9(X3,'F_0',X2)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12516,f416])).
fof(f12516,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12512,f11558])).
fof(f12512,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl0),
  inference(resolution,[],[f12506,f320])).
fof(f12506,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f12494,f328])).
fof(f12494,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f11945,f317])).
fof(f11945,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl0),
  inference(superposition,[],[f10443,f335])).
fof(f10443,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(backward_demodulation,[],[f4883,f4881])).
fof(f4881,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | $spl0),
  inference(resolution,[],[f467,f411])).
fof(f4883,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | $spl0),
  inference(resolution,[],[f467,f416])).
fof(f12268,plain,(
  ( ! [X111,X112] : (sP6(X111,X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl5),
  inference(resolution,[],[f11960,f332])).
fof(f11960,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl5),
  inference(resolution,[],[f472,f435])).
fof(f58788,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X4) = s(bool,X5)) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f58714])).
fof(f58714,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X4) = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18887,f36864])).
fof(f36864,plain,(
  ( ! [X120,X119] : (~sP9('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X120,X119) | s(bool,X119) = s(bool,X120)) ) | $spl94),
  inference(resolution,[],[f14077,f337])).
fof(f18887,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,'const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP9('const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),sK13)),X179,X180) | s(bool,X179) = s(bool,X180)) ) | $spl0),
  inference(resolution,[],[f4804,f335])).
fof(f58791,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X0))) ) | $spl286),
  inference(cnf_transformation,[],[f58791_D])).
fof(f58791_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X0))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f58792,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X2) = s(bool,X3) | p(s(bool,X3))) ) | ($spl0 | $spl94)),
  inference(subsumption_resolution,[],[f58713,f388])).
fof(f58713,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X2) = s(bool,X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18887,f36851])).
fof(f58790,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(subsumption_resolution,[],[f58712,f388])).
fof(f58712,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18887,f49495])).
fof(f56815,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X11),s(fun(sK12,bool),sK13)))) | p(s(bool,X12)) | p(s(bool,'const.bool.IN_2'(s(sK12,X11),s(fun(sK12,bool),sK14))))) ) | ($spl0 | $spl264)),
  inference(resolution,[],[f53287,f43013])).
fof(f56791,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f56743])).
fof(f56743,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18871,f36849])).
fof(f36849,plain,(
  ( ! [X90,X89] : (~sP5(X89,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X90) | ~p(s(bool,X90)) | p(s(bool,X89))) ) | $spl94),
  inference(resolution,[],[f14077,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f54,f147,f146])).
fof(f147,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f54,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUBSET_INTER','thm.sat.sat_23_prove_ltb')).
fof(f18871,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'const.bool.IN_2'(s(sK12,X130),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP5(X131,'const.bool.IN_2'(s(sK12,X130),s(fun(sK12,bool),sK13)),X132) | p(s(bool,X131))) ) | $spl0),
  inference(resolution,[],[f4804,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f56551,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X1)) | p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18870,f36835])).
fof(f36835,plain,(
  ( ! [X61,X62] : (~sP5('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X62,X61) | p(s(bool,X61))) ) | $spl94),
  inference(resolution,[],[f14077,f300])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f18870,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,'const.bool.IN_2'(s(sK12,X127),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP5('const.bool.IN_2'(s(sK12,X127),s(fun(sK12,bool),sK13)),X128,X129) | p(s(bool,X128))) ) | $spl0),
  inference(resolution,[],[f4804,f311])).
fof(f56591,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f56550])).
fof(f56550,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18870,f36836])).
fof(f36836,plain,(
  ( ! [X64,X63] : (~sP5('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X63,X64) | p(s(bool,X63))) ) | $spl94),
  inference(resolution,[],[f14077,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f56062,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,X10))) ) | ($spl264 | $spl276)),
  inference(forward_demodulation,[],[f56060,f53360])).
fof(f56060,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl276),
  inference(duplicate_literal_removal,[],[f56006])).
fof(f56006,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl276),
  inference(resolution,[],[f53637,f338])).
fof(f55984,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f55936])).
fof(f55936,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18869,f36850])).
fof(f36850,plain,(
  ( ! [X92,X91] : (~sP5(X91,X92,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X92)) | p(s(bool,X91))) ) | $spl94),
  inference(resolution,[],[f14077,f315])).
fof(f18869,plain,(
  ( ! [X125,X126,X124] : (p(s(bool,'const.bool.IN_2'(s(sK12,X124),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP5(X125,X126,'const.bool.IN_2'(s(sK12,X124),s(fun(sK12,bool),sK13))) | p(s(bool,X125))) ) | $spl0),
  inference(resolution,[],[f4804,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f55758,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl264 | $spl270)),
  inference(subsumption_resolution,[],[f55757,f416])).
fof(f55757,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl264 | $spl270)),
  inference(forward_demodulation,[],[f55699,f53360])).
fof(f55699,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl270),
  inference(resolution,[],[f53564,f338])).
fof(f55565,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,X10))) ) | ($spl264 | $spl278)),
  inference(forward_demodulation,[],[f55563,f53360])).
fof(f55563,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl278),
  inference(duplicate_literal_removal,[],[f55506])).
fof(f55506,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl278),
  inference(resolution,[],[f53661,f338])).
fof(f55332,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f55292])).
fof(f55292,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2)) | p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18868,f36835])).
fof(f18868,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,'const.bool.IN_2'(s(sK12,X121),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP5('const.bool.IN_2'(s(sK12,X121),s(fun(sK12,bool),sK13)),X122,X123) | p(s(bool,X123))) ) | $spl0),
  inference(resolution,[],[f4804,f310])).
fof(f55291,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18868,f36836])).
fof(f55175,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f55148])).
fof(f55148,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18862,f36844])).
fof(f36844,plain,(
  ( ! [X80,X79] : (sP5(X79,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X80) | ~p(s(bool,X80)) | ~p(s(bool,X79))) ) | $spl94),
  inference(resolution,[],[f14077,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f18862,plain,(
  ( ! [X103,X105,X104] : (~sP5(X104,'const.bool.IN_2'(s(sK12,X103),s(fun(sK12,bool),sK13)),X105) | ~p(s(bool,X104)) | p(s(bool,'const.bool.IN_2'(s(sK12,X103),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f301])).
fof(f55128,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f55101])).
fof(f55101,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18861,f36845])).
fof(f36845,plain,(
  ( ! [X81,X82] : (sP5(X81,X82,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X82)) | ~p(s(bool,X81))) ) | $spl94),
  inference(resolution,[],[f14077,f312])).
fof(f18861,plain,(
  ( ! [X101,X102,X100] : (~sP5(X101,X102,'const.bool.IN_2'(s(sK12,X100),s(fun(sK12,bool),sK13))) | ~p(s(bool,X101)) | p(s(bool,'const.bool.IN_2'(s(sK12,X100),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f300])).
fof(f55063,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18846,f36829])).
fof(f36829,plain,(
  ( ! [X50,X49] : (sP3('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X49,X50) | ~p(s(bool,X50))) ) | $spl94),
  inference(resolution,[],[f14077,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f52,f144,f143])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUBSET_INTER','thm.sat.sat_25_prove_ltb')).
fof(f18846,plain,(
  ( ! [X57,X56,X55] : (~sP3('const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),sK13)),X56,X57) | ~p(s(bool,X56)) | p(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f283])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f170])).
fof(f170,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f55086,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f55062])).
fof(f55062,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18846,f36831])).
fof(f36831,plain,(
  ( ! [X54,X53] : (sP3('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X53,X54) | ~p(s(bool,X53))) ) | $spl94),
  inference(resolution,[],[f14077,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f54897,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f54874])).
fof(f54874,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18845,f36829])).
fof(f18845,plain,(
  ( ! [X54,X52,X53] : (~sP3('const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),sK13)),X54,X53) | ~p(s(bool,X53)) | p(s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f282])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f54873,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18845,f36831])).
fof(f54383,plain,(
  ( ! [X350,X349] : (p(s(bool,X349)) | p(s(bool,X350)) | sP9(X350,X349,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl284),
  inference(cnf_transformation,[],[f54383_D])).
fof(f54383_D,plain,(
  ( ! [X350,X349] : (p(s(bool,X349)) | p(s(bool,X350)) | sP9(X350,X349,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f54382,plain,(
  ( ! [X350,X349,X348] : (~p(s(bool,X348)) | p(s(bool,X349)) | sP9(X350,X349,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X350))) ) | ($spl0 | $spl264)),
  inference(subsumption_resolution,[],[f54203,f416])).
fof(f54203,plain,(
  ( ! [X350,X349,X348] : (s(bool,'T_0') != s(bool,X348) | ~p(s(bool,X348)) | p(s(bool,X349)) | sP9(X350,X349,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X350))) ) | ($spl0 | $spl264)),
  inference(superposition,[],[f36725,f53360])).
fof(f36725,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X6)) | p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl0),
  inference(superposition,[],[f18043,f335])).
fof(f18043,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f4850,f4862])).
fof(f4862,plain,(
  ( ! [X116,X115] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X115,X116) | s(bool,X115) != s(bool,X116)) ) | $spl0),
  inference(resolution,[],[f467,f336])).
fof(f4850,plain,(
  ( ! [X92,X91] : (~sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))),X91,X92) | p(s(bool,X92)) | ~p(s(bool,X91))) ) | $spl0),
  inference(resolution,[],[f467,f319])).
fof(f54381,plain,(
  ( ! [X344,X343] : (p(s(bool,X343)) | p(s(bool,X344)) | sP9(X344,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X343)) ) | $spl282),
  inference(cnf_transformation,[],[f54381_D])).
fof(f54381_D,plain,(
  ( ! [X344,X343] : (p(s(bool,X343)) | p(s(bool,X344)) | sP9(X344,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X343)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f54380,plain,(
  ( ! [X344,X343,X342] : (~p(s(bool,X342)) | p(s(bool,X343)) | sP9(X344,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X343) | p(s(bool,X344))) ) | ($spl0 | $spl264)),
  inference(subsumption_resolution,[],[f54201,f416])).
fof(f54201,plain,(
  ( ! [X344,X343,X342] : (s(bool,'T_0') != s(bool,X342) | ~p(s(bool,X342)) | p(s(bool,X343)) | sP9(X344,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X343) | p(s(bool,X344))) ) | ($spl0 | $spl264)),
  inference(superposition,[],[f36724,f53360])).
fof(f36724,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f18043,f335])).
fof(f53936,plain,(
  ( ! [X405,X404] : (sP1('T_0',X404,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl280),
  inference(cnf_transformation,[],[f53936_D])).
fof(f53936_D,plain,(
  ( ! [X405,X404] : (sP1('T_0',X404,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f53928,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP1('T_0',X61,'const.bool.IN_2'(s(sK12,X62),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X62),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f53918])).
fof(f53918,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | sP1('T_0',X61,'const.bool.IN_2'(s(sK12,X62),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X62),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X60))) ) | $spl0),
  inference(superposition,[],[f18840,f416])).
fof(f18840,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X38)) | sP1(X38,X39,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f117,f141,f140])).
fof(f141,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUBSET_INTER','thm.sat.sat_27_prove_ltb')).
fof(f53929,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP1('T_0',X58,'const.bool.IN_2'(s(sK12,X59),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X59),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f53917])).
fof(f53917,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP1('T_0',X58,'const.bool.IN_2'(s(sK12,X59),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X59),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X57))) ) | $spl0),
  inference(superposition,[],[f18840,f418])).
fof(f53948,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP1('T_0',X55,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f53916,f411])).
fof(f53916,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP1('T_0',X55,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'F_0') = s(bool,X54)) ) | $spl0),
  inference(superposition,[],[f18840,f442])).
fof(f53861,plain,(
  ( ! [X420,X418,X419] : (sP1('T_0',X418,'const.bool.IN_2'(s(sK12,X419),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X419),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X420))) ) | $spl0),
  inference(resolution,[],[f18840,f408])).
fof(f408,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) => p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_CLAUSES|split|0')).
fof(f53860,plain,(
  ( ! [X416,X415,X417] : (sP1('T_0',X415,'const.bool.IN_2'(s(sK12,X416),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X416),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X417))) ) | $spl0),
  inference(resolution,[],[f18840,f406])).
fof(f406,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.IMP_CLAUSES|split|2')).
fof(f53858,plain,(
  ( ! [X412,X410,X411] : (sP1('T_0',X410,'const.bool.IN_2'(s(sK12,X411),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X411),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X412))) ) | $spl0),
  inference(resolution,[],[f18840,f403])).
fof(f53857,plain,(
  ( ! [X408,X407,X409] : (sP1('T_0',X407,'const.bool.IN_2'(s(sK12,X408),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X408),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X409))) ) | $spl0),
  inference(resolution,[],[f18840,f402])).
fof(f53856,plain,(
  ( ! [X405,X404,X406] : (sP1('T_0',X404,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X405),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X406))) ) | $spl0),
  inference(resolution,[],[f18840,f399])).
fof(f53663,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18832,f36815])).
fof(f36815,plain,(
  ( ! [X21,X22] : (sP1('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X21,X22) | p(s(bool,X22))) ) | $spl94),
  inference(resolution,[],[f14077,f277])).
fof(f18832,plain,(
  ( ! [X14,X15,X13] : (~sP1('const.bool.IN_2'(s(sK12,X13),s(fun(sK12,bool),sK13)),X14,X15) | ~p(s(bool,X14)) | p(s(bool,'const.bool.IN_2'(s(sK12,X13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f141])).
fof(f53684,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f53662])).
fof(f53662,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18832,f36817])).
fof(f36817,plain,(
  ( ! [X26,X25] : (sP1('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X25,X26) | ~p(s(bool,X25))) ) | $spl94),
  inference(resolution,[],[f14077,f278])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f168])).
fof(f53563,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl269),
  inference(cnf_transformation,[],[f53563_D])).
fof(f53563_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f53661,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X16))) ) | $spl278),
  inference(cnf_transformation,[],[f53661_D])).
fof(f53661_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X16))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f53647,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl114 | $spl264)),
  inference(resolution,[],[f53338,f27497])).
fof(f53338,plain,(
  ( ! [X116,X115] : (sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X115,X116)) ) | $spl264),
  inference(resolution,[],[f53275,f334])).
fof(f53566,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl272),
  inference(cnf_transformation,[],[f53566_D])).
fof(f53566_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f53646,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl106 | $spl264)),
  inference(resolution,[],[f53338,f27389])).
fof(f53637,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8))) ) | $spl276),
  inference(cnf_transformation,[],[f53637_D])).
fof(f53637_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f53622,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl96 | $spl264)),
  inference(resolution,[],[f53337,f26605])).
fof(f53619,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53337,f12103])).
fof(f53567,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl274),
  inference(cnf_transformation,[],[f53567_D])).
fof(f53567_D,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f53545,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X31))) ) | ($spl0 | $spl264)),
  inference(resolution,[],[f53332,f48293])).
fof(f48293,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X2,X0,X3) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f48292,f334])).
fof(f48292,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP6(X0,X3,X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f48291,f328])).
fof(f48291,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f48190])).
fof(f48190,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0)) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl0),
  inference(resolution,[],[f36749,f320])).
fof(f36749,plain,(
  ( ! [X10,X8,X11,X9] : (~sP9(X11,X9,X8) | ~p(s(bool,X8)) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(superposition,[],[f18043,f337])).
fof(f53332,plain,(
  ( ! [X103,X104] : (sP8(X103,X104,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl264),
  inference(resolution,[],[f53275,f324])).
fof(f53564,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13))) ) | $spl270),
  inference(cnf_transformation,[],[f53564_D])).
fof(f53564_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f53541,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl96 | $spl264)),
  inference(resolution,[],[f53332,f26604])).
fof(f53540,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53332,f21264])).
fof(f53539,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53332,f21261])).
fof(f53537,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl264)),
  inference(resolution,[],[f53332,f12101])).
fof(f53522,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f53501])).
fof(f53501,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X2))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18831,f36815])).
fof(f18831,plain,(
  ( ! [X12,X10,X11] : (~sP1('const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),sK13)),X12,X11) | p(s(bool,X11)) | p(s(bool,'const.bool.IN_2'(s(sK12,X10),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f265])).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f163])).
fof(f53500,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18831,f36817])).
fof(f53433,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f53405])).
fof(f53405,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18830,f36814])).
fof(f36814,plain,(
  ( ! [X19,X20] : (sP1(X19,X20,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X20)) | p(s(bool,X19))) ) | $spl94),
  inference(resolution,[],[f14077,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f168])).
fof(f18830,plain,(
  ( ! [X8,X7,X9] : (~sP1(X8,X9,'const.bool.IN_2'(s(sK12,X7),s(fun(sK12,bool),sK13))) | p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(sK12,X7),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f265])).
fof(f53273,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),sK13)))) | p(s(bool,X30))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f43013,f36811])).
fof(f53272,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13)))) | p(s(bool,X28))) ) | ($spl0 | $spl2 | $spl41)),
  inference(resolution,[],[f43013,f36622])).
fof(f53271,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.bool.IN_2'(s(sK12,X25),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X25),s(fun(sK12,bool),sK13)))) | p(s(bool,X26))) ) | $spl0),
  inference(resolution,[],[f43013,f4810])).
fof(f53269,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X22))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f43013,f37928])).
fof(f53268,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.bool.IN_2'(s(sK12,X19),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X19),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X20))) ) | ($spl0 | $spl43)),
  inference(resolution,[],[f43013,f16478])).
fof(f53267,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X17),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X18))) ) | ($spl0 | $spl41)),
  inference(resolution,[],[f43013,f14347])).
fof(f53262,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13))))) ) | $spl0),
  inference(resolution,[],[f43013,f272])).
fof(f53275,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl264),
  inference(cnf_transformation,[],[f53275_D])).
fof(f53275_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f53261,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f43013,f36812])).
fof(f52558,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18935,f12282])).
fof(f12282,plain,(
  ( ! [X130] : (sP11('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X130)) ) | $spl5),
  inference(resolution,[],[f11960,f377])).
fof(f18935,plain,(
  ( ! [X202,X203] : (~sP11('const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),sK13)),X203) | p(s(bool,X203)) | p(s(bool,'const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f18898,f427])).
fof(f18898,plain,(
  ( ! [X202,X203] : (p(s(bool,'const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,'F_0')) | p(s(bool,X203)) | ~sP11('const.bool.IN_2'(s(sK12,X202),s(fun(sK12,bool),sK13)),X203)) ) | $spl0),
  inference(resolution,[],[f4804,f375])).
fof(f52541,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f18934,f12281])).
fof(f12281,plain,(
  ( ! [X129] : (sP11(X129,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl5),
  inference(resolution,[],[f11960,f376])).
fof(f18934,plain,(
  ( ! [X200,X201] : (~sP11(X201,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),sK13))) | p(s(bool,X201)) | p(s(bool,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f18897,f427])).
fof(f18897,plain,(
  ( ! [X200,X201] : (p(s(bool,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,'F_0')) | p(s(bool,X201)) | ~sP11(X201,'const.bool.IN_2'(s(sK12,X200),s(fun(sK12,bool),sK13)))) ) | $spl0),
  inference(resolution,[],[f4804,f375])).
fof(f51208,plain,(
  ( ! [X318] : (sP10('const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl260),
  inference(cnf_transformation,[],[f51208_D])).
fof(f51208_D,plain,(
  ( ! [X318] : (sP10('const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f51200,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP10('const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f51190])).
fof(f51190,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP10('const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X44))) ) | $spl0),
  inference(superposition,[],[f18892,f416])).
fof(f18892,plain,(
  ( ! [X191,X190] : (~p(s(bool,X191)) | sP10('const.bool.IN_2'(s(sK12,X190),s(fun(sK12,bool),sK13)),X191) | p(s(bool,'const.bool.IN_2'(s(sK12,X190),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f51201,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP10('const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f51189])).
fof(f51189,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP10('const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X42))) ) | $spl0),
  inference(superposition,[],[f18892,f418])).
fof(f51221,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP10('const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f51188,f411])).
fof(f51188,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP10('const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'F_0') = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f18892,f442])).
fof(f51027,plain,(
  ( ! [X337,X338] : (~sP11(X337,X338) | p(s(bool,X337)) | p(s(bool,X338))) ) | $spl256),
  inference(cnf_transformation,[],[f51027_D])).
fof(f51027_D,plain,(
  ( ! [X337,X338] : (~sP11(X337,X338) | p(s(bool,X337)) | p(s(bool,X338))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f51212,plain,(
  ( ! [X336] : (sP10('const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13)),'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl262),
  inference(cnf_transformation,[],[f51212_D])).
fof(f51212_D,plain,(
  ( ! [X336] : (sP10('const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13)),'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f51145,plain,(
  ( ! [X344,X343,X342] : (sP10('const.bool.IN_2'(s(sK12,X342),s(fun(sK12,bool),sK13)),'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X342),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X343)) | p(s(bool,X344)) | ~sP11(X343,X344)) ) | $spl0),
  inference(resolution,[],[f18892,f375])).
fof(f51143,plain,(
  ( ! [X337,X336,X338] : (sP10('const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13)),'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X337)) | p(s(bool,X338)) | ~sP11(X337,X338)) ) | $spl0),
  inference(resolution,[],[f18892,f379])).
fof(f51135,plain,(
  ( ! [X327,X328] : (sP10('const.bool.IN_2'(s(sK12,X327),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X327),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X328))) ) | $spl0),
  inference(resolution,[],[f18892,f408])).
fof(f51134,plain,(
  ( ! [X325,X326] : (sP10('const.bool.IN_2'(s(sK12,X325),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X325),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X326))) ) | $spl0),
  inference(resolution,[],[f18892,f406])).
fof(f51132,plain,(
  ( ! [X323,X322] : (sP10('const.bool.IN_2'(s(sK12,X322),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X322),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X323))) ) | $spl0),
  inference(resolution,[],[f18892,f403])).
fof(f51131,plain,(
  ( ! [X321,X320] : (sP10('const.bool.IN_2'(s(sK12,X320),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X320),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X321))) ) | $spl0),
  inference(resolution,[],[f18892,f402])).
fof(f51130,plain,(
  ( ! [X318,X319] : (sP10('const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13)),'T_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X319))) ) | $spl0),
  inference(resolution,[],[f18892,f399])).
fof(f51023,plain,(
  ( ! [X318] : (sP10('T_0','const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl254),
  inference(cnf_transformation,[],[f51023_D])).
fof(f51023_D,plain,(
  ( ! [X318] : (sP10('T_0','const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f51015,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP10('T_0','const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f51005])).
fof(f51005,plain,(
  ( ! [X45,X44] : (~p(s(bool,X44)) | sP10('T_0','const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X44))) ) | $spl0),
  inference(superposition,[],[f18891,f416])).
fof(f18891,plain,(
  ( ! [X189,X188] : (~p(s(bool,X189)) | sP10(X189,'const.bool.IN_2'(s(sK12,X188),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X188),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f361])).
fof(f361,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f206])).
fof(f51016,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP10('T_0','const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f51004])).
fof(f51004,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP10('T_0','const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X42))) ) | $spl0),
  inference(superposition,[],[f18891,f418])).
fof(f51037,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP10('T_0','const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f51003,f411])).
fof(f51003,plain,(
  ( ! [X41,X40] : (~p(s(bool,X40)) | sP10('T_0','const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'F_0') = s(bool,X40)) ) | $spl0),
  inference(superposition,[],[f18891,f442])).
fof(f51028,plain,(
  ( ! [X336] : (sP10('F_0','const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl258),
  inference(cnf_transformation,[],[f51028_D])).
fof(f51028_D,plain,(
  ( ! [X336] : (sP10('F_0','const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f50960,plain,(
  ( ! [X344,X343,X342] : (sP10('F_0','const.bool.IN_2'(s(sK12,X342),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X342),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X343)) | p(s(bool,X344)) | ~sP11(X343,X344)) ) | $spl0),
  inference(resolution,[],[f18891,f375])).
fof(f50958,plain,(
  ( ! [X337,X336,X338] : (sP10('F_0','const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X336),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X337)) | p(s(bool,X338)) | ~sP11(X337,X338)) ) | $spl0),
  inference(resolution,[],[f18891,f379])).
fof(f50950,plain,(
  ( ! [X327,X328] : (sP10('T_0','const.bool.IN_2'(s(sK12,X327),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X327),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X328))) ) | $spl0),
  inference(resolution,[],[f18891,f408])).
fof(f50949,plain,(
  ( ! [X325,X326] : (sP10('T_0','const.bool.IN_2'(s(sK12,X325),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X325),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X326))) ) | $spl0),
  inference(resolution,[],[f18891,f406])).
fof(f50947,plain,(
  ( ! [X323,X322] : (sP10('T_0','const.bool.IN_2'(s(sK12,X322),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X322),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X323))) ) | $spl0),
  inference(resolution,[],[f18891,f403])).
fof(f50946,plain,(
  ( ! [X321,X320] : (sP10('T_0','const.bool.IN_2'(s(sK12,X320),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X320),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X321))) ) | $spl0),
  inference(resolution,[],[f18891,f402])).
fof(f50945,plain,(
  ( ! [X318,X319] : (sP10('T_0','const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),sK13))) | p(s(bool,'const.bool.IN_2'(s(sK12,X318),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X319))) ) | $spl0),
  inference(resolution,[],[f18891,f399])).
fof(f50360,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f50347])).
fof(f50347,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18890,f36867])).
fof(f36867,plain,(
  ( ! [X123] : (sP10('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X123) | p(s(bool,X123))) ) | $spl94),
  inference(resolution,[],[f14077,f361])).
fof(f18890,plain,(
  ( ! [X187,X186] : (~sP10('const.bool.IN_2'(s(sK12,X186),s(fun(sK12,bool),sK13)),X187) | p(s(bool,X187)) | p(s(bool,'const.bool.IN_2'(s(sK12,X186),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f50340,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(duplicate_literal_removal,[],[f50327])).
fof(f50327,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | p(s(bool,X0))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f18889,f36868])).
fof(f36868,plain,(
  ( ! [X124] : (sP10(X124,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X124))) ) | $spl94),
  inference(resolution,[],[f14077,f364])).
fof(f18889,plain,(
  ( ! [X185,X184] : (~sP10(X185,'const.bool.IN_2'(s(sK12,X184),s(fun(sK12,bool),sK13))) | p(s(bool,X185)) | p(s(bool,'const.bool.IN_2'(s(sK12,X184),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f4804,f359])).
fof(f45130,plain,(
  ( ! [X15,X16] : (p(s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),sK13)))) ) | $spl244),
  inference(cnf_transformation,[],[f45130_D])).
fof(f45130_D,plain,(
  ( ! [X15,X16] : (p(s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),sK13)))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f49948,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),sK13))) | ~p(s(bool,X36)) | p(s(bool,'const.bool.IN_2'(s(sK12,X37),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | $spl0),
  inference(resolution,[],[f47004,f18879])).
fof(f18879,plain,(
  ( ! [X156,X154,X155] : (p(s(bool,'const.bool.IN_2'(s(sK12,X154),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP8(X155,'const.bool.IN_2'(s(sK12,X154),s(fun(sK12,bool),sK13)),X156)) ) | $spl0),
  inference(resolution,[],[f4804,f325])).
fof(f47004,plain,(
  ( ! [X6,X4,X7,X5] : (~sP8(X6,X7,X5) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~p(s(bool,X4))) ) | $spl0),
  inference(subsumption_resolution,[],[f47003,f332])).
fof(f47003,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~sP8(X6,X7,X5) | ~sP6(X7,X5,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f46999,f328])).
fof(f46999,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~sP8(X6,X7,X5) | ~sP7(X5,X7,X6) | ~sP6(X7,X5,X6)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f46909])).
fof(f46909,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | p(s(bool,X5)) | ~sP8(X6,X7,X5) | ~sP7(X5,X7,X6) | ~sP6(X7,X5,X6)) ) | $spl0),
  inference(resolution,[],[f36727,f321])).
fof(f36727,plain,(
  ( ! [X14,X12,X15,X13] : (~sP9(X15,X12,X13) | ~p(s(bool,X14)) | p(s(bool,X12)) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f18043,f337])).
fof(f49599,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36858,f18882])).
fof(f18882,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,'const.bool.IN_2'(s(sK12,X163),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP7(X164,'const.bool.IN_2'(s(sK12,X163),s(fun(sK12,bool),sK13)),X165)) ) | $spl0),
  inference(resolution,[],[f4804,f329])).
fof(f49550,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36857,f18883])).
fof(f18883,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,'const.bool.IN_2'(s(sK12,X166),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP7('const.bool.IN_2'(s(sK12,X166),s(fun(sK12,bool),sK13)),X167,X168)) ) | $spl0),
  inference(resolution,[],[f4804,f330])).
fof(f49516,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36855,f18879])).
fof(f36855,plain,(
  ( ! [X101,X102] : (~sP8(X101,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X102) | p(s(bool,X102)) | ~p(s(bool,X101))) ) | $spl94),
  inference(resolution,[],[f14077,f323])).
fof(f49491,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36854,f18880])).
fof(f18880,plain,(
  ( ! [X158,X159,X157] : (p(s(bool,'const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP8('const.bool.IN_2'(s(sK12,X157),s(fun(sK12,bool),sK13)),X158,X159)) ) | $spl0),
  inference(resolution,[],[f4804,f326])).
fof(f48721,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36842,f18866])).
fof(f18866,plain,(
  ( ! [X116,X117,X115] : (p(s(bool,'const.bool.IN_2'(s(sK12,X115),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP4(X116,'const.bool.IN_2'(s(sK12,X115),s(fun(sK12,bool),sK13)),X117)) ) | $spl0),
  inference(resolution,[],[f4804,f308])).
fof(f36842,plain,(
  ( ! [X76,X75] : (~sP4(X75,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X76) | p(s(bool,X76)) | ~p(s(bool,X75))) ) | $spl94),
  inference(resolution,[],[f14077,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f181])).
fof(f48696,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36841,f18867])).
fof(f18867,plain,(
  ( ! [X118,X120,X119] : (p(s(bool,'const.bool.IN_2'(s(sK12,X118),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP4('const.bool.IN_2'(s(sK12,X118),s(fun(sK12,bool),sK13)),X119,X120)) ) | $spl0),
  inference(resolution,[],[f4804,f309])).
fof(f36841,plain,(
  ( ! [X74,X73] : (~sP4('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X73,X74) | p(s(bool,X74)) | ~p(s(bool,X73))) ) | $spl94),
  inference(resolution,[],[f14077,f306])).
fof(f48640,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36825,f18853])).
fof(f18853,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,'const.bool.IN_2'(s(sK12,X76),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP2('const.bool.IN_2'(s(sK12,X76),s(fun(sK12,bool),sK13)),X77,X78)) ) | $spl0),
  inference(resolution,[],[f4804,f292])).
fof(f36825,plain,(
  ( ! [X41,X42] : (~sP2('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X41,X42) | p(s(bool,X42)) | p(s(bool,X41))) ) | $spl94),
  inference(resolution,[],[f14077,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f174])).
fof(f48471,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36812,f18836])).
fof(f48444,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))) ) | ($spl0 | $spl94)),
  inference(resolution,[],[f36811,f18837])).
fof(f18837,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | sP0('const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),sK13)),X29,X30)) ) | $spl0),
  inference(resolution,[],[f4804,f275])).
fof(f47672,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl250),
  inference(cnf_transformation,[],[f47672_D])).
fof(f47672_D,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f47671,plain,(
  ( ! [X109,X107,X108] : (s(bool,'F_0') != s(bool,X107) | ~p(s(bool,X108)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X109))) ) | ($spl0 | $spl21)),
  inference(subsumption_resolution,[],[f47629,f411])).
fof(f47629,plain,(
  ( ! [X109,X107,X108] : (s(bool,'F_0') != s(bool,X107) | ~p(s(bool,X108)) | p(s(bool,X107)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | p(s(bool,X109))) ) | ($spl0 | $spl21)),
  inference(superposition,[],[f36748,f37913])).
fof(f36748,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X4)) | p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl0),
  inference(superposition,[],[f18043,f335])).
fof(f47665,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(subsumption_resolution,[],[f47617,f411])).
fof(f47617,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(superposition,[],[f36748,f442])).
fof(f47115,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X108)) ) | $spl248),
  inference(cnf_transformation,[],[f47115_D])).
fof(f47115_D,plain,(
  ( ! [X109,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X108)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f47114,plain,(
  ( ! [X109,X107,X108] : (s(bool,'F_0') != s(bool,X107) | ~p(s(bool,X108)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X108) | p(s(bool,X109))) ) | ($spl0 | $spl21)),
  inference(subsumption_resolution,[],[f47071,f411])).
fof(f47071,plain,(
  ( ! [X109,X107,X108] : (s(bool,'F_0') != s(bool,X107) | ~p(s(bool,X108)) | p(s(bool,X107)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X108) | p(s(bool,X109))) ) | ($spl0 | $spl21)),
  inference(superposition,[],[f36747,f37913])).
fof(f36747,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X0)) | p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f18043,f335])).
fof(f47107,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(subsumption_resolution,[],[f47059,f411])).
fof(f47059,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(superposition,[],[f36747,f442])).
fof(f30064,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl134),
  inference(cnf_transformation,[],[f30064_D])).
fof(f30064_D,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f46367,plain,(
  ( ! [X105,X106,X104] : (~p(s(bool,X104)) | p(s(bool,X105)) | sP9(X106,X105,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X106))) ) | ($spl0 | $spl94)),
  inference(subsumption_resolution,[],[f46324,f416])).
fof(f46324,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | ~p(s(bool,X104)) | p(s(bool,X105)) | sP9(X106,X105,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X106))) ) | ($spl0 | $spl94)),
  inference(superposition,[],[f36725,f36884])).
fof(f30056,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl130),
  inference(cnf_transformation,[],[f30056_D])).
fof(f30056_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f46361,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(subsumption_resolution,[],[f46312,f416])).
fof(f46312,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(superposition,[],[f36725,f442])).
fof(f45677,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) | $spl246),
  inference(cnf_transformation,[],[f45677_D])).
fof(f45677_D,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f45678,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | sP9(X63,X61,X62) | p(s(bool,X63))) ) | $spl0),
  inference(subsumption_resolution,[],[f45647,f411])).
fof(f45647,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | sP9(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(superposition,[],[f36724,f442])).
fof(f45662,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | sP9(X59,X57,X58) | p(s(bool,X59))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f45646])).
fof(f45646,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | sP9(X59,X57,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl0),
  inference(superposition,[],[f36724,f418])).
fof(f45663,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | sP9(X55,X53,X54) | p(s(bool,X55))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f45645])).
fof(f45645,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | sP9(X55,X53,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl0),
  inference(superposition,[],[f36724,f416])).
fof(f28522,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) | $spl124),
  inference(cnf_transformation,[],[f28522_D])).
fof(f28522_D,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | p(s(bool,X109)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f45676,plain,(
  ( ! [X105,X106,X104] : (~p(s(bool,X104)) | p(s(bool,X105)) | sP9(X106,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X105) | p(s(bool,X106))) ) | ($spl0 | $spl94)),
  inference(subsumption_resolution,[],[f45633,f416])).
fof(f45633,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') != s(bool,X104) | ~p(s(bool,X104)) | p(s(bool,X105)) | sP9(X106,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X105) | p(s(bool,X106))) ) | ($spl0 | $spl94)),
  inference(superposition,[],[f36724,f36884])).
fof(f28512,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl116),
  inference(cnf_transformation,[],[f28512_D])).
fof(f28512_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f45670,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(subsumption_resolution,[],[f45621,f416])).
fof(f45621,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl0),
  inference(superposition,[],[f36724,f442])).
fof(f45103,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),sK13)))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f18883,f13485])).
fof(f44730,plain,(
  ( ! [X48,X46] : (p(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'T_0') != s(bool,X48) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),sK13)))) ) | $spl242),
  inference(cnf_transformation,[],[f44730_D])).
fof(f44730_D,plain,(
  ( ! [X48,X46] : (p(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'T_0') != s(bool,X48) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),sK13)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f44702,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X47)) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X48)) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f18879,f21063])).
fof(f44728,plain,(
  ( ! [X27] : (p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13)))) ) | $spl238),
  inference(cnf_transformation,[],[f44728_D])).
fof(f44728_D,plain,(
  ( ! [X27] : (p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f44700,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X41),s(fun(sK12,bool),sK13))) | ~p(s(bool,X42))) ) | $spl0),
  inference(resolution,[],[f18879,f12706])).
fof(f12706,plain,(
  ( ! [X8,X9] : (~sP8(X8,X9,'F_0') | s(bool,'T_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl0),
  inference(subsumption_resolution,[],[f12705,f332])).
fof(f12705,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | ~sP8(X8,X9,'F_0') | ~sP6(X9,'F_0',X8)) ) | $spl0),
  inference(subsumption_resolution,[],[f12704,f397])).
fof(f12704,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'F_0')) | ~sP8(X8,X9,'F_0') | ~sP6(X9,'F_0',X8)) ) | $spl0),
  inference(subsumption_resolution,[],[f12688,f330])).
fof(f12688,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'F_0')) | ~sP8(X8,X9,'F_0') | ~sP7('F_0',X9,X8) | ~sP6(X9,'F_0',X8)) ) | $spl0),
  inference(resolution,[],[f11947,f321])).
fof(f11947,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f10443,f337])).
fof(f44729,plain,(
  ( ! [X30] : (p(s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),sK13)))) ) | $spl240),
  inference(cnf_transformation,[],[f44729_D])).
fof(f44729_D,plain,(
  ( ! [X30] : (p(s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),sK13)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f44695,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X31)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),sK13))) | p(s(bool,X32))) ) | ($spl0 | $spl152)),
  inference(resolution,[],[f18879,f40490])).
fof(f44694,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),sK13)))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f18879,f21264])).
fof(f41592,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl94 | $spl224)),
  inference(subsumption_resolution,[],[f41591,f416])).
fof(f41591,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl94 | $spl224)),
  inference(forward_demodulation,[],[f41559,f36884])).
fof(f41559,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl224),
  inference(resolution,[],[f36995,f338])).
fof(f39798,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,X10))) ) | ($spl2 | $spl41 | $spl214)),
  inference(forward_demodulation,[],[f39796,f36695])).
fof(f36695,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ($spl2 | $spl41)),
  inference(resolution,[],[f36441,f416])).
fof(f39796,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl214),
  inference(duplicate_literal_removal,[],[f39773])).
fof(f39773,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl214),
  inference(resolution,[],[f36592,f338])).
fof(f39737,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl2 | $spl41 | $spl180)),
  inference(subsumption_resolution,[],[f39736,f416])).
fof(f39736,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl2 | $spl41 | $spl180)),
  inference(forward_demodulation,[],[f39709,f36695])).
fof(f39709,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl180),
  inference(resolution,[],[f36522,f338])).
fof(f39337,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,X10))) ) | ($spl2 | $spl41 | $spl220)),
  inference(forward_demodulation,[],[f39335,f36695])).
fof(f39335,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl220),
  inference(duplicate_literal_removal,[],[f39309])).
fof(f39309,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X10)) | s(bool,X11) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | $spl220),
  inference(resolution,[],[f36614,f338])).
fof(f38192,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl236),
  inference(cnf_transformation,[],[f38192_D])).
fof(f38192_D,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f38184,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl21)),
  inference(resolution,[],[f37976,f13485])).
fof(f37976,plain,(
  ( ! [X109,X110] : (sP7('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X109,X110)) ) | $spl21),
  inference(resolution,[],[f37914,f330])).
fof(f38107,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl234),
  inference(cnf_transformation,[],[f38107_D])).
fof(f38107_D,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f38101,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,X30)) ) | ($spl8 | $spl21)),
  inference(resolution,[],[f37972,f21063])).
fof(f37972,plain,(
  ( ! [X101,X102] : (sP8(X101,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X102)) ) | $spl21),
  inference(resolution,[],[f37914,f325])).
fof(f6302,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl22),
  inference(cnf_transformation,[],[f6302_D])).
fof(f6302_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f38100,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X27) | sP9(X28,'F_0','const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl8 | $spl21)),
  inference(resolution,[],[f37972,f20975])).
fof(f37918,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl233),
  inference(cnf_transformation,[],[f37918_D])).
fof(f37918_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f37919,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | $spl21),
  inference(subsumption_resolution,[],[f37906,f416])).
fof(f37906,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))))) ) | $spl21),
  inference(superposition,[],[f6195,f391])).
fof(f37916,plain,(
  ( ! [X10] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X10))) ) | $spl21),
  inference(subsumption_resolution,[],[f37905,f416])).
fof(f37905,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | ~p(s(bool,X10))) ) | $spl21),
  inference(superposition,[],[f6195,f391])).
fof(f36269,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl163),
  inference(cnf_transformation,[],[f36269_D])).
fof(f36269_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f37068,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16))) ) | $spl230),
  inference(cnf_transformation,[],[f37068_D])).
fof(f37068_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f37058,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl114)),
  inference(resolution,[],[f36862,f27497])).
fof(f36862,plain,(
  ( ! [X116,X115] : (sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X115,X116)) ) | $spl94),
  inference(resolution,[],[f14077,f334])).
fof(f36997,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl226),
  inference(cnf_transformation,[],[f36997_D])).
fof(f36997_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f37057,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl106)),
  inference(resolution,[],[f36862,f27389])).
fof(f37048,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8))) ) | $spl228),
  inference(cnf_transformation,[],[f37048_D])).
fof(f37048_D,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f37038,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl96)),
  inference(resolution,[],[f36861,f26605])).
fof(f37035,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36861,f12103])).
fof(f36995,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13))) ) | $spl224),
  inference(cnf_transformation,[],[f36995_D])).
fof(f36995_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f36987,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl94 | $spl96)),
  inference(resolution,[],[f36856,f26604])).
fof(f36856,plain,(
  ( ! [X103,X104] : (sP8(X103,X104,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl94),
  inference(resolution,[],[f14077,f324])).
fof(f36986,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36856,f21264])).
fof(f36985,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36856,f21261])).
fof(f36983,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl8 | $spl94)),
  inference(resolution,[],[f36856,f12101])).
fof(f36603,plain,(
  ( ! [X19,X18] : (sP9(X18,X19,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl114)),
  inference(resolution,[],[f530,f27497])).
fof(f36602,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl106)),
  inference(resolution,[],[f530,f27389])).
fof(f36580,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X15) | p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f529,f26605])).
fof(f36577,plain,(
  ( ! [X8,X9] : (sP9(X8,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f529,f12103])).
fof(f36510,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f524,f26604])).
fof(f36509,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f524,f21264])).
fof(f36508,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f524,f21261])).
fof(f36506,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f524,f12101])).
fof(f14077,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl94),
  inference(cnf_transformation,[],[f14077_D])).
fof(f14077_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f36395,plain,(
  ( ! [X64] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X64))) ) | ($spl0 | $spl2 | $spl5 | $spl8)),
  inference(duplicate_literal_removal,[],[f36392])).
fof(f36392,plain,(
  ( ! [X64] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X64)) | p(s(bool,X64))) ) | ($spl0 | $spl2 | $spl5 | $spl8)),
  inference(resolution,[],[f1663,f12807])).
fof(f36422,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | ($spl0 | $spl2 | $spl5 | $spl8)),
  inference(duplicate_literal_removal,[],[f36356])).
fof(f36356,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X9))) ) | ($spl0 | $spl2 | $spl5 | $spl8)),
  inference(resolution,[],[f1663,f12829])).
fof(f12829,plain,(
  ( ! [X11] : (sP9(X11,'F_0','const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X11))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12270,f12576])).
fof(f12576,plain,(
  ( ! [X6,X7] : (~sP6(X7,'F_0',X6) | p(s(bool,X6)) | sP9(X6,'F_0',X7)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12573,f326])).
fof(f12573,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f12564,f322])).
fof(f12270,plain,(
  ( ! [X116,X115] : (sP6('const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X115,X116)) ) | $spl5),
  inference(resolution,[],[f11960,f334])).
fof(f36324,plain,(
  $false | ($spl0 | $spl90)),
  inference(trivial_inequality_removal,[],[f36320])).
fof(f36320,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl90)),
  inference(superposition,[],[f14054,f4883])).
fof(f36323,plain,(
  $false | $spl90),
  inference(equality_resolution,[],[f14054])).
fof(f13605,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X3) != s(bool,X5)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13582])).
fof(f13582,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,X5)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl8),
  inference(resolution,[],[f13521,f338])).
fof(f13521,plain,(
  ( ! [X10,X11,X9] : (sP9(X9,X10,X11) | p(s(bool,X10)) | p(s(bool,X9)) | ~p(s(bool,X11))) ) | $spl8),
  inference(resolution,[],[f13415,f334])).
fof(f13415,plain,(
  ( ! [X6,X8,X7] : (~sP6(X8,X6,X7) | sP9(X7,X6,X8) | p(s(bool,X6)) | p(s(bool,X7))) ) | $spl8),
  inference(resolution,[],[f12102,f326])).
fof(f36282,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl8),
  inference(subsumption_resolution,[],[f16107,f415])).
fof(f16107,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl8),
  inference(superposition,[],[f12161,f391])).
fof(f36281,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl8),
  inference(subsumption_resolution,[],[f16108,f415])).
fof(f16108,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl8),
  inference(superposition,[],[f12161,f391])).
fof(f36279,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X16))) ) | $spl24),
  inference(subsumption_resolution,[],[f17380,f411])).
fof(f17380,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl24),
  inference(superposition,[],[f11865,f388])).
fof(f36278,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19))) ) | $spl24),
  inference(subsumption_resolution,[],[f17381,f411])).
fof(f17381,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl24),
  inference(superposition,[],[f11865,f388])).
fof(f19257,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl0 | $spl5 | $spl8)),
  inference(duplicate_literal_removal,[],[f19245])).
fof(f19245,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12829,f338])).
fof(f19611,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | ($spl0 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f19601,f412])).
fof(f19601,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | s(bool,'F_0') != s(bool,X3)) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12807,f338])).
fof(f20090,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | p(s(bool,X11))) ) | $spl8),
  inference(resolution,[],[f13485,f330])).
fof(f36263,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl8),
  inference(subsumption_resolution,[],[f20147,f416])).
fof(f20147,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl8),
  inference(superposition,[],[f12160,f391])).
fof(f36262,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl8),
  inference(subsumption_resolution,[],[f20148,f416])).
fof(f20148,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl8),
  inference(superposition,[],[f12160,f391])).
fof(f36228,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X21)) | p(s(bool,X20))) ) | $spl8),
  inference(subsumption_resolution,[],[f22468,f388])).
fof(f22468,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X21) != s(bool,X22) | s(bool,X20) != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X21)) | p(s(bool,X20))) ) | $spl8),
  inference(superposition,[],[f16088,f388])).
fof(f16088,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | ~p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f12161,f418])).
fof(f36227,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl8),
  inference(subsumption_resolution,[],[f22469,f388])).
fof(f22469,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl8),
  inference(superposition,[],[f16088,f388])).
fof(f36226,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl8),
  inference(subsumption_resolution,[],[f22470,f391])).
fof(f22470,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,X28) != s(bool,X31) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f16088,f391])).
fof(f36225,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl8),
  inference(subsumption_resolution,[],[f22471,f391])).
fof(f22471,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,X32) != s(bool,X35) | p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl8),
  inference(superposition,[],[f16088,f391])).
fof(f36224,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X38) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl8),
  inference(subsumption_resolution,[],[f22472,f412])).
fof(f22472,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl8),
  inference(superposition,[],[f16088,f412])).
fof(f36223,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,X41) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl8),
  inference(subsumption_resolution,[],[f22473,f412])).
fof(f22473,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X39) != s(bool,X41) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl8),
  inference(superposition,[],[f16088,f414])).
fof(f36221,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X53) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl8),
  inference(subsumption_resolution,[],[f22477,f412])).
fof(f22477,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,X52) | s(bool,X51) != s(bool,X53) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl8),
  inference(superposition,[],[f16088,f442])).
fof(f36220,plain,(
  ( ! [X54,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'T_0') != s(bool,X56) | p(s(bool,X55)) | s(bool,'F_0') = s(bool,X54)) ) | $spl8),
  inference(subsumption_resolution,[],[f22478,f415])).
fof(f22478,plain,(
  ( ! [X54,X56,X55] : (s(bool,X54) != s(bool,X55) | s(bool,'T_0') != s(bool,X56) | p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,'F_0') = s(bool,X54)) ) | $spl8),
  inference(superposition,[],[f16088,f442])).
fof(f36219,plain,(
  ( ! [X59,X57,X58] : (s(bool,X57) != s(bool,X58) | s(bool,'T_0') != s(bool,X59) | p(s(bool,X58)) | ~p(s(bool,X57))) ) | $spl8),
  inference(subsumption_resolution,[],[f22479,f415])).
fof(f22479,plain,(
  ( ! [X59,X57,X58] : (s(bool,X57) != s(bool,X58) | s(bool,'T_0') != s(bool,X59) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~p(s(bool,X57))) ) | $spl8),
  inference(superposition,[],[f16088,f418])).
fof(f36217,plain,(
  ( ! [X61,X62,X60] : (s(bool,X60) != s(bool,X61) | s(bool,'T_0') != s(bool,X62) | p(s(bool,X61)) | ~p(s(bool,X60))) ) | $spl8),
  inference(subsumption_resolution,[],[f22480,f415])).
fof(f22480,plain,(
  ( ! [X61,X62,X60] : (s(bool,X60) != s(bool,X61) | s(bool,'T_0') != s(bool,X62) | p(s(bool,X61)) | ~p(s(bool,X62)) | ~p(s(bool,X60))) ) | $spl8),
  inference(superposition,[],[f16088,f416])).
fof(f26568,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f26506])).
fof(f26506,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl96),
  inference(resolution,[],[f26344,f338])).
fof(f26591,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl8 | $spl96)),
  inference(resolution,[],[f26570,f13485])).
fof(f27350,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl106),
  inference(duplicate_literal_removal,[],[f27289])).
fof(f27289,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl106),
  inference(resolution,[],[f26912,f338])).
fof(f29457,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl116),
  inference(duplicate_literal_removal,[],[f29389])).
fof(f29389,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl116),
  inference(resolution,[],[f28512,f338])).
fof(f29762,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl106 | $spl116)),
  inference(duplicate_literal_removal,[],[f29700])).
fof(f29700,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl106 | $spl116)),
  inference(resolution,[],[f29481,f338])).
fof(f29481,plain,(
  ( ! [X26,X24,X25] : (sP9(X25,X24,X26) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | s(bool,'T_0') = s(bool,X26)) ) | ($spl106 | $spl116)),
  inference(resolution,[],[f29460,f27389])).
fof(f36176,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl128),
  inference(subsumption_resolution,[],[f29927,f412])).
fof(f29927,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl128),
  inference(superposition,[],[f29848,f388])).
fof(f36175,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl128),
  inference(subsumption_resolution,[],[f29928,f412])).
fof(f29928,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl128),
  inference(superposition,[],[f29848,f388])).
fof(f30634,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X15) != s(bool,X17)) ) | $spl130),
  inference(duplicate_literal_removal,[],[f30566])).
fof(f30566,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl130),
  inference(resolution,[],[f30056,f338])).
fof(f31014,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl96 | $spl130)),
  inference(duplicate_literal_removal,[],[f30952])).
fof(f30952,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl96 | $spl130)),
  inference(resolution,[],[f30654,f338])).
fof(f30654,plain,(
  ( ! [X23,X21,X22] : (sP9(X22,X23,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X22)) | s(bool,'T_0') = s(bool,X23)) ) | ($spl96 | $spl130)),
  inference(resolution,[],[f30636,f26605])).
fof(f31390,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13)) ) | $spl136),
  inference(duplicate_literal_removal,[],[f31290])).
fof(f31290,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,X14))) ) | $spl136),
  inference(resolution,[],[f31176,f336])).
fof(f31975,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13)) ) | $spl146),
  inference(duplicate_literal_removal,[],[f31873])).
fof(f31873,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13) | ~p(s(bool,X14))) ) | $spl146),
  inference(resolution,[],[f31531,f336])).
fof(f31526,plain,(
  ( ! [X37,X38,X34] : (~sP9(X38,X37,X34) | ~p(s(bool,X34)) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl144),
  inference(cnf_transformation,[],[f31526_D])).
fof(f31526_D,plain,(
  ( ! [X37,X38,X34] : (~sP9(X38,X37,X34) | ~p(s(bool,X34)) | p(s(bool,X37)) | ~p(s(bool,X38))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f32085,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~sP9(X31,X30,X28) | ~p(s(bool,X31)) | ~p(s(bool,X28))) ) | $spl152),
  inference(subsumption_resolution,[],[f32050,f415])).
fof(f32050,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,X30)) | ~sP9(X31,X30,X28) | ~p(s(bool,X31)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl152),
  inference(superposition,[],[f31538,f391])).
fof(f32086,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | p(s(bool,X34)) | ~sP9(X35,X34,X32) | ~p(s(bool,X35)) | ~p(s(bool,X32))) ) | $spl152),
  inference(subsumption_resolution,[],[f32051,f415])).
fof(f32051,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | p(s(bool,X34)) | ~sP9(X35,X34,X32) | ~p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl152),
  inference(superposition,[],[f31538,f391])).
fof(f36125,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f36124,f5079])).
fof(f5079,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),X0))))))) ) | $spl3),
  inference(resolution,[],[f4679,f437])).
fof(f4679,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) | $spl3),
  inference(resolution,[],[f470,f436])).
fof(f36124,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f36101,f402])).
fof(f36101,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ($spl0 | $spl3)),
  inference(superposition,[],[f4967,f4883])).
fof(f4967,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),X2))))) ) | $spl3),
  inference(resolution,[],[f4678,f434])).
fof(f4678,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl3),
  inference(resolution,[],[f470,f435])).
fof(f31538,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~sP9(X55,X54,X53) | ~p(s(bool,X55))) ) | $spl152),
  inference(cnf_transformation,[],[f31538_D])).
fof(f31538_D,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~sP9(X55,X54,X53) | ~p(s(bool,X55))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f31539,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X62,X61)) ) | $spl3),
  inference(subsumption_resolution,[],[f31505,f411])).
fof(f31505,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X62,X61) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12935,f442])).
fof(f12935,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,X14) | ~p(s(bool,X14)) | p(s(bool,X12)) | ~p(s(bool,X15)) | ~sP9(X15,X12,X13)) ) | $spl3),
  inference(superposition,[],[f11899,f337])).
fof(f11899,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl3),
  inference(subsumption_resolution,[],[f11898,f403])).
fof(f11898,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl3),
  inference(forward_demodulation,[],[f10271,f5046])).
fof(f5046,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl3),
  inference(resolution,[],[f4678,f416])).
fof(f10271,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl3),
  inference(resolution,[],[f5025,f319])).
fof(f5025,plain,(
  ( ! [X118,X117] : (sP9('const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)),X117,X118) | s(bool,X117) != s(bool,X118)) ) | $spl3),
  inference(resolution,[],[f4678,f336])).
fof(f31521,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X58,X57)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f31504])).
fof(f31504,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X58,X57) | ~p(s(bool,X56))) ) | $spl3),
  inference(superposition,[],[f12935,f418])).
fof(f31522,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X54,X53)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f31503])).
fof(f31503,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X54,X53) | ~p(s(bool,X52))) ) | $spl3),
  inference(superposition,[],[f12935,f416])).
fof(f31537,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl150),
  inference(cnf_transformation,[],[f31537_D])).
fof(f31537_D,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f31536,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f31491,f416])).
fof(f31491,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,X107) | ~p(s(bool,X107)) | p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f12935,f14166])).
fof(f14166,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl3 | $spl5)),
  inference(resolution,[],[f14073,f416])).
fof(f14073,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl3 | $spl5)),
  inference(resolution,[],[f12281,f4783])).
fof(f4783,plain,(
  ( ! [X127] : (~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X127) | p(s(bool,X127))) ) | $spl3),
  inference(subsumption_resolution,[],[f4748,f427])).
fof(f4748,plain,(
  ( ! [X127] : (p(s(bool,'F_0')) | p(s(bool,X127)) | ~sP11('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)),X127)) ) | $spl3),
  inference(resolution,[],[f470,f375])).
fof(f31533,plain,(
  ( ! [X70,X71] : (~sP9(X71,X70,'T_0') | ~p(s(bool,X71)) | p(s(bool,X70))) ) | $spl148),
  inference(cnf_transformation,[],[f31533_D])).
fof(f31533_D,plain,(
  ( ! [X70,X71] : (~sP9(X71,X70,'T_0') | ~p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f31535,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP9(X79,X78,'T_0') | ~p(s(bool,X76))) ) | $spl3),
  inference(subsumption_resolution,[],[f31482,f11899])).
fof(f31482,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP9(X79,X78,'T_0') | ~p(s(bool,X76))) ) | $spl3),
  inference(superposition,[],[f12935,f416])).
fof(f31534,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP9(X75,X74,'T_0') | ~p(s(bool,X72))) ) | $spl3),
  inference(subsumption_resolution,[],[f31481,f11899])).
fof(f31481,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP9(X75,X74,'T_0') | ~p(s(bool,X72))) ) | $spl3),
  inference(superposition,[],[f12935,f418])).
fof(f31532,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~sP9(X71,X70,'T_0') | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(subsumption_resolution,[],[f31480,f12967])).
fof(f12967,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl3),
  inference(subsumption_resolution,[],[f12922,f399])).
fof(f12922,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | ~p(s(bool,'T_0')) | p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f11899,f442])).
fof(f31480,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~sP9(X71,X70,'T_0') | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f12935,f442])).
fof(f31531,plain,(
  ( ! [X62,X60,X63] : (~sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl146),
  inference(cnf_transformation,[],[f31531_D])).
fof(f31531_D,plain,(
  ( ! [X62,X60,X63] : (~sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X62)) | ~p(s(bool,X63))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f31530,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(subsumption_resolution,[],[f31478,f416])).
fof(f31478,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12935,f442])).
fof(f31529,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X58,X56) | ~p(s(bool,X56))) ) | $spl3),
  inference(subsumption_resolution,[],[f31477,f416])).
fof(f31477,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X58,X56) | ~p(s(bool,X56))) ) | $spl3),
  inference(superposition,[],[f12935,f418])).
fof(f31528,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X54,X52) | ~p(s(bool,X52))) ) | $spl3),
  inference(subsumption_resolution,[],[f31476,f416])).
fof(f31476,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X54,X52) | ~p(s(bool,X52))) ) | $spl3),
  inference(superposition,[],[f12935,f416])).
fof(f31527,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP9(X43,X42,X39) | ~p(s(bool,X40)) | ~p(s(bool,X39))) ) | $spl3),
  inference(subsumption_resolution,[],[f31473,f11899])).
fof(f31473,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP9(X43,X42,X39) | ~p(s(bool,X40)) | ~p(s(bool,X39))) ) | $spl3),
  inference(superposition,[],[f12935,f391])).
fof(f31525,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP9(X38,X37,X34) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl3),
  inference(subsumption_resolution,[],[f31472,f11899])).
fof(f31472,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP9(X38,X37,X34) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl3),
  inference(superposition,[],[f12935,f391])).
fof(f31183,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~sP9(X55,X53,X54) | ~p(s(bool,X55)) | p(s(bool,X54))) ) | $spl142),
  inference(cnf_transformation,[],[f31183_D])).
fof(f31183_D,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~sP9(X55,X53,X54) | ~p(s(bool,X55)) | p(s(bool,X54))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f31184,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X61,X62)) ) | $spl3),
  inference(subsumption_resolution,[],[f31151,f411])).
fof(f31151,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X61,X62) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12934,f442])).
fof(f12934,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X8)) | ~p(s(bool,X11)) | ~sP9(X11,X9,X8)) ) | $spl3),
  inference(superposition,[],[f11899,f337])).
fof(f31167,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X57,X58)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f31150])).
fof(f31150,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | ~p(s(bool,X59)) | ~sP9(X59,X57,X58) | ~p(s(bool,X56))) ) | $spl3),
  inference(superposition,[],[f12934,f418])).
fof(f31168,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X53,X54)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f31149])).
fof(f31149,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | ~p(s(bool,X55)) | ~sP9(X55,X53,X54) | ~p(s(bool,X52))) ) | $spl3),
  inference(superposition,[],[f12934,f416])).
fof(f31182,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) | $spl140),
  inference(cnf_transformation,[],[f31182_D])).
fof(f31182_D,plain,(
  ( ! [X109,X108] : (p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f31181,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f31137,f416])).
fof(f31137,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,X107) | ~p(s(bool,X107)) | p(s(bool,X108)) | ~p(s(bool,X109)) | ~sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108)) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f12934,f14166])).
fof(f31178,plain,(
  ( ! [X70,X71] : (~sP9(X71,'T_0',X70) | ~p(s(bool,X71)) | p(s(bool,X70))) ) | $spl138),
  inference(cnf_transformation,[],[f31178_D])).
fof(f31178_D,plain,(
  ( ! [X70,X71] : (~sP9(X71,'T_0',X70) | ~p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f31180,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP9(X79,'T_0',X78) | ~p(s(bool,X76))) ) | $spl3),
  inference(subsumption_resolution,[],[f31128,f11899])).
fof(f31128,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP9(X79,'T_0',X78) | ~p(s(bool,X76))) ) | $spl3),
  inference(superposition,[],[f12934,f416])).
fof(f31179,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP9(X75,'T_0',X74) | ~p(s(bool,X72))) ) | $spl3),
  inference(subsumption_resolution,[],[f31127,f11899])).
fof(f31127,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP9(X75,'T_0',X74) | ~p(s(bool,X72))) ) | $spl3),
  inference(superposition,[],[f12934,f418])).
fof(f31177,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~sP9(X71,'T_0',X70) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(subsumption_resolution,[],[f31126,f12967])).
fof(f31126,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~sP9(X71,'T_0',X70) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f12934,f442])).
fof(f31176,plain,(
  ( ! [X62,X60,X63] : (~sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl136),
  inference(cnf_transformation,[],[f31176_D])).
fof(f31176_D,plain,(
  ( ! [X62,X60,X63] : (~sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X62)) | ~p(s(bool,X63))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f31175,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(subsumption_resolution,[],[f31124,f416])).
fof(f31124,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,X63)) | ~sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12934,f442])).
fof(f30063,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X109))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f30017,f416])).
fof(f30017,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,X107) | ~p(s(bool,X107)) | p(s(bool,X108)) | sP9(X109,X108,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))) | p(s(bool,X109))) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f12933,f14166])).
fof(f12933,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X6)) | p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl3),
  inference(superposition,[],[f11899,f335])).
fof(f30058,plain,(
  ( ! [X70,X71] : (sP9(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl132),
  inference(cnf_transformation,[],[f30058_D])).
fof(f30058_D,plain,(
  ( ! [X70,X71] : (sP9(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f30060,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl3),
  inference(subsumption_resolution,[],[f30008,f11899])).
fof(f30008,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl3),
  inference(superposition,[],[f12933,f416])).
fof(f30059,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP9(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl3),
  inference(subsumption_resolution,[],[f30007,f11899])).
fof(f30007,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | p(s(bool,X74)) | sP9(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl3),
  inference(superposition,[],[f12933,f418])).
fof(f30057,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | sP9(X71,X70,'T_0') | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(subsumption_resolution,[],[f30006,f12967])).
fof(f30006,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,X70,'T_0') | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f12933,f442])).
fof(f30055,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(subsumption_resolution,[],[f30004,f416])).
fof(f30004,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12933,f442])).
fof(f29848,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17) | s(bool,'F_0') = s(bool,X15)) ) | $spl128),
  inference(cnf_transformation,[],[f29848_D])).
fof(f29848_D,plain,(
  ( ! [X17,X15] : (s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17) | s(bool,'F_0') = s(bool,X15)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f29846,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X15) != s(bool,X17)) ) | ($spl114 | $spl116)),
  inference(duplicate_literal_removal,[],[f29783])).
fof(f29783,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | ($spl114 | $spl116)),
  inference(resolution,[],[f29482,f338])).
fof(f29482,plain,(
  ( ! [X28,X29,X27] : (sP9(X28,X27,X29) | s(bool,'F_0') = s(bool,X27) | p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29)) ) | ($spl114 | $spl116)),
  inference(resolution,[],[f29460,f27497])).
fof(f29122,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'F_0') = s(bool,X38)) ) | $spl120),
  inference(subsumption_resolution,[],[f29074,f28515])).
fof(f29074,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'F_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl120),
  inference(superposition,[],[f28515,f442])).
fof(f29116,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl120),
  inference(duplicate_literal_removal,[],[f29073])).
fof(f29073,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X36)) ) | $spl120),
  inference(superposition,[],[f28515,f442])).
fof(f29120,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl120),
  inference(subsumption_resolution,[],[f29072,f411])).
fof(f29072,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,'F_0') = s(bool,X34) | ~p(s(bool,X34))) ) | $spl120),
  inference(superposition,[],[f28515,f418])).
fof(f29119,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl120),
  inference(subsumption_resolution,[],[f29071,f411])).
fof(f29071,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,'F_0') = s(bool,X32) | ~p(s(bool,X32))) ) | $spl120),
  inference(superposition,[],[f28515,f416])).
fof(f29118,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl120),
  inference(subsumption_resolution,[],[f29068,f411])).
fof(f29068,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl120),
  inference(superposition,[],[f28515,f391])).
fof(f29117,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl120),
  inference(subsumption_resolution,[],[f29067,f411])).
fof(f29067,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'F_0') = s(bool,X22) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl120),
  inference(superposition,[],[f28515,f391])).
fof(f28806,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl122),
  inference(subsumption_resolution,[],[f28759,f28517])).
fof(f28759,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl122),
  inference(superposition,[],[f28517,f442])).
fof(f28805,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl122),
  inference(subsumption_resolution,[],[f28758,f28517])).
fof(f28758,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl122),
  inference(superposition,[],[f28517,f442])).
fof(f28801,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl122),
  inference(duplicate_literal_removal,[],[f28757])).
fof(f28757,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl122),
  inference(superposition,[],[f28517,f418])).
fof(f28802,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl122),
  inference(duplicate_literal_removal,[],[f28756])).
fof(f28756,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl122),
  inference(superposition,[],[f28517,f416])).
fof(f28745,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28662,f4227])).
fof(f28662,plain,(
  ( ! [X285,X287,X286] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X285,bool),X286),s(fun(X285,bool),X287))))) ) | $spl16),
  inference(resolution,[],[f4227,f435])).
fof(f28744,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28648,f4227])).
fof(f28648,plain,(
  ( ! [X241,X239,X240] : (p(s(bool,'const.bool.IN_2'(s(X239,sK16(X240,X241,X239)),s(fun(X239,bool),X241))))) ) | $spl16),
  inference(resolution,[],[f4227,f435])).
fof(f28739,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28639,f4227])).
fof(f28639,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f4227,f430])).
fof(f430,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.NOT_CLAUSES|split|2')).
fof(f28738,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28638,f4227])).
fof(f28638,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f4227,f428])).
fof(f428,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_INTER','thm.bool.NOT_CLAUSES|split|1')).
fof(f28737,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28637,f4227])).
fof(f28637,plain,(
  ( ! [X232] : (p(s(bool,X232))) ) | $spl16),
  inference(resolution,[],[f4227,f408])).
fof(f28736,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28635,f4227])).
fof(f28635,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f4227,f405])).
fof(f405,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f225])).
fof(f28735,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28633,f4227])).
fof(f28633,plain,(
  ( ! [X229] : (p(s(bool,X229))) ) | $spl16),
  inference(resolution,[],[f4227,f402])).
fof(f28734,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f28632,f4227])).
fof(f28632,plain,(
  ( ! [X228] : (p(s(bool,X228))) ) | $spl16),
  inference(resolution,[],[f4227,f399])).
fof(f28640,plain,(
  $false | $spl16),
  inference(resolution,[],[f4227,f432])).
fof(f432,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INTER',aHL_TRUTH)).
fof(f28641,plain,(
  $false | $spl16),
  inference(resolution,[],[f4227,f433])).
fof(f28521,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108) | p(s(bool,X109))) ) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f28473,f416])).
fof(f28473,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,X107) | ~p(s(bool,X107)) | p(s(bool,X108)) | sP9(X109,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)),X108) | p(s(bool,X109))) ) | ($spl3 | $spl5)),
  inference(superposition,[],[f12932,f14166])).
fof(f12932,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f11899,f335])).
fof(f28514,plain,(
  ( ! [X70,X71] : (sP9(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl118),
  inference(cnf_transformation,[],[f28514_D])).
fof(f28514_D,plain,(
  ( ! [X70,X71] : (sP9(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f28518,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl3),
  inference(subsumption_resolution,[],[f28464,f11899])).
fof(f28464,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl3),
  inference(superposition,[],[f12932,f416])).
fof(f28516,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP9(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl3),
  inference(subsumption_resolution,[],[f28463,f11899])).
fof(f28463,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | ~p(s(bool,X73)) | p(s(bool,X74)) | sP9(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl3),
  inference(superposition,[],[f12932,f418])).
fof(f28513,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | sP9(X71,'T_0',X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(subsumption_resolution,[],[f28462,f12967])).
fof(f28462,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,'T_0',X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f12932,f442])).
fof(f28511,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(subsumption_resolution,[],[f28460,f416])).
fof(f28460,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl3),
  inference(superposition,[],[f12932,f442])).
fof(f27179,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl110),
  inference(duplicate_literal_removal,[],[f27137])).
fof(f27137,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl110),
  inference(superposition,[],[f26915,f442])).
fof(f27184,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl110),
  inference(subsumption_resolution,[],[f27136,f26915])).
fof(f27136,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X36)) ) | $spl110),
  inference(superposition,[],[f26915,f442])).
fof(f27183,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl110),
  inference(subsumption_resolution,[],[f27133,f415])).
fof(f27133,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') = s(bool,X30) | p(s(bool,X30))) ) | $spl110),
  inference(superposition,[],[f26915,f414])).
fof(f27182,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl110),
  inference(subsumption_resolution,[],[f27132,f415])).
fof(f27132,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X28) | p(s(bool,X28))) ) | $spl110),
  inference(superposition,[],[f26915,f412])).
fof(f27181,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl110),
  inference(subsumption_resolution,[],[f27129,f415])).
fof(f27129,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl110),
  inference(superposition,[],[f26915,f388])).
fof(f27180,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl110),
  inference(subsumption_resolution,[],[f27128,f415])).
fof(f27128,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl110),
  inference(superposition,[],[f26915,f388])).
fof(f26997,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38))) ) | $spl104),
  inference(subsumption_resolution,[],[f26977,f26907])).
fof(f26977,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl104),
  inference(superposition,[],[f26907,f442])).
fof(f26996,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36))) ) | $spl104),
  inference(subsumption_resolution,[],[f26976,f26907])).
fof(f26976,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl104),
  inference(superposition,[],[f26907,f442])).
fof(f26992,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl104),
  inference(duplicate_literal_removal,[],[f26973])).
fof(f26973,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl104),
  inference(superposition,[],[f26907,f414])).
fof(f26993,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl104),
  inference(duplicate_literal_removal,[],[f26972])).
fof(f26972,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl104),
  inference(superposition,[],[f26907,f412])).
fof(f26923,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) | $spl114),
  inference(cnf_transformation,[],[f26923_D])).
fof(f26923_D,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f26925,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(subsumption_resolution,[],[f26886,f416])).
fof(f26886,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X64)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(superposition,[],[f12909,f442])).
fof(f12909,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X4)) | p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl3),
  inference(superposition,[],[f11899,f335])).
fof(f26901,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X48)) | sP9(X51,X50,X49) | p(s(bool,X51))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f26882])).
fof(f26882,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X48)) | sP9(X51,X50,X49) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl3),
  inference(superposition,[],[f12909,f414])).
fof(f26902,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X44)) | sP9(X47,X46,X45) | p(s(bool,X47))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f26881])).
fof(f26881,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X44)) | sP9(X47,X46,X45) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl3),
  inference(superposition,[],[f12909,f412])).
fof(f26922,plain,(
  ( ! [X111,X112] : (sP9(X112,X111,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X112)) | ~p(s(bool,X111))) ) | $spl112),
  inference(cnf_transformation,[],[f26922_D])).
fof(f26922_D,plain,(
  ( ! [X111,X112] : (sP9(X112,X111,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X112)) | ~p(s(bool,X111))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f26921,plain,(
  ( ! [X111,X112,X110] : (s(bool,'F_0') != s(bool,X110) | ~p(s(bool,X111)) | sP9(X112,X111,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X112))) ) | ($spl3 | $spl43)),
  inference(subsumption_resolution,[],[f26872,f411])).
fof(f26872,plain,(
  ( ! [X111,X112,X110] : (s(bool,'F_0') != s(bool,X110) | ~p(s(bool,X111)) | p(s(bool,X110)) | sP9(X112,X111,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | p(s(bool,X112))) ) | ($spl3 | $spl43)),
  inference(superposition,[],[f12909,f15564])).
fof(f26914,plain,(
  ( ! [X83,X82] : (sP9(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) | $spl108),
  inference(cnf_transformation,[],[f26914_D])).
fof(f26914_D,plain,(
  ( ! [X83,X82] : (sP9(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f26917,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | ~p(s(bool,X90)) | p(s(bool,X89)) | sP9(X91,X90,'F_0') | p(s(bool,X91))) ) | $spl3),
  inference(subsumption_resolution,[],[f26865,f11899])).
fof(f26865,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | ~p(s(bool,X90)) | p(s(bool,X89)) | sP9(X91,X90,'F_0') | p(s(bool,X91)) | p(s(bool,X88))) ) | $spl3),
  inference(superposition,[],[f12909,f412])).
fof(f26916,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X86)) | p(s(bool,X85)) | sP9(X87,X86,'F_0') | p(s(bool,X87))) ) | $spl3),
  inference(subsumption_resolution,[],[f26864,f11899])).
fof(f26864,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X86)) | p(s(bool,X85)) | sP9(X87,X86,'F_0') | p(s(bool,X87)) | p(s(bool,X84))) ) | $spl3),
  inference(superposition,[],[f12909,f414])).
fof(f26913,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X82)) | sP9(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl3),
  inference(subsumption_resolution,[],[f26863,f12985])).
fof(f12985,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(subsumption_resolution,[],[f12949,f397])).
fof(f12949,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(superposition,[],[f11899,f442])).
fof(f26863,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X82)) | p(s(bool,X81)) | sP9(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl3),
  inference(superposition,[],[f12909,f442])).
fof(f26911,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(subsumption_resolution,[],[f26859,f411])).
fof(f26859,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(superposition,[],[f12909,f442])).
fof(f26906,plain,(
  ( ! [X28,X24,X27] : (sP9(X28,X27,X24) | p(s(bool,X24)) | p(s(bool,X28)) | ~p(s(bool,X27))) ) | $spl102),
  inference(cnf_transformation,[],[f26906_D])).
fof(f26906_D,plain,(
  ( ! [X28,X24,X27] : (sP9(X28,X27,X24) | p(s(bool,X24)) | p(s(bool,X28)) | ~p(s(bool,X27))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f26910,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl3),
  inference(subsumption_resolution,[],[f26855,f411])).
fof(f26855,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X49)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl3),
  inference(superposition,[],[f12909,f414])).
fof(f26909,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl3),
  inference(subsumption_resolution,[],[f26854,f411])).
fof(f26854,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X45)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl3),
  inference(superposition,[],[f12909,f412])).
fof(f26908,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl3),
  inference(subsumption_resolution,[],[f26851,f11899])).
fof(f26851,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl3),
  inference(superposition,[],[f12909,f388])).
fof(f26905,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl3),
  inference(subsumption_resolution,[],[f26850,f11899])).
fof(f26850,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl3),
  inference(superposition,[],[f12909,f388])).
fof(f26492,plain,(
  $false | ($spl3 | $spl98)),
  inference(trivial_inequality_removal,[],[f26482])).
fof(f26482,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl98)),
  inference(superposition,[],[f26345,f4757])).
fof(f4757,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))) | $spl3),
  inference(resolution,[],[f470,f412])).
fof(f26491,plain,(
  $false | ($spl41 | $spl98)),
  inference(trivial_inequality_removal,[],[f26483])).
fof(f26483,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl41 | $spl98)),
  inference(superposition,[],[f26345,f14335])).
fof(f14335,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl41),
  inference(trivial_inequality_removal,[],[f14333])).
fof(f14333,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl41),
  inference(superposition,[],[f13004,f442])).
fof(f26490,plain,(
  $false | ($spl3 | $spl98)),
  inference(trivial_inequality_removal,[],[f26485])).
fof(f26485,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl3 | $spl98)),
  inference(superposition,[],[f26345,f5159])).
fof(f5159,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl3),
  inference(resolution,[],[f4679,f412])).
fof(f26489,plain,(
  $false | ($spl43 | $spl98)),
  inference(trivial_inequality_removal,[],[f26487])).
fof(f26487,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl43 | $spl98)),
  inference(superposition,[],[f26345,f15564])).
fof(f26488,plain,(
  $false | $spl98),
  inference(equality_resolution,[],[f26345])).
fof(f26353,plain,(
  ( ! [X111,X112] : (sP9(X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X111) | p(s(bool,X112)) | ~p(s(bool,X111))) ) | $spl100),
  inference(cnf_transformation,[],[f26353_D])).
fof(f26353_D,plain,(
  ( ! [X111,X112] : (sP9(X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X111) | p(s(bool,X112)) | ~p(s(bool,X111))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f26352,plain,(
  ( ! [X111,X112,X110] : (s(bool,'F_0') != s(bool,X110) | ~p(s(bool,X111)) | sP9(X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X111) | p(s(bool,X112))) ) | ($spl3 | $spl43)),
  inference(subsumption_resolution,[],[f26306,f411])).
fof(f26306,plain,(
  ( ! [X111,X112,X110] : (s(bool,'F_0') != s(bool,X110) | ~p(s(bool,X111)) | p(s(bool,X110)) | sP9(X112,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X111) | p(s(bool,X112))) ) | ($spl3 | $spl43)),
  inference(superposition,[],[f12908,f15564])).
fof(f12908,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X0)) | p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f11899,f335])).
fof(f26343,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(subsumption_resolution,[],[f26293,f411])).
fof(f26293,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl3),
  inference(superposition,[],[f12908,f442])).
fof(f13006,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | $spl43),
  inference(cnf_transformation,[],[f13006_D])).
fof(f13006_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f14670,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X14))) ) | ($spl0 | $spl41)),
  inference(resolution,[],[f14530,f12706])).
fof(f14530,plain,(
  ( ! [X101,X102] : (sP8(X101,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X102)) ) | $spl41),
  inference(resolution,[],[f14341,f325])).
fof(f14341,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl41),
  inference(resolution,[],[f14336,f436])).
fof(f13023,plain,(
  ( ! [X11] : (sP9(X11,'F_0','const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | $spl46),
  inference(cnf_transformation,[],[f13023_D])).
fof(f13023_D,plain,(
  ( ! [X11] : (sP9(X11,'F_0','const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f14669,plain,(
  ( ! [X13] : (sP9(X13,'F_0','const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl8 | $spl41)),
  inference(resolution,[],[f14530,f12519])).
fof(f12519,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12518,f427])).
fof(f12518,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f12513,f11558])).
fof(f12513,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP6(X4,'F_0',X5)) ) | $spl0),
  inference(resolution,[],[f12506,f321])).
fof(f14339,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | $spl41),
  inference(subsumption_resolution,[],[f14328,f416])).
fof(f14328,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))))) ) | $spl41),
  inference(superposition,[],[f13004,f391])).
fof(f14338,plain,(
  ( ! [X10] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X10))) ) | $spl41),
  inference(subsumption_resolution,[],[f14327,f416])).
fof(f14327,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,X10))) ) | $spl41),
  inference(superposition,[],[f13004,f391])).
fof(f14084,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f14080,f427])).
fof(f14080,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f12282,f375])).
fof(f14083,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f14079,f427])).
fof(f14079,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f12282,f379])).
fof(f14078,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14072,f427])).
fof(f14072,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl5),
  inference(resolution,[],[f12281,f375])).
fof(f14075,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f14071,f427])).
fof(f14071,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl5),
  inference(resolution,[],[f12281,f379])).
fof(f14053,plain,(
  ( ! [X4,X5] : (~sP7(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~sP6(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X4,X5)) ) | $spl88),
  inference(cnf_transformation,[],[f14053_D])).
fof(f14053_D,plain,(
  ( ! [X4,X5] : (~sP7(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~sP6(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X4,X5)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f14041,plain,(
  ( ! [X4,X5] : (sP9('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X4,X5) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~sP7(X4,X5,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~sP6(X5,X4,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl5),
  inference(resolution,[],[f12012,f322])).
fof(f12012,plain,(
  ( ! [X101,X100] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X100,X101)) ) | $spl5),
  inference(resolution,[],[f472,f326])).
fof(f13004,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl41),
  inference(cnf_transformation,[],[f13004_D])).
fof(f13004_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f14032,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~p(s(bool,X14))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f12011,f12706])).
fof(f12011,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)),X99)) ) | $spl5),
  inference(resolution,[],[f472,f325])).
fof(f13022,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | $spl44),
  inference(cnf_transformation,[],[f13022_D])).
fof(f13022_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f14031,plain,(
  ( ! [X13] : (sP9(X13,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12011,f12519])).
fof(f12995,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | ~p(s(bool,X9))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f12706,f12011])).
fof(f12997,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | ~p(s(bool,X11))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f12706,f12381])).
fof(f12381,plain,(
  ( ! [X101,X102] : (sP8(X101,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)),X102)) ) | $spl5),
  inference(resolution,[],[f11961,f325])).
fof(f11961,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl5),
  inference(resolution,[],[f472,f436])).
fof(f13972,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) | $spl74),
  inference(cnf_transformation,[],[f13972_D])).
fof(f13972_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f13973,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f13967,f427])).
fof(f13967,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f5630,f375])).
fof(f5630,plain,(
  ( ! [X129] : (sP11(X129,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f4964,f376])).
fof(f4964,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15)))) | ($spl3 | $spl4)),
  inference(resolution,[],[f4678,f581])).
fof(f581,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK15))))) ) | $spl4),
  inference(resolution,[],[f471,f434])).
fof(f13970,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f13966,f427])).
fof(f13966,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK15))))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f5630,f379])).
fof(f13739,plain,(
  $false | ($spl5 | $spl11)),
  inference(subsumption_resolution,[],[f1772,f12039])).
fof(f12039,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | $spl5),
  inference(resolution,[],[f472,f412])).
fof(f13016,plain,(
  ( ! [X11] : (sP9(X11,'F_0','const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK15,sK13,sK12)),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12519,f12381])).
fof(f13014,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl0 | $spl5 | $spl8)),
  inference(resolution,[],[f12519,f12011])).
fof(f1221,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X27))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1217])).
fof(f1217,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,X27) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X27))) ) | $spl8),
  inference(superposition,[],[f1003,f414])).
fof(f1222,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X26)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1216])).
fof(f1216,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,X26) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X26)) ) | $spl8),
  inference(superposition,[],[f1003,f442])).
fof(f1220,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f1218])).
fof(f1218,plain,(
  ( ! [X28] : (s(bool,X28) != s(bool,X28) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f1003,f412])).
fof(f5342,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f5339])).
fof(f5339,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,X31) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl8),
  inference(superposition,[],[f1003,f412])).
fof(f5343,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl8),
  inference(trivial_inequality_removal,[],[f5338])).
fof(f5338,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,X30) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl8),
  inference(superposition,[],[f1003,f414])).
fof(f5344,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(trivial_inequality_removal,[],[f5337])).
fof(f5337,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,X29) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(superposition,[],[f1003,f442])).
fof(f11939,plain,(
  ( ! [X21,X19] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X19)) | s(bool,X19) != s(bool,X21)) ) | $spl38),
  inference(cnf_transformation,[],[f11939_D])).
fof(f11939_D,plain,(
  ( ! [X21,X19] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X19)) | s(bool,X19) != s(bool,X21)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f11940,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl18),
  inference(subsumption_resolution,[],[f10232,f412])).
fof(f10232,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl18),
  inference(superposition,[],[f4610,f388])).
fof(f11938,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl18),
  inference(subsumption_resolution,[],[f10233,f412])).
fof(f10233,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl18),
  inference(superposition,[],[f4610,f388])).
fof(f1003,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl8),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f10255,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') != s(bool,X44) | s(bool,'T_0') = s(bool,X43)) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10242])).
fof(f10242,plain,(
  ( ! [X43,X44] : (s(bool,X43) != s(bool,X43) | s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') != s(bool,X44) | s(bool,'T_0') = s(bool,X43)) ) | $spl18),
  inference(superposition,[],[f4610,f442])).
fof(f10254,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,'F_0') != s(bool,X46) | p(s(bool,X45))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10243])).
fof(f10243,plain,(
  ( ! [X45,X46] : (s(bool,X45) != s(bool,X45) | s(bool,'T_0') != s(bool,X46) | s(bool,'F_0') != s(bool,X46) | p(s(bool,X45))) ) | $spl18),
  inference(superposition,[],[f4610,f414])).
fof(f10253,plain,(
  ( ! [X47,X48] : (s(bool,'T_0') != s(bool,X48) | s(bool,'F_0') != s(bool,X48) | p(s(bool,X47))) ) | $spl18),
  inference(trivial_inequality_removal,[],[f10244])).
fof(f10244,plain,(
  ( ! [X47,X48] : (s(bool,X47) != s(bool,X47) | s(bool,'T_0') != s(bool,X48) | s(bool,'F_0') != s(bool,X48) | p(s(bool,X47))) ) | $spl18),
  inference(superposition,[],[f4610,f412])).
fof(f11903,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,X35)) ) | $spl34),
  inference(cnf_transformation,[],[f11903_D])).
fof(f11903_D,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,X35)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f10238,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X33) | s(bool,X32) != s(bool,X33) | ~p(s(bool,X32))) ) | $spl18),
  inference(superposition,[],[f4610,f416])).
fof(f10239,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X35) | s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | $spl18),
  inference(superposition,[],[f4610,f418])).
fof(f11901,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37)) ) | $spl18),
  inference(subsumption_resolution,[],[f10240,f4610])).
fof(f10240,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl18),
  inference(superposition,[],[f4610,f442])).
fof(f11900,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl30),
  inference(cnf_transformation,[],[f11900_D])).
fof(f11900_D,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f10247,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X51) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | ($spl3 | $spl18)),
  inference(superposition,[],[f4610,f5046])).
fof(f10654,plain,(
  ( ! [X317] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X317) | s(bool,X317) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl0 | $spl18)),
  inference(superposition,[],[f4610,f4883])).
fof(f11865,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl24),
  inference(cnf_transformation,[],[f11865_D])).
fof(f11865_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f11864,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f11829])).
fof(f11829,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl8),
  inference(resolution,[],[f11501,f338])).
fof(f6195,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | $spl21),
  inference(cnf_transformation,[],[f6195_D])).
fof(f6195_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f6298,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f5378,f5131])).
fof(f5131,plain,(
  ( ! [X101,X102] : (sP8(X101,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14)),X102)) ) | $spl3),
  inference(resolution,[],[f4679,f325])).
fof(f5378,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f5377,f427])).
fof(f5377,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f5374,f330])).
fof(f5374,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) ) | $spl4),
  inference(resolution,[],[f5364,f321])).
fof(f5364,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f5353,f332])).
fof(f5353,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl4),
  inference(resolution,[],[f4663,f316])).
fof(f6189,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,sK13,sK12)),s(fun(sK12,bool),sK14))) | ~p(s(bool,X9))) ) | ($spl3 | $spl4)),
  inference(resolution,[],[f5537,f5131])).
fof(f5537,plain,(
  ( ! [X12,X13] : (~sP8(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl4),
  inference(subsumption_resolution,[],[f5536,f332])).
fof(f5536,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f5535,f397])).
fof(f5535,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(subsumption_resolution,[],[f5523,f330])).
fof(f5523,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP7('F_0',X13,X12) | ~sP6(X13,'F_0',X12)) ) | $spl4),
  inference(resolution,[],[f4665,f321])).
fof(f5275,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f5264])).
fof(f5264,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f4662,f338])).
fof(f4662,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl4),
  inference(superposition,[],[f1821,f335])).
fof(f4610,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl18),
  inference(cnf_transformation,[],[f4610_D])).
fof(f4610_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4609,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f4588])).
fof(f4588,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl8),
  inference(resolution,[],[f1188,f338])).
fof(f1188,plain,(
  ( ! [X6,X8,X7] : (sP9(X8,X6,'F_0') | s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | p(s(bool,X8))) ) | $spl8),
  inference(superposition,[],[f1003,f335])).
fof(f4216,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | ~p(s(bool,X9))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f746,f1724])).
fof(f1724,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl2),
  inference(subsumption_resolution,[],[f1723,f332])).
fof(f1723,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl2),
  inference(subsumption_resolution,[],[f1722,f397])).
fof(f1722,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl2),
  inference(subsumption_resolution,[],[f1712,f330])).
fof(f1712,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) ) | $spl2),
  inference(resolution,[],[f982,f321])).
fof(f982,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl2),
  inference(superposition,[],[f973,f337])).
fof(f4215,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f746,f1244])).
fof(f1244,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f1243,f427])).
fof(f1243,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) ) | $spl2),
  inference(subsumption_resolution,[],[f1240,f330])).
fof(f1240,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) ) | $spl2),
  inference(resolution,[],[f1236,f321])).
fof(f1765,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) ) | ($spl2 | $spl4)),
  inference(resolution,[],[f1650,f638])).
fof(f1650,plain,(
  ( ! [X2,X3] : (~sP6(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f1645])).
fof(f1645,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~sP6(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl2),
  inference(resolution,[],[f1638,f338])).
fof(f1638,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X7,'F_0',X6)) ) | $spl2),
  inference(subsumption_resolution,[],[f1636,f326])).
fof(f1636,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) ) | $spl2),
  inference(resolution,[],[f1632,f322])).
fof(f1173,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1103,f1002])).
fof(f1103,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(resolution,[],[f1002,f431])).
fof(f431,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f233])).
fof(f1172,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1102,f1002])).
fof(f1102,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(resolution,[],[f1002,f429])).
fof(f429,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f232])).
fof(f1171,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1100,f1002])).
fof(f1100,plain,(
  ( ! [X238] : (~p(s(bool,X238))) ) | $spl6),
  inference(resolution,[],[f1002,f425])).
fof(f1170,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1098,f1002])).
fof(f1098,plain,(
  ( ! [X236] : (~p(s(bool,X236))) ) | $spl6),
  inference(resolution,[],[f1002,f397])).
fof(f1101,plain,(
  $false | $spl6),
  inference(resolution,[],[f1002,f427])).
fof(f1105,plain,(
  $false | ($spl1 | $spl6)),
  inference(resolution,[],[f1002,f468])).
fof(f1001,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f996])).
fof(f996,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f979,f338])).
fof(f979,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl2),
  inference(superposition,[],[f973,f335])).
fof(f467,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | $spl0),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f469,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | $spl2),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f262,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) <~> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.SUBSET_INTER','thm.pred_set.SUBSET_INTER')).
fof(f263,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))),
  inference(cnf_transformation,[],[f160])).
fof(f468,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) | $spl1),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f470,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | $spl3),
  inference(cnf_transformation,[],[f470_D])).
fof(f470_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f264,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK15)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.INTER_2'(s(fun(sK12,bool),sK14),s(fun(sK12,bool),sK15))))))),
  inference(cnf_transformation,[],[f160])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1854
% Passive clauses: 11150
% Generated clauses: 84507
% Final active clauses: 708
% Final passive clauses: 1932
% Input formulas: 45
% Initial clauses: 201
% 
% Selected by SInE selection: 44
% SInE iterations: 1
% 
% Duplicate literals: 7185
% Trivial inequalities: 421
% Fw subsumption resolutions: 14830
% Fw demodulations: 13333
% Bw demodulations: 30
% 
% Simple tautologies: 10526
% Equational tautologies: 238
% Forward subsumptions: 27763
% 
% Binary resolution: 38094
% Factoring: 5
% Forward superposition: 4644
% Backward superposition: 4612
% Equality resolution: 60
% 
% Split clauses: 966
% Split components: 2094
% SAT solver clauses: 1040
% SAT solver unit clauses: 17
% SAT solver binary clauses: 11
% 
% Sat splits: 966
% Sat splitting refutations: 29
% 
% Memory used [KB]: 18038
% Time elapsed: 1.362 s
% ------------------------------
% ------------------------------
% Success in time 13.034 s
