Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (102438)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 708
% Passive clauses: 8927
% Generated clauses: 19706
% Final active clauses: 646
% Final passive clauses: 7619
% Input formulas: 38
% Initial clauses: 157
% 
% Duplicate literals: 376
% Trivial inequalities: 124
% Fw subsumption resolutions: 1855
% Fw demodulations: 511
% Bw demodulations: 60
% 
% Simple tautologies: 207
% Equational tautologies: 64
% Forward subsumptions: 7048
% 
% Binary resolution: 9957
% Factoring: 1161
% Forward superposition: 2937
% Backward superposition: 2276
% Self superposition: 31
% Equality factoring: 28
% Equality resolution: 42
% 
% Split clauses: 466
% Split components: 944
% SAT solver clauses: 846
% SAT solver unit clauses: 27
% SAT solver binary clauses: 196
% 
% Sat splits: 466
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 840
% TWLsolver calls for satisfiability: 434
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (102456)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 425
% Passive clauses: 2733
% Generated clauses: 14710
% Final active clauses: 370
% Final passive clauses: 1055
% Input formulas: 38
% Initial clauses: 157
% 
% Duplicate literals: 234
% Trivial inequalities: 54
% Fw subsumption resolutions: 1113
% Fw demodulations: 1205
% Bw demodulations: 1271
% Condensations: 395
% Global subsumptions: 14
% 
% Simple tautologies: 229
% Equational tautologies: 19
% Forward subsumptions: 5707
% Bw demodulations to eq. taut.: 35
% 
% Binary resolution: 3829
% Factoring: 266
% Forward superposition: 2026
% Backward superposition: 1331
% Self superposition: 31
% Equality factoring: 28
% Equality resolution: 15
% 
% Split clauses: 69
% Split components: 143
% SAT solver clauses: 3022
% SAT solver unit clauses: 181
% SAT solver binary clauses: 73
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 10
% 
% Sat splits: 69
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2966
% TWLsolver calls for satisfiability: 1673
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f88924,plain,(
  $false),
  inference(subsumption_resolution,[],[f88923,f374])).
fof(f374,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f214,f211])).
fof(f211,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.bool.TRUTH')).
fof(f214,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f88923,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f88922,f413])).
fof(f413,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),sK14)))),
  inference(unit_resulting_resolution,[],[f209,f230])).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f209,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),sK14)))) & s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))) = s(sK12,'const.pred_set.chooser_2'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))),s('type.num.num',sK13)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f98])).
fof(f98,plain,(
  ? [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))) != s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s('type.num.num',X1))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ~! [X0] : ! [X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => ~s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))) = s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s('type.num.num',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => ~s(X6,'const.pred_set.CHOICE_1'(s(fun(X6,bool),X8))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X8))),s('type.num.num',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => ~s(X6,'const.pred_set.CHOICE_1'(s(fun(X6,bool),X8))) = s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X8))),s('type.num.num',X7))))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.pred_set_436_prove_ltb')).
fof(f88922,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),sK14))))),
  inference(forward_demodulation,[],[f88921,f1467])).
fof(f1467,plain,(
  ( ! [X4,X3] : (s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),X4))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),'const.pred_set.REST_1'(s(fun(X3,bool),X4)))))) )),
  inference(superposition,[],[f255,f253])).
fof(f253,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X1))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X8] : s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,'const.pred_set.CHOICE_1'(s(fun(X6,bool),X8)))))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.REST_DEF')).
fof(f255,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1)))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X12,X8] : s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X12)))))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.FINITE_DELETE')).
fof(f88921,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))))))),
  inference(subsumption_resolution,[],[f88920,f374])).
fof(f88920,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))))))),
  inference(forward_demodulation,[],[f88903,f5984])).
fof(f5984,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14)))))),
  inference(unit_resulting_resolution,[],[f405,f1852,f629])).
fof(f629,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f623,f374])).
fof(f623,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f260,f218])).
fof(f218,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.pred_set_436_prove_ltb',aHL_BOOL_CASES)).
fof(f260,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f1852,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))))))),
  inference(superposition,[],[f1185,f253])).
fof(f1185,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),'const.pred_set.DELETE_2'(s(fun(sK12,bool),X0),s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))))))))) )),
  inference(forward_demodulation,[],[f1173,f210])).
fof(f210,plain,(
  s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))) = s(sK12,'const.pred_set.chooser_2'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))),s('type.num.num',sK13)))),
  inference(cnf_transformation,[],[f142])).
fof(f1173,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,'const.pred_set.chooser_2'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))),s('type.num.num',sK13))),s(fun(sK12,bool),'const.pred_set.DELETE_2'(s(fun(sK12,bool),X0),s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))))))))) )),
  inference(unit_resulting_resolution,[],[f210,f258])).
fof(f258,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X8,X12,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X13)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8)))) & ~s(X6,X12) = s(X6,X13)))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.IN_DELETE')).
fof(f405,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f374,f225])).
fof(f225,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f88903,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,'const.pred_set.CHOICE_1'(s(fun(sK12,bool),sK14))),s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14)))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK12,bool),'const.pred_set.REST_1'(s(fun(sK12,bool),sK14))))))),
  inference(superposition,[],[f1961,f210])).
fof(f1961,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num',X2))),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(subsumption_resolution,[],[f1922,f211])).
fof(f1922,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X1),s('type.num.num',X2))),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f256,f562])).
fof(f562,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f252,f226])).
fof(f226,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f145])).
fof(f252,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X8))))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.SUBSET_REFL')).
fof(f256,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X2),s('type.num.num',X1))),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X2),s('type.num.num',X1))),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X2),s('type.num.num',X1))),s(fun(X0,bool),X3)))))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X2),s('type.num.num',X1))),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2,X3] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.bool.IN_2'(s(X0,'const.pred_set.chooser_2'(s(fun(X0,bool),X2),s('type.num.num',X1))),s(fun(X0,bool),X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X7,X8,X11] : ((~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X11))))) => p(s(bool,'const.bool.IN_2'(s(X6,'const.pred_set.chooser_2'(s(fun(X6,bool),X8),s('type.num.num',X7))),s(fun(X6,bool),X11)))))),
  file('thm.pred_set.pred_set_436_prove_ltb','thm.pred_set.pred_set_435_prove_ltb')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1613
% Passive clauses: 7708
% Generated clauses: 91202
% Final active clauses: 999
% Final passive clauses: 1676
% Input formulas: 38
% Initial clauses: 157
% 
% Duplicate literals: 1740
% Trivial inequalities: 206
% Fw subsumption resolutions: 2951
% Fw demodulations: 19900
% Bw demodulations: 2104
% Condensations: 1390
% Global subsumptions: 131
% 
% Simple tautologies: 153
% Equational tautologies: 1199
% Forward subsumptions: 53832
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 177
% 
% Unit resulting resolution: 34295
% Factoring: 4
% Forward superposition: 14938
% Backward superposition: 10171
% Self superposition: 80
% Equality factoring: 42
% Equality resolution: 10
% 
% Split clauses: 186
% Split components: 382
% SAT solver clauses: 10002
% SAT solver unit clauses: 3539
% SAT solver binary clauses: 3228
% 
% Sat splits: 186
% Sat splitting refutations: 1789
% 
% Memory used [KB]: 28784
% Time elapsed: 2.731 s
% ------------------------------
% ------------------------------
% Success in time 4.089 s
