Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.sum
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.sum
% SZS output start Proof for thm.sum
fof(f5173,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5172,f1915,f145,f242,f144,f240,f238,f4723,f4545,f146,f243,f241,f239])).
fof(f239,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | $spl1),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f241,plain,(
  p(s(bool,sK2)) | $spl2),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  p(s(bool,sK2)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f243,plain,(
  s(sK1,sK4) != s(sK1,sK5) | $spl5),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(sK1,sK4) != s(sK1,sK5) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f146,plain,(
  s(sK1,sK4) != s(sK1,sK5) | p(s(bool,sK2)) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | (~p(s(bool,sK2)) & s(sK1,sK4) = s(sK1,sK5))) & (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | p(s(bool,sK2)) | s(sK1,sK4) != s(sK1,sK5))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f102])).
fof(f102,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) | (~p(s(bool,X2)) & s(X1,X4) = s(X1,X5))) & (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) | p(s(bool,X2)) | s(X1,X4) != s(X1,X5)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) | (~p(s(bool,X2)) & s(X1,X4) = s(X1,X5))) & (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) | (p(s(bool,X2)) | s(X1,X4) != s(X1,X5))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) <~> (~p(s(bool,X2)) & s(X1,X4) = s(X1,X5)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) <=> (~p(s(bool,X2)) & s(X1,X4) = s(X1,X5)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))))) <=> (~p(s(bool,X2)) & s(X1,X4) = s(X1,X5)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X11))) = s('type.sum.sum'(X6,X7),'const.bool.COND_3'(s(bool,X8),s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X9))),s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X10))))) <=> (~p(s(bool,X8)) & s(X7,X10) = s(X7,X11)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X11))) = s('type.sum.sum'(X6,X7),'const.bool.COND_3'(s(bool,X8),s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X9))),s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X10))))) <=> (~p(s(bool,X8)) & s(X7,X10) = s(X7,X11)))),
  file('thm.sum.cond_sum_expand|split|3','thm.sum.cond_sum_expand|split|3')).
fof(f4545,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f4544,f243])).
fof(f4544,plain,(
  s(sK1,sK4) = s(sK1,sK5) | ($spl0 | $spl3)),
  inference(equality_resolution,[],[f1918])).
fof(f1918,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X0))) | s(sK1,sK5) = s(sK1,X0)) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f218,f1917])).
fof(f1917,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f1916,f201])).
fof(f201,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X15] : ! [X13,X14] : s(X15,X14) = s(X15,'const.bool.COND_3'(s(bool,'F_0'),s(X15,X13),s(X15,X14)))),
  file('thm.sum.cond_sum_expand|split|3','thm.bool.COND_CLAUSES|split|1')).
fof(f1916,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f238,f256])).
fof(f256,plain,(
  s(bool,'F_0') = s(bool,sK2) | $spl3),
  inference(resolution,[],[f162,f240])).
fof(f162,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.sum.cond_sum_expand|split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f218,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X15,X19] : ! [X10,X9] : (s('type.sum.sum'(X15,X19),'const.sum.INR_1'(s(X19,X9))) = s('type.sum.sum'(X15,X19),'const.sum.INR_1'(s(X19,X10))) <=> s(X19,X9) = s(X19,X10))),
  file('thm.sum.cond_sum_expand|split|3','thm.sum.INR_|49||49|')).
fof(f4723,plain,(
  $false | ($spl2 | $spl23)),
  inference(subsumption_resolution,[],[f4699,f487])).
fof(f487,plain,(
  s(bool,'T_0') != s(bool,sK2) | $spl23),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  s(bool,'T_0') != s(bool,sK2) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f4699,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl2),
  inference(resolution,[],[f241,f158])).
fof(f158,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.sum.cond_sum_expand|split|3','thm.bool.EQ_CLAUSES|split|1')).
fof(f238,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | $spl0),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f240,plain,(
  ~p(s(bool,sK2)) | $spl3),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  ~p(s(bool,sK2)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f144,plain,(
  ~p(s(bool,sK2)) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f242,plain,(
  s(sK1,sK4) = s(sK1,sK5) | $spl4),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s(sK1,sK4) = s(sK1,sK5) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f145,plain,(
  s(sK1,sK4) = s(sK1,sK5) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4)))))),
  inference(cnf_transformation,[],[f103])).
fof(f1915,plain,(
  $false | ($spl1 | $spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f1914])).
fof(f1914,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))) | ($spl1 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f1913,f242])).
fof(f1913,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f1912,f201])).
fof(f1912,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f239,f256])).
fof(f5172,plain,(
  $false | ($spl0 | $spl22)),
  inference(subsumption_resolution,[],[f5171,f220])).
fof(f220,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ! [X2,X3] : ~s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X15,X19] : ! [X9,X10] : ~s('type.sum.sum'(X15,X19),'const.sum.INL_1'(s(X15,X9))) = s('type.sum.sum'(X15,X19),'const.sum.INR_1'(s(X19,X10)))),
  file('thm.sum.cond_sum_expand|split|3','thm.sum.sum_distinct')).
fof(f5171,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) | ($spl0 | $spl22)),
  inference(forward_demodulation,[],[f5170,f202])).
fof(f202,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15] : ! [X13,X14] : s(X15,X13) = s(X15,'const.bool.COND_3'(s(bool,'T_0'),s(X15,X13),s(X15,X14)))),
  file('thm.sum.cond_sum_expand|split|3','thm.bool.COND_CLAUSES|split|0')).
fof(f5170,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK5))) = s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK3))),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK4))))) | ($spl0 | $spl22)),
  inference(forward_demodulation,[],[f238,f486])).
fof(f486,plain,(
  s(bool,'T_0') = s(bool,sK2) | $spl22),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  s(bool,'T_0') = s(bool,sK2) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
% SZS output end Proof for thm.sum
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 55
% Passive clauses: 1187
% Generated clauses: 5635
% Final active clauses: 38
% Final passive clauses: 557
% Input formulas: 37
% Initial clauses: 91
% 
% Duplicate literals: 264
% Trivial inequalities: 76
% Fw subsumption resolutions: 629
% Fw demodulations: 613
% Bw demodulations: 2
% 
% Simple tautologies: 164
% Equational tautologies: 51
% Forward subsumptions: 1996
% 
% Binary resolution: 70
% Factoring: 1
% Forward superposition: 1182
% Backward superposition: 1616
% Self superposition: 103
% Equality factoring: 96
% Equality resolution: 16
% 
% Split clauses: 585
% Split components: 1181
% SAT solver clauses: 712
% SAT solver unit clauses: 11
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 585
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 706
% TWLsolver calls for satisfiability: 39
% 
% Memory used [KB]: 2686
% Time elapsed: 0.134 s
% ------------------------------
% ------------------------------
% Success in time 0.129 s
