Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.transc
% (57005)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 163
% Passive clauses: 1946
% Generated clauses: 21575
% Final active clauses: 151
% Final passive clauses: 1593
% Input formulas: 39
% Initial clauses: 54
% 
% Duplicate literals: 1442
% Trivial inequalities: 104
% Fw subsumption resolutions: 2864
% Fw demodulations: 3319
% Bw demodulations: 10
% 
% Simple tautologies: 343
% Equational tautologies: 20
% Forward subsumptions: 10219
% 
% Binary resolution: 358
% Forward superposition: 8039
% Backward superposition: 5091
% Self superposition: 150
% Equality factoring: 34
% Equality resolution: 26
% 
% Split clauses: 799
% Split components: 1604
% SAT solver clauses: 1508
% SAT solver unit clauses: 8
% SAT solver binary clauses: 53
% 
% Sat splits: 799
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1507
% TWLsolver calls for satisfiability: 201
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f3414,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3401,f118,f176,f120,f178,f169,f212,f232,f234,f116,f174,f131,f188,f247,f249,f135,f190,f228,f230,f117,f172,f123,f181,f342,f349,f224,f225,f128,f186])).
fof(f186,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))) ) | $spl22),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f128,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X6))),s('type.realax.real',X6))))),
  file('thm.transc.DIFF_COMPOSITE|split|7','thm.transc.DIFF_EXP')).
fof(f225,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) ) | $spl78),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f224,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))) ) | ($spl12 | $spl22)),
  inference(backward_demodulation,[],[f181,f186])).
fof(f349,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) ) | $spl100),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f342,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) ) | ($spl26 | $spl78)),
  inference(resolution,[],[f225,f188])).
fof(f181,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))) ) | $spl12),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f123,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13] : s('type.realax.real','const.transc.exp_1'(s('type.realax.real',X13))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',X13)))),
  file('thm.transc.DIFF_COMPOSITE|split|7','ARITY_0_1_const.transc.exp')).
fof(f172,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f117,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X1),s('type.realax.real',X0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X0))))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X1),s('type.realax.real',X0)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X2),s('type.realax.real',X0))))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6))))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6)))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6))))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|7','thm.transc.DIFF_COMPOSITE|split|7')).
fof(f230,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | $spl83),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f228,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f172,f181])).
fof(f190,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.transc.DIFF_COMPOSITE|split|7','thm.bool.EQ_CLAUSES|split|2')).
fof(f249,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | $spl86),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f247,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s(fun('type.realax.real','type.realax.real'),sK2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | ($spl30 | $spl83)),
  inference(resolution,[],[f230,f190])).
fof(f188,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.transc.DIFF_COMPOSITE|split|7','thm.bool.EQ_CLAUSES|split|0')).
fof(f174,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | $spl2),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f116,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f100])).
fof(f234,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) | $spl84),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f232,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0))) | ($spl2 | $spl26)),
  inference(resolution,[],[f188,f174])).
fof(f212,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) ) | $spl66),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4)))))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10,X8,X11,X7,X6] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X11),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X6)))))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X8),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X8))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|7','thm.lim.DIFF_CHAIN')).
fof(f178,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f120,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|7',aHL_FALSITY)).
fof(f176,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f118,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|7',aHL_TRUTH)).
fof(f3401,plain,(
  $false | ($spl4 | $spl7 | $spl66 | $spl84 | $spl86 | $spl100)),
  inference(subsumption_resolution,[],[f3400,f176])).
fof(f3400,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl7 | $spl66 | $spl84 | $spl86 | $spl100)),
  inference(forward_demodulation,[],[f3399,f349])).
fof(f3399,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl4 | $spl7 | $spl66 | $spl84 | $spl86)),
  inference(subsumption_resolution,[],[f3398,f176])).
fof(f3398,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl7 | $spl66 | $spl84 | $spl86)),
  inference(forward_demodulation,[],[f3397,f234])).
fof(f3397,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl7 | $spl66 | $spl86)),
  inference(subsumption_resolution,[],[f3319,f178])).
fof(f3319,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.exp_0'),s('type.realax.real','const.transc.exp_1'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK2),s('type.realax.real',sK0)))))) | ($spl66 | $spl86)),
  inference(superposition,[],[f212,f249])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 68
% Passive clauses: 360
% Generated clauses: 3254
% Final active clauses: 56
% Final passive clauses: 253
% Input formulas: 39
% Initial clauses: 54
% 
% Duplicate literals: 154
% Trivial inequalities: 38
% Fw subsumption resolutions: 165
% Fw demodulations: 316
% Bw demodulations: 47
% Condensations: 153
% Global subsumptions: 4
% 
% Simple tautologies: 83
% Equational tautologies: 18
% Forward subsumptions: 1425
% 
% Binary resolution: 42
% Forward superposition: 1194
% Backward superposition: 616
% Self superposition: 68
% Equality factoring: 30
% Equality resolution: 9
% 
% Split clauses: 33
% Split components: 71
% SAT solver clauses: 591
% SAT solver unit clauses: 77
% SAT solver binary clauses: 30
% SAT solver learnt clauses: 9
% 
% Sat splits: 33
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 574
% TWLsolver calls for satisfiability: 234
% 
% Memory used [KB]: 5373
% Time elapsed: 0.156 s
% ------------------------------
% ------------------------------
% Success in time 0.775 s
