Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.measure
% (30829)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 704
% Passive clauses: 3360
% Generated clauses: 16309
% Final active clauses: 691
% Final passive clauses: 2421
% Input formulas: 260
% Initial clauses: 573
% 
% Duplicate literals: 202
% Trivial inequalities: 72
% Fw subsumption resolutions: 1044
% Fw demodulations: 957
% Bw demodulations: 12
% 
% Simple tautologies: 259
% Equational tautologies: 23
% Forward subsumptions: 10264
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 4461
% Factoring: 938
% Forward superposition: 4225
% Backward superposition: 3689
% Self superposition: 34
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 191
% Split components: 384
% SAT solver clauses: 317
% SAT solver unit clauses: 15
% SAT solver binary clauses: 53
% 
% Sat splits: 191
% 
% TWLsolver clauses: 317
% TWLsolver calls for satisfiability: 278
% 
% Memory used [KB]: 8059
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (30844)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 470
% Passive clauses: 2037
% Generated clauses: 12153
% Final active clauses: 451
% Final passive clauses: 1529
% Input formulas: 260
% Initial clauses: 573
% 
% Duplicate literals: 97
% Trivial inequalities: 29
% Fw subsumption resolutions: 535
% Fw demodulations: 1418
% Bw demodulations: 51
% Condensations: 47
% Global subsumptions: 24
% 
% Simple tautologies: 268
% Equational tautologies: 9
% Forward subsumptions: 5407
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 1898
% Factoring: 73
% Forward superposition: 3328
% Backward superposition: 2014
% Self superposition: 41
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 2852
% SAT solver unit clauses: 667
% SAT solver binary clauses: 168
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 20
% 
% Sat splits: 20
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 2758
% TWLsolver calls for satisfiability: 5438
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (30858)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2357
% Passive clauses: 15376
% Generated clauses: 231571
% Final active clauses: 2124
% Final passive clauses: 4798
% Input formulas: 260
% Initial clauses: 573
% Discarded non-redundant clauses: 35924
% 
% Duplicate literals: 11
% Trivial inequalities: 46
% Fw subsumption resolutions: 1373
% Fw demodulations: 79638
% Bw demodulations: 267
% Condensations: 26
% 
% Simple tautologies: 115
% Equational tautologies: 88
% Forward subsumptions: 62658
% Fw demodulations to eq. taut.: 44323
% Bw demodulations to eq. taut.: 317
% 
% Unit resulting resolution: 118697
% Factoring: 1
% Forward superposition: 16143
% Backward superposition: 17134
% Self superposition: 154
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 16976
% SAT solver unit clauses: 7784
% SAT solver binary clauses: 7205
% 
% Sat splits: 19
% Sat splitting refutations: 1560
% 
% Memory used [KB]: 114369
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f17109,plain,(
  $false),
  inference(subsumption_resolution,[],[f16954,f563])).
fof(f563,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))))),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14)))) & p(s(bool,'const.bool.IN_2'(s(fun(sK13,bool),sK15),s(fun(fun(sK13,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14)))))) & (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK15)))) & ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f375])).
fof(f375,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))))),
  file('thm.measure.MEASURE_SPACE_IN_MSPACE','thm.measure.MEASURE_SPACE_IN_MSPACE')).
fof(f16954,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))))),
  inference(resolution,[],[f2239,f810])).
fof(f810,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X0))))) )),
  inference(resolution,[],[f562,f582])).
fof(f582,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f203])).
fof(f203,axiom,(
  ! [X6] : ! [X11,X13] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X11),s(fun(X6,bool),X13)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X11)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X13))))))),
  file('thm.measure.MEASURE_SPACE_IN_MSPACE','thm.pred_set.SUBSET_DEF')).
fof(f562,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK15))))),
  inference(cnf_transformation,[],[f442])).
fof(f2239,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))))),
  inference(subsumption_resolution,[],[f2091,f560])).
fof(f560,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))),
  inference(cnf_transformation,[],[f442])).
fof(f2091,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK15),s(fun(sK13,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))),
  inference(resolution,[],[f561,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))))),
  inference(flattening,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))))),
  inference(ennf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7,X11] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X11),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X11),s(fun(X6,bool),'const.measure.m_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))))))),
  file('thm.measure.MEASURE_SPACE_IN_MSPACE','thm.measure.MEASURABLE_SETS_SUBSET_SPACE')).
fof(f561,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK13,bool),sK15),s(fun(fun(sK13,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(sK13,bool),'type.pair.prod'(fun(fun(sK13,bool),bool),fun(fun(sK13,bool),'type.realax.real'))),sK14))))))),
  inference(cnf_transformation,[],[f442])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 603
% Passive clauses: 6771
% Generated clauses: 16269
% Final active clauses: 596
% Final passive clauses: 6145
% Input formulas: 260
% Initial clauses: 241
% 
% Selected by SInE selection: 65
% SInE iterations: 1
% 
% Duplicate literals: 596
% Fw subsumption resolutions: 967
% Fw demodulations: 2
% 
% Simple tautologies: 939
% Equational tautologies: 73
% Forward subsumptions: 5623
% 
% Binary resolution: 11934
% Factoring: 11
% Forward superposition: 2636
% Equality resolution: 12
% 
% Split clauses: 1144
% Split components: 2725
% SAT solver clauses: 1271
% SAT solver unit clauses: 16
% SAT solver binary clauses: 243
% 
% Sat splits: 1144
% Sat splitting refutations: 16
% 
% Memory used [KB]: 7291
% Time elapsed: 0.390 s
% ------------------------------
% ------------------------------
% Success in time 12.135 s
