Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (96577)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 770
% Passive clauses: 8441
% Generated clauses: 18540
% Final active clauses: 739
% Final passive clauses: 7234
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 259
% Trivial inequalities: 39
% Fw subsumption resolutions: 1168
% Fw demodulations: 1410
% Bw demodulations: 30
% 
% Simple tautologies: 232
% Equational tautologies: 62
% Forward subsumptions: 6647
% 
% Binary resolution: 9867
% Factoring: 1460
% Forward superposition: 2132
% Backward superposition: 1781
% Self superposition: 23
% Equality factoring: 4
% Equality resolution: 19
% 
% Split clauses: 237
% Split components: 477
% SAT solver clauses: 556
% SAT solver unit clauses: 15
% SAT solver binary clauses: 73
% 
% Sat splits: 237
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 555
% TWLsolver calls for satisfiability: 362
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (96599)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 431
% Passive clauses: 2426
% Generated clauses: 13265
% Final active clauses: 427
% Final passive clauses: 1994
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 256
% Trivial inequalities: 38
% Fw subsumption resolutions: 555
% Fw demodulations: 1025
% Bw demodulations: 5
% Condensations: 517
% Global subsumptions: 23
% 
% Simple tautologies: 287
% Equational tautologies: 17
% Forward subsumptions: 3736
% 
% Binary resolution: 2390
% Factoring: 99
% Forward superposition: 2795
% Backward superposition: 2798
% Self superposition: 44
% Equality factoring: 28
% Equality resolution: 11
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 2946
% SAT solver unit clauses: 265
% SAT solver binary clauses: 110
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 19
% 
% Sat splits: 44
% 
% TWLsolver clauses: 2855
% TWLsolver calls for satisfiability: 2449
% 
% Memory used [KB]: 16119
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (96629)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1380
% Passive clauses: 7080
% Generated clauses: 218376
% Final active clauses: 1132
% Final passive clauses: 1983
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 450
% Trivial inequalities: 94
% Fw subsumption resolutions: 3887
% Fw demodulations: 109442
% Bw demodulations: 3460
% Condensations: 1466
% Global subsumptions: 2812
% 
% Simple tautologies: 124
% Equational tautologies: 95
% Forward subsumptions: 63177
% Fw demodulations to eq. taut.: 20486
% Bw demodulations to eq. taut.: 289
% 
% Unit resulting resolution: 76843
% Factoring: 1
% Forward superposition: 12863
% Backward superposition: 5072
% Self superposition: 58
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 103
% Split components: 211
% SAT solver clauses: 16474
% SAT solver unit clauses: 7330
% SAT solver binary clauses: 5166
% 
% Sat splits: 103
% Sat splitting refutations: 6373
% 
% Memory used [KB]: 271722
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (96927)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2233
% Passive clauses: 35416
% Generated clauses: 251893
% Final active clauses: 2086
% Final passive clauses: 30666
% Input formulas: 62
% Initial clauses: 204
% 
% Selected by SInE selection: 50
% SInE iterations: 1
% 
% Duplicate literals: 6299
% Trivial inequalities: 992
% Fw subsumption resolutions: 35047
% Fw demodulations: 10772
% Bw demodulations: 122
% 
% Simple tautologies: 7939
% Equational tautologies: 882
% Forward subsumptions: 137673
% 
% Binary resolution: 185245
% Factoring: 214
% Forward superposition: 6927
% Backward superposition: 4945
% Self superposition: 9
% Equality factoring: 2
% Equality resolution: 76
% 
% Split clauses: 18188
% Split components: 39790
% SAT solver clauses: 19235
% SAT solver unit clauses: 46
% SAT solver binary clauses: 779
% 
% Sat splits: 18188
% Sat splitting refutations: 80
% 
% Memory used [KB]: 66907
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (97023)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 321
% Passive clauses: 3120
% Generated clauses: 12513
% Final active clauses: 311
% Final passive clauses: 2697
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 182
% Trivial inequalities: 29
% Fw subsumption resolutions: 796
% Fw demodulations: 2676
% Bw demodulations: 39
% Global subsumptions: 27
% 
% Simple tautologies: 232
% Equational tautologies: 19
% Forward subsumptions: 4421
% Fw demodulations to eq. taut.: 876
% 
% Unit resulting resolution: 2707
% Factoring: 2
% Forward superposition: 3087
% Backward superposition: 2516
% Self superposition: 119
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 8
% Split components: 17
% SAT solver clauses: 3260
% SAT solver unit clauses: 381
% SAT solver binary clauses: 430
% 
% Sat splits: 8
% Sat splitting refutations: 103
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (97059)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1174
% Passive clauses: 4971
% Generated clauses: 50424
% Final active clauses: 1153
% Final passive clauses: 3878
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 2825
% Trivial inequalities: 149
% Fw subsumption resolutions: 5166
% Bw subsumption resolutions: 16
% Fw demodulations: 10116
% Bw demodulations: 12
% Condensations: 3676
% Global subsumptions: 2
% 
% Simple tautologies: 3765
% Equational tautologies: 170
% Forward subsumptions: 19782
% Backward subsumptions: 93
% 
% Binary resolution: 19567
% Factoring: 1
% Forward superposition: 3746
% Backward superposition: 5284
% Self superposition: 6
% Equality factoring: 6
% Equality resolution: 9
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4977
% SAT solver unit clauses: 234
% SAT solver binary clauses: 1172
% 
% Sat splits: 2
% 
% Memory used [KB]: 20212
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f23510,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f343,f600,f342,f602,f341,f604,f914,f961,f908,f963,f906,f905,f904,f967,f903,f965,f902,f901,f969,f899,f898,f897,f973,f896,f971,f895,f894,f975,f892,f891,f890,f979,f889,f977,f888,f887,f981,f874,f873,f864,f863,f862,f853,f852,f851,f842,f841,f840,f833,f832,f831,f826,f825,f816,f815,f814,f805,f804,f803,f1003,f790,f789,f784,f783,f781,f780,f779,f1065,f778,f1063,f777,f776,f1067,f768,f1069,f766,f765,f764,f1073,f763,f1071,f762,f761,f1075,f757,f1077,f742,f1079,f739,f738,f737,f1083,f736,f1081,f735,f734,f1085,f731,f730,f729,f728,f727,f726,f725,f724,f723,f1093,f722,f721,f720,f1089,f719,f718,f717,f716,f715,f714,f1097,f713,f1087,f712,f711,f1091,f710,f709,f708,f707,f1095,f706,f705,f1099,f703,f702,f701,f1103,f700,f1101,f699,f698,f1105,f681,f680,f679,f1109,f678,f1107,f677,f676,f1111,f674,f673,f672,f1115,f671,f1113,f670,f669,f1117,f668,f667,f666,f1121,f665,f1119,f664,f663,f1123,f662,f661,f660,f659,f658,f657,f656,f655,f654,f1133,f653,f1131,f652,f651,f650,f1129,f649,f1125,f648,f647,f1127,f641,f640,f638,f637,f609,f608,f960,f607,f606,f917,f918,f1176,f1135,f1195,f1194,f1201,f1193,f1203,f1187,f1186,f1205,f1185,f1207,f1519,f1566,f1513,f1568,f1511,f1510,f1509,f1572,f1508,f1570,f1507,f1506,f1574,f1504,f1503,f1502,f1578,f1501,f1576,f1500,f1499,f1580,f1497,f1496,f1495,f1584,f1494,f1582,f1493,f1492,f1586,f1479,f1478,f1469,f1458,f1447,f1438,f1437,f1436,f1431,f1430,f1421,f1410,f1395,f1394,f1389,f1388,f1386,f1385,f1384,f1668,f1383,f1666,f1382,f1381,f1670,f1373,f1672,f1371,f1370,f1369,f1676,f1368,f1674,f1367,f1366,f1678,f1362,f1680,f1347,f1682,f1344,f1343,f1342,f1686,f1341,f1684,f1340,f1339,f1688,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1696,f1327,f1326,f1325,f1692,f1324,f1323,f1322,f1321,f1320,f1319,f1700,f1318,f1690,f1317,f1316,f1694,f1315,f1314,f1313,f1312,f1698,f1311,f1310,f1702,f1308,f1307,f1306,f1706,f1305,f1704,f1304,f1303,f1708,f1300,f1710,f1299,f1284,f1283,f1282,f1716,f1281,f1714,f1280,f1279,f1718,f1277,f1276,f1275,f1722,f1274,f1720,f1273,f1272,f1724,f1271,f1270,f1269,f1728,f1268,f1726,f1267,f1266,f1712,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1738,f1256,f1736,f1255,f1254,f1253,f1734,f1252,f1730,f1251,f1250,f1732,f1244,f1243,f1241,f1240,f1212,f1211,f1565,f1210,f1209,f1522,f1523,f1781,f1740,f2191,f2190,f2183,f2182,f2158,f2157,f2156,f2155,f2264,f2154,f2153,f2262,f2152,f2151,f2266,f2148,f2147,f2146,f2145,f2268,f2142,f2141,f2140,f2139,f2270,f2136,f2135,f2134,f2133,f2272,f2124,f2123,f2122,f2121,f2120,f2119,f2118,f2117,f2100,f2099,f2096,f2095,f2084,f2083,f2082,f2081,f2072,f2070,f2060,f2058,f2048,f2047,f2046,f2045,f2044,f2043,f2042,f2041,f2032,f2031,f2030,f2029,f2028,f2027,f2026,f2025,f2016,f2014,f2004,f2002,f1992,f1991,f1990,f1989,f1988,f1987,f1986,f1985,f1978,f1977,f1976,f1975,f2362,f1972,f1971,f1970,f1969,f2366,f1968,f1967,f1966,f2364,f1965,f2368,f1962,f1961,f1960,f1959,f2370,f1948,f1947,f1946,f1945,f2374,f1944,f1943,f1942,f2372,f1941,f2376,f1940,f1939,f1938,f1937,f2384,f1936,f1935,f1934,f2380,f1933,f2388,f1932,f2378,f1931,f2386,f1930,f2382,f1929,f2390,f1927,f1926,f1925,f1924,f2394,f1923,f1922,f2392,f1921,f1920,f2396,f1915,f1914,f1913,f1912,f2398,f1897,f1896,f1895,f1894,f2400,f1879,f2402,f1878,f2404,f1877,f1876,f1875,f1874,f2412,f1873,f1872,f2408,f1871,f1870,f2416,f1869,f2406,f1868,f2410,f1867,f2414,f1866,f2418,f1863,f1862,f1861,f1860,f2420,f1857,f1856,f1855,f1854,f2422,f1853,f1852,f1851,f1850,f2424,f1849,f1848,f1847,f1846,f1845,f1844,f1843,f1842,f1841,f1840,f2434,f1839,f1838,f2432,f1837,f1836,f2430,f1835,f2426,f1834,f2428,f1833,f1832,f1831,f1830,f1829,f1828,f2438,f1827,f2436,f1826,f2440,f1824,f1822,f1793,f1792,f1791,f1790,f2241,f1787,f2484,f2200,f2201,f2442,f2600,f2599,f2598,f2603,f2612,f2604,f2610,f2605,f2608,f2615,f2614,f2618,f2823,f2822,f2821,f2825,f2835,f2826,f2833,f2827,f2831,f3002,f3008,f3010,f3066,f3072,f3077,f3094,f3100,f3105,f3128,f3132,f3138,f3142,f3148,f3152,f3163,f3167,f3180,f3179,f3186,f3195,f3194,f3198,f3206,f3205,f3209,f3634,f3640,f3642,f3694,f3700,f3705,f3722,f3728,f3733,f3756,f3760,f3766,f3770,f3776,f3780,f4452,f4451,f5484,f3971,f3970,f6381,f7572,f7571,f8983,f6930,f6929,f10283,f11730,f10872,f10871,f11380,f11736,f11381,f11734,f11382,f11741,f11740,f11742,f11732,f12608,f11750,f11749,f12258,f12614,f12259,f12612,f12260,f12619,f12618,f12620,f12610,f13186,f13231,f13185,f13233,f13181,f13235,f13180,f13237,f13176,f13268,f13144,f13274,f13140,f13139,f13138,f13280,f13135,f13278,f13134,f13133,f13284,f13129,f13128,f13127,f13290,f13124,f13288,f13123,f13122,f13294,f13118,f13117,f13116,f13300,f13113,f13298,f13112,f13111,f13304,f13087,f13086,f13085,f13084,f13083,f13082,f13081,f13069,f13055,f13041,f13027,f13015,f13014,f13013,f13006,f13005,f12993,f12979,f12965,f12939,f12938,f12937,f12936,f12935,f12934,f12933,f12926,f12925,f12919,f12918,f12917,f13456,f12916,f13454,f12915,f12914,f13458,f12904,f13462,f12900,f12899,f12898,f13468,f12895,f13466,f12894,f12893,f13472,f12863,f13482,f12840,f13488,f12825,f12824,f12823,f13496,f12822,f13494,f12821,f12820,f13498,f12817,f12816,f12815,f12814,f12813,f12812,f12811,f12810,f12809,f13506,f12808,f12807,f12806,f13502,f12805,f12804,f12803,f12802,f12801,f12800,f13510,f12799,f13500,f12798,f12797,f13504,f12796,f12795,f12794,f12793,f13508,f12792,f12791,f13512,f12787,f12786,f12785,f13518,f12782,f13516,f12781,f12780,f13522,f12774,f12773,f12772,f13530,f12770,f12769,f12768,f13534,f12766,f13528,f12765,f12764,f13536,f12762,f13532,f12761,f12760,f13538,f12756,f13540,f12754,f12749,f13544,f12748,f13546,f12742,f13548,f12741,f13550,f12733,f12732,f12731,f13556,f12728,f13554,f12727,f12726,f13560,f12722,f12721,f12720,f13566,f12717,f13564,f12716,f12715,f13570,f12714,f13572,f12713,f13574,f12712,f13576,f12711,f13578,f12710,f12709,f12708,f13582,f12707,f13580,f12706,f12705,f13542,f12704,f12703,f12702,f12701,f12700,f12699,f12698,f12697,f12696,f13592,f12695,f13590,f12694,f12693,f12692,f13588,f12691,f13584,f12690,f12689,f13586,f12683,f12682,f12680,f12679,f12629,f12628,f13242,f12627,f12626,f13190,f13191,f13679,f13594,f14155,f14161,f14168,f14195,f14199,f14207,f14211,f14346,f14352,f14359,f14386,f14390,f14398,f14402,f15104,f15103,f15102,f15100,f15099,f15098,f15097,f15156,f15096,f15158,f15091,f15090,f15089,f15087,f15086,f15085,f15084,f15166,f15083,f15168,f15079,f15199,f15047,f15205,f15043,f15042,f15041,f15211,f15038,f15209,f15037,f15036,f15215,f15032,f15031,f15030,f15221,f15027,f15219,f15026,f15025,f15225,f15021,f15020,f15019,f15231,f15016,f15229,f15015,f15014,f15235,f14990,f14989,f14988,f14987,f14986,f14985,f14984,f14972,f14958,f14944,f14930,f14918,f14917,f14916,f14909,f14908,f14896,f14882,f14868,f14842,f14841,f14840,f14839,f14838,f14837,f14836,f14829,f14828,f14822,f14821,f14820,f15387,f14819,f15385,f14818,f14817,f15389,f14807,f15393,f14803,f14802,f14801,f15399,f14798,f15397,f14797,f14796,f15403,f14762,f15417,f14739,f15423,f14720,f14719,f14718,f15435,f14717,f15433,f14716,f14715,f15437,f14712,f14711,f14710,f14709,f14708,f14707,f14706,f14705,f14704,f15445,f14703,f14702,f14701,f15441,f14700,f14699,f14698,f14697,f14696,f14695,f15449,f14694,f15439,f14693,f14692,f15443,f14691,f14690,f14689,f14688,f15447,f14687,f14686,f15451,f14682,f14681,f14680,f15457,f14677,f15455,f14676,f14675,f15461,f14669,f14668,f14667,f15154,f14665,f14664,f14663,f15164,f14661,f15150,f14660,f14659,f15152,f14657,f15160,f14656,f14655,f15162,f14651,f15467,f14649,f14644,f15471,f14643,f15473,f14637,f15475,f14636,f15477,f14628,f14627,f14626,f15483,f14623,f15481,f14622,f14621,f15487,f14617,f14616,f14615,f15493,f14612,f15491,f14611,f14610,f15497,f14609,f15499,f14608,f15501,f14607,f15503,f14606,f15505,f14605,f14604,f14603,f15509,f14602,f15507,f14601,f14600,f15469,f14599,f14598,f14597,f14596,f14595,f14594,f14593,f14592,f14591,f15519,f14590,f15517,f14589,f14588,f14587,f15515,f14586,f15511,f14585,f14584,f15513,f14578,f14577,f14575,f14574,f14524,f14523,f15173,f14522,f14521,f15109,f15110,f15606,f15521,f15632,f15638,f15631,f15640,f15652,f15653,f15662,f15661,f15660,f15665,f15659,f15667,f15942,f15943,f15655,f15945,f16193,f16199,f16201,f16248,f16254,f16256,f16265,f16271,f16273,f16313,f16315,f16320,f16322,f16327,f16329,f16334,f16338,f16343,f16347,f16358,f16366,f16354,f16369,f17445,f17444,f19137,f16631,f16630,f20761,f21370,f21369,f21376,f21368,f21367,f21378,f21461,f21469,f21458,f21471,f21823,f21822,f21826,f21821,f21820,f21828,f21916,f21915,f21919,f21956,f21962,f21965,f21984,f21986,f21992,f21994,f22004,f22006,f22767,f22766,f22764,f22763,f22762,f22761,f22812,f22760,f22810,f22759,f22814,f22758,f22757,f22818,f22756,f22816,f22755,f22820,f22749,f22748,f22824,f22747,f22822,f22746,f22826,f22745,f22828,f22744,f22830,f22742,f22741,f22834,f22740,f22832,f22739,f22836,f22738,f22838,f22737,f22840,f22844,f22724,f22723,f22863,f22715,f22714,f22869,f22873,f22875,f22877,f22879,f22881,f22679,f22678,f22677,f22676,f22903,f22672,f22671,f22901,f22670,f22669,f22908,f22664,f22663,f22662,f22661,f22912,f22656,f22655,f22654,f22653,f22916,f22648,f22647,f22646,f22645,f22920,f22924,f22928,f22932,f22936,f22628,f22627,f22626,f22625,f22624,f22623,f22622,f22621,f22620,f22619,f22618,f22617,f22616,f22615,f22614,f22613,f22952,f22592,f22591,f22588,f22587,f22975,f22979,f22574,f22573,f22572,f22571,f22997,f22558,f22556,f23012,f22542,f22540,f23027,f22526,f22524,f23036,f23040,f23044,f23048,f22506,f22505,f22504,f22503,f22502,f22501,f22500,f22499,f22498,f22497,f22496,f22495,f23064,f23068,f23072,f23076,f22478,f22477,f22476,f22475,f22474,f22473,f22472,f22471,f22470,f22469,f22468,f22467,f22466,f22465,f22464,f22463,f23098,f22450,f22448,f23113,f22434,f22432,f23128,f22418,f22416,f23137,f23141,f23145,f23149,f22846,f22398,f22397,f22396,f22395,f22394,f22393,f22392,f22391,f22390,f22389,f22388,f22387,f22373,f22372,f22371,f22370,f23168,f22364,f22363,f22362,f22361,f23173,f22359,f22358,f22357,f23171,f22356,f23176,f22351,f22350,f22349,f22348,f23180,f22344,f22343,f23184,f22319,f22318,f22317,f22316,f23196,f22313,f22312,f22311,f23194,f22310,f23200,f22309,f22308,f22307,f22306,f23208,f22303,f22302,f22301,f23204,f22300,f23214,f22297,f23202,f22296,f23212,f22295,f23206,f22294,f23218,f22289,f22288,f22287,f22286,f23225,f22282,f22281,f23223,f22280,f22279,f23230,f22275,f22274,f22806,f22270,f22804,f22269,f22808,f22250,f22249,f22248,f22247,f23238,f22218,f22217,f22216,f22215,f23250,f22196,f23256,f22195,f23258,f22194,f23260,f22188,f23262,f22187,f23264,f22186,f23266,f22184,f23268,f22182,f23270,f22181,f23272,f22179,f23274,f22172,f22171,f22170,f22169,f23288,f22162,f22161,f23284,f22160,f22159,f23298,f22152,f23282,f22151,f23286,f22150,f23296,f22149,f23306,f22144,f22143,f22142,f22141,f23310,f22136,f22135,f22134,f22133,f23314,f22132,f22131,f23316,f22130,f22129,f22128,f23318,f22127,f23320,f22126,f22125,f23322,f22124,f22123,f22122,f23324,f22121,f23326,f22120,f22119,f22118,f22117,f23328,f22116,f22115,f22114,f22113,f22112,f22111,f22110,f22109,f22108,f22107,f23338,f22106,f22105,f23336,f22104,f22103,f23334,f22102,f23330,f22101,f23332,f22096,f22095,f22090,f22089,f22084,f22083,f23346,f22078,f23344,f22077,f23352,f22075,f22073,f22016,f22015,f22014,f22013,f22859,f22009,f22769,f23507,f22771,f22772,f22871,f23509])).
fof(f23509,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f22871,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl764),
  inference(cnf_transformation,[],[f22871_D])).
fof(f22871_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl764),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl764])])).
fof(f22772,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(resolution,[],[f2484,f368])).
fof(f368,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.EQ_CLAUSES|split|2')).
fof(f22771,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(resolution,[],[f2484,f366])).
fof(f366,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.EQ_CLAUSES|split|3')).
fof(f23507,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl888),
  inference(cnf_transformation,[],[f23507_D])).
fof(f23507_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) <=> ~$spl888),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl888])])).
fof(f22769,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl307),
  inference(resolution,[],[f2484,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X8,X11] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11)))) <=> ! [X20] : (~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X20)))) => ? [X21] : ~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,happ(s(fun(X6,X7),X8),s(X6,X21))))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f22009,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f476])).
fof(f22859,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl763),
  inference(cnf_transformation,[],[f22859_D])).
fof(f22859_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl763),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl763])])).
fof(f22013,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f361])).
fof(f361,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.EQ_CLAUSES|split|1')).
fof(f22014,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f363])).
fof(f363,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.EQ_CLAUSES|split|0')).
fof(f22015,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f366])).
fof(f22016,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f368])).
fof(f22073,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f431])).
fof(f431,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.IMP_ANTISYM_AX')).
fof(f352,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9',aHL_FALSITY)).
fof(f22075,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f431])).
fof(f23352,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl887),
  inference(cnf_transformation,[],[f23352_D])).
fof(f23352_D,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl887),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl887])])).
fof(f22077,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f23344,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl883),
  inference(cnf_transformation,[],[f23344_D])).
fof(f23344_D,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl883),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl883])])).
fof(f22078,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f437])).
fof(f23346,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl885),
  inference(cnf_transformation,[],[f23346_D])).
fof(f23346_D,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl885),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl885])])).
fof(f22083,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f437])).
fof(f22084,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f437])).
fof(f22089,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f441])).
fof(f441,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f178,f191])).
fof(f178,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7,X6] : ((~(p(s(bool,X6)) | p(s(bool,X7))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X7)) => p(s(bool,'F_0')))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_9_prove_ltb')).
fof(f22090,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f441])).
fof(f22095,plain,(
  ~sP0('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f441])).
fof(f22096,plain,(
  ~sP0(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f441])).
fof(f23332,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl875),
  inference(cnf_transformation,[],[f23332_D])).
fof(f23332_D,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl875),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl875])])).
fof(f22101,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f452])).
fof(f452,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f23330,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl873),
  inference(cnf_transformation,[],[f23330_D])).
fof(f23330_D,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl873),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl873])])).
fof(f22102,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f452])).
fof(f23334,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl877),
  inference(cnf_transformation,[],[f23334_D])).
fof(f23334_D,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl877),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl877])])).
fof(f22103,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f452])).
fof(f22104,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f452])).
fof(f23336,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl878),
  inference(cnf_transformation,[],[f23336_D])).
fof(f23336_D,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl878),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl878])])).
fof(f22105,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f454])).
fof(f454,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f345,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.TRUTH')).
fof(f22106,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f454])).
fof(f344,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9',aHL_TRUTH)).
fof(f23338,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl880),
  inference(cnf_transformation,[],[f23338_D])).
fof(f23338_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl880),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl880])])).
fof(f22107,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f457])).
fof(f457,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f22108,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f457])).
fof(f22109,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f458])).
fof(f458,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f127,f193])).
fof(f127,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X25,X26] : ((p(s(bool,X26)) <=> ~p(s(bool,X25))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_29_prove_ltb')).
fof(f22110,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f458])).
fof(f22111,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f459])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f22112,plain,(
  sP1('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f459])).
fof(f22113,plain,(
  ~sP1('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f268])).
fof(f22114,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f461])).
fof(f22115,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f461])).
fof(f22116,plain,(
  ~sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f461])).
fof(f23328,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl871),
  inference(cnf_transformation,[],[f23328_D])).
fof(f23328_D,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl871),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl871])])).
fof(f22117,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f488])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f291])).
fof(f291,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f22118,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f488])).
fof(f22119,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f488])).
fof(f22120,plain,(
  ~sP5('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f488])).
fof(f23326,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl868),
  inference(cnf_transformation,[],[f23326_D])).
fof(f23326_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl868),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl868])])).
fof(f22121,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f683,f682,f1858,f2484,f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f1858,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f294])).
fof(f294,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f682,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f497])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f297])).
fof(f297,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f683,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f501])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f23324,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl866),
  inference(cnf_transformation,[],[f23324_D])).
fof(f23324_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl866),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl866])])).
fof(f22122,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1286,f1285,f1858,f2484,f490])).
fof(f1285,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f497])).
fof(f1286,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f501])).
fof(f22123,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f683,f1864,f1858,f2484,f490])).
fof(f1864,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f298])).
fof(f22124,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1286,f1864,f1858,f2484,f490])).
fof(f23322,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl864),
  inference(cnf_transformation,[],[f23322_D])).
fof(f23322_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl864),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl864])])).
fof(f22125,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12737,f1864,f1858,f2484,f490])).
fof(f12737,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f501])).
fof(f22126,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12737,f12736,f1858,f2484,f490])).
fof(f12736,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f497])).
fof(f23320,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl862),
  inference(cnf_transformation,[],[f23320_D])).
fof(f23320_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl862),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl862])])).
fof(f22127,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f684,f1864,f675,f2484,f491])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f675,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f295])).
fof(f684,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f23318,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl860),
  inference(cnf_transformation,[],[f23318_D])).
fof(f23318_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl860),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl860])])).
fof(f22128,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1287,f1864,f1278,f2484,f491])).
fof(f1278,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f493])).
fof(f1287,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f502])).
fof(f22129,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f684,f1864,f1858,f2484,f491])).
fof(f22130,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f1287,f1864,f1858,f2484,f491])).
fof(f23316,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl858),
  inference(cnf_transformation,[],[f23316_D])).
fof(f23316_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl858),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl858])])).
fof(f22131,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12738,f1864,f1858,f2484,f491])).
fof(f12738,plain,(
  ( ! [X0,X1] : (sP2(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f502])).
fof(f22132,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307 | $spl362)),
  inference(unit_resulting_resolution,[],[f12738,f1864,f12725,f2484,f491])).
fof(f12725,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f493])).
fof(f23314,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl857),
  inference(cnf_transformation,[],[f23314_D])).
fof(f23314_D,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl857),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl857])])).
fof(f22133,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f492])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f22134,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f492])).
fof(f22135,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f492])).
fof(f22136,plain,(
  ~sP4('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f492])).
fof(f23310,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl855),
  inference(cnf_transformation,[],[f23310_D])).
fof(f23310_D,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl855),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl855])])).
fof(f22141,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f22142,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f496])).
fof(f22143,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f496])).
fof(f22144,plain,(
  ~sP3('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f496])).
fof(f23306,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl853),
  inference(cnf_transformation,[],[f23306_D])).
fof(f23306_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') <=> ~$spl853),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl853])])).
fof(f22149,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f23296,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl849),
  inference(cnf_transformation,[],[f23296_D])).
fof(f23296_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl849),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl849])])).
fof(f22150,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f23286,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl845),
  inference(cnf_transformation,[],[f23286_D])).
fof(f23286_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl845),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl845])])).
fof(f22151,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f23282,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl841),
  inference(cnf_transformation,[],[f23282_D])).
fof(f23282_D,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl841),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl841])])).
fof(f22152,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f23298,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl851),
  inference(cnf_transformation,[],[f23298_D])).
fof(f23298_D,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl851),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl851])])).
fof(f22159,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f22160,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f23284,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl843),
  inference(cnf_transformation,[],[f23284_D])).
fof(f23284_D,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl843),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl843])])).
fof(f22161,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f22162,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f23288,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl847),
  inference(cnf_transformation,[],[f23288_D])).
fof(f23288_D,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl847),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl847])])).
fof(f22169,plain,(
  ~sP2('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f500])).
fof(f22170,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f500])).
fof(f22171,plain,(
  ~sP2('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f500])).
fof(f22172,plain,(
  ~sP2(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f500])).
fof(f23274,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') | $spl838),
  inference(cnf_transformation,[],[f23274_D])).
fof(f23274_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') <=> ~$spl838),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl838])])).
fof(f22179,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','F_0') | ($spl29 | $spl307)),
  inference(unit_resulting_resolution,[],[f1003,f2484,f504])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f144,f198,f197,f196,f195])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> s(bool,X24) = s(bool,X25)) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X25)) | ~p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_2|49|_prove_ltb')).
fof(f23272,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl836),
  inference(cnf_transformation,[],[f23272_D])).
fof(f23272_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl836),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl836])])).
fof(f22181,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl211 | $spl307)),
  inference(unit_resulting_resolution,[],[f2241,f2484,f504])).
fof(f23270,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl834),
  inference(cnf_transformation,[],[f23270_D])).
fof(f23270_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') <=> ~$spl834),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl834])])).
fof(f22182,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | ($spl29 | $spl307)),
  inference(unit_resulting_resolution,[],[f1003,f2484,f504])).
fof(f23268,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl832),
  inference(cnf_transformation,[],[f23268_D])).
fof(f23268_D,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl832),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl832])])).
fof(f22184,plain,(
  sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | ($spl211 | $spl307)),
  inference(unit_resulting_resolution,[],[f2241,f2484,f504])).
fof(f23266,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl831),
  inference(cnf_transformation,[],[f23266_D])).
fof(f23266_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl831),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl831])])).
fof(f22186,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f1135,f2484,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f301])).
fof(f23264,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl829),
  inference(cnf_transformation,[],[f23264_D])).
fof(f23264_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl829),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl829])])).
fof(f22187,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f1740,f2484,f507])).
fof(f23262,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl827),
  inference(cnf_transformation,[],[f23262_D])).
fof(f23262_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl827),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl827])])).
fof(f22188,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f2442,f2484,f507])).
fof(f23260,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl825),
  inference(cnf_transformation,[],[f23260_D])).
fof(f23260_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl825),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl825])])).
fof(f22194,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f1135,f2484,f507])).
fof(f23258,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl823),
  inference(cnf_transformation,[],[f23258_D])).
fof(f23258_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl823),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl823])])).
fof(f22195,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f1740,f2484,f507])).
fof(f23256,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl821),
  inference(cnf_transformation,[],[f23256_D])).
fof(f23256_D,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl821),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl821])])).
fof(f22196,plain,(
  ~sP5(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f2442,f2484,f507])).
fof(f23250,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl819),
  inference(cnf_transformation,[],[f23250_D])).
fof(f23250_D,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl819),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl819])])).
fof(f22215,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f22216,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f515])).
fof(f22217,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f515])).
fof(f22218,plain,(
  ~sP6('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f515])).
fof(f23238,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl817),
  inference(cnf_transformation,[],[f23238_D])).
fof(f23238_D,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') <=> ~$spl817),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl817])])).
fof(f22247,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f146,f201,f200])).
fof(f201,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) & p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_23_prove_ltb')).
fof(f22248,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f524])).
fof(f22249,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f524])).
fof(f22250,plain,(
  ~sP7(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f524])).
fof(f22808,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl726),
  inference(cnf_transformation,[],[f22808_D])).
fof(f22808_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl726),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl726])])).
fof(f22269,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f1928,f2484,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f1928,plain,(
  ( ! [X0,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f22804,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl722),
  inference(cnf_transformation,[],[f22804_D])).
fof(f22804_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl722),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl722])])).
fof(f22270,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1928,f2484,f531])).
fof(f22806,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl724),
  inference(cnf_transformation,[],[f22806_D])).
fof(f22806_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl724),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl724])])).
fof(f22274,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f1928,f2484,f531])).
fof(f22275,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1928,f2484,f531])).
fof(f23230,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl815),
  inference(cnf_transformation,[],[f23230_D])).
fof(f23230_D,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl815),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl815])])).
fof(f22279,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f22280,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f532])).
fof(f23223,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl811),
  inference(cnf_transformation,[],[f23223_D])).
fof(f23223_D,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl811),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl811])])).
fof(f22281,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f532])).
fof(f22282,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f532])).
fof(f23225,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl813),
  inference(cnf_transformation,[],[f23225_D])).
fof(f23225_D,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl813),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl813])])).
fof(f22286,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f532])).
fof(f22287,plain,(
  ~sP8('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f532])).
fof(f22288,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f532])).
fof(f22289,plain,(
  ~sP8('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f532])).
fof(f23218,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl808),
  inference(cnf_transformation,[],[f23218_D])).
fof(f23218_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl808),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl808])])).
fof(f22294,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f148,f204,f203])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) | p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_25_prove_ltb')).
fof(f23206,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl800),
  inference(cnf_transformation,[],[f23206_D])).
fof(f23206_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl800),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl800])])).
fof(f22295,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f23212,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl804),
  inference(cnf_transformation,[],[f23212_D])).
fof(f23212_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl804),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl804])])).
fof(f22296,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f23202,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl796),
  inference(cnf_transformation,[],[f23202_D])).
fof(f23202_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl796])])).
fof(f22297,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f23214,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl806),
  inference(cnf_transformation,[],[f23214_D])).
fof(f23214_D,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl806),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl806])])).
fof(f22300,plain,(
  sP9('F_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f23204,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl798),
  inference(cnf_transformation,[],[f23204_D])).
fof(f23204_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl798),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl798])])).
fof(f22301,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f22302,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f22303,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f23208,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl802),
  inference(cnf_transformation,[],[f23208_D])).
fof(f23208_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') <=> ~$spl802),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl802])])).
fof(f22306,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f352,f2484,f536])).
fof(f22307,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f352,f2484,f536])).
fof(f22308,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f2484,f2484,f536])).
fof(f22309,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2484,f536])).
fof(f23200,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl795),
  inference(cnf_transformation,[],[f23200_D])).
fof(f23200_D,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl795),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl795])])).
fof(f22310,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f23194,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl791),
  inference(cnf_transformation,[],[f23194_D])).
fof(f23194_D,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl791),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl791])])).
fof(f22311,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f539])).
fof(f22312,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f539])).
fof(f22313,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f539])).
fof(f23196,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl793),
  inference(cnf_transformation,[],[f23196_D])).
fof(f23196_D,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl793),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl793])])).
fof(f22316,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f539])).
fof(f22317,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f539])).
fof(f22318,plain,(
  ~sP9('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f539])).
fof(f22319,plain,(
  ~sP9('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f539])).
fof(f23184,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl788),
  inference(cnf_transformation,[],[f23184_D])).
fof(f23184_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl788),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl788])])).
fof(f22343,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1964,f2484,f548])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f318])).
fof(f318,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1964,plain,(
  ( ! [X0,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f321])).
fof(f321,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f22344,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f1964,f2484,f548])).
fof(f23180,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl787),
  inference(cnf_transformation,[],[f23180_D])).
fof(f23180_D,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl787),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl787])])).
fof(f22348,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f22349,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f549])).
fof(f22350,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f549])).
fof(f22351,plain,(
  ~sP10('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f549])).
fof(f23176,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl784),
  inference(cnf_transformation,[],[f23176_D])).
fof(f23176_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl784),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl784])])).
fof(f22356,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f188,f207,f206])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) => p(s(bool,X24)))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_27_prove_ltb')).
fof(f23171,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl780),
  inference(cnf_transformation,[],[f23171_D])).
fof(f23171_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl780),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl780])])).
fof(f22357,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f2484,f553])).
fof(f22358,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f553])).
fof(f22359,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f2484,f553])).
fof(f23173,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl782),
  inference(cnf_transformation,[],[f23173_D])).
fof(f23173_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') <=> ~$spl782),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl782])])).
fof(f22361,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f553])).
fof(f22362,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f553])).
fof(f22363,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f553])).
fof(f22364,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f553])).
fof(f23168,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl779),
  inference(cnf_transformation,[],[f23168_D])).
fof(f23168_D,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl779),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl779])])).
fof(f22370,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f556])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f22371,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f556])).
fof(f22372,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f556])).
fof(f22373,plain,(
  ~sP11('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f556])).
fof(f22387,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f352,f2442,f2484,f560])).
fof(f560,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13,X18,X16,X19] : ((s(bool,X13) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X16) = s(bool,X19))) => ((p(s(bool,X13)) => p(s(bool,X16))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.IMP_CONG')).
fof(f22388,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f352,f2442,f2484,f560])).
fof(f22389,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f2484,f560])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X14,X15] : s(X6,X14) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X14),s(X6,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.COND_CLAUSES|split|0')).
fof(f22390,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f2484,f560])).
fof(f22391,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f421,f2484,f560])).
fof(f22392,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f421,f2484,f560])).
fof(f22393,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f2484,f560])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X14,X15] : s(X6,X15) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X14),s(X6,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.COND_CLAUSES|split|1')).
fof(f22394,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f2484,f560])).
fof(f22395,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f422,f2484,f560])).
fof(f22396,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f422,f2484,f560])).
fof(f22397,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f560])).
fof(f591,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8,X9] : s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(fun(X4,X1),X6),s(fun(X4,X2),X7),s(fun(X4,X3),X8))),s(X4,X9))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),happ(s(fun(X1,fun(X2,fun(X3,X0))),X5),s(X1,happ(s(fun(X4,X1),X6),s(X4,X9))))),s(X2,happ(s(fun(X4,X2),X7),s(X4,X9))))),s(X3,happ(s(fun(X4,X3),X8),s(X4,X9)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X31,X32,X33,X34,X35] : ! [X36,X37,X38,X39,X40] : s(X31,happ(s(fun(X35,X31),'_dst_x0x1x1x1_4'(s(fun(X32,fun(X33,fun(X34,X31))),X36),s(fun(X35,X32),X37),s(fun(X35,X33),X38),s(fun(X35,X34),X39))),s(X35,X40))) = s(X31,happ(s(fun(X34,X31),happ(s(fun(X33,fun(X34,X31)),happ(s(fun(X32,fun(X33,fun(X34,X31))),X36),s(X32,happ(s(fun(X35,X32),X37),s(X35,X40))))),s(X33,happ(s(fun(X35,X33),X38),s(X35,X40))))),s(X34,happ(s(fun(X35,X34),X39),s(X35,X40)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','_dst_x0x1x1x1_ax')).
fof(f22398,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f591,f2484,f560])).
fof(f22846,plain,(
  ~p(s(bool,'T_0')) | $spl761),
  inference(cnf_transformation,[],[f22846_D])).
fof(f22846_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl761),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl761])])).
fof(f23149,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23148,f421])).
fof(f23148,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23147,f11795])).
fof(f11795,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f1208,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f1208,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0))))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7,X6] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X22),s(fun(X7,bool),X23)))) <=> ! [X21] : p(s(bool,happ(s(fun(X7,bool),X23),s(X7,happ(s(fun(X6,X7),X22),s(X6,X21)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f23147,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23146,f10917])).
fof(f10917,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f605,f434])).
fof(f605,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,X0))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f463])).
fof(f23146,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(forward_demodulation,[],[f22411,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X28,X29,X30] : s(X6,'const.bool.COND_3'(s(bool,X28),s(X6,X29),s(X6,X30))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),happ(s(fun(bool,fun(X6,fun(X6,X6))),'const.bool.COND_0'),s(bool,X28))),s(X6,X29))),s(X6,X30)))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','ARITY_0_3_const.bool.COND')).
fof(f22411,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2442,f591,f2484,f560])).
fof(f23145,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23144,f421])).
fof(f23144,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23143,f11795])).
fof(f23143,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23142,f10917])).
fof(f23142,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22412,f426])).
fof(f22412,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f421,f591,f2484,f560])).
fof(f23141,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23140,f421])).
fof(f23140,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23139,f11795])).
fof(f23139,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23138,f10917])).
fof(f23138,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22413,f426])).
fof(f22413,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f422,f591,f2484,f560])).
fof(f23137,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23136,f421])).
fof(f23136,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23135,f11795])).
fof(f23135,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23134,f10917])).
fof(f23134,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22414,f426])).
fof(f22414,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f560])).
fof(f22416,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f421,f2484,f560])).
fof(f22418,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f421,f2484,f560])).
fof(f23128,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23127,f421])).
fof(f23127,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23126,f11795])).
fof(f23126,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23125,f10917])).
fof(f23125,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22424,f426])).
fof(f22424,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f421,f2484,f560])).
fof(f22432,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f422,f2484,f560])).
fof(f22434,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f422,f2484,f560])).
fof(f23113,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23112,f421])).
fof(f23112,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23111,f11795])).
fof(f23111,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23110,f10917])).
fof(f23110,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22440,f426])).
fof(f22440,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f422,f2484,f560])).
fof(f22448,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f591,f2484,f560])).
fof(f22450,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f591,f2484,f560])).
fof(f23098,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23097,f421])).
fof(f23097,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23096,f11795])).
fof(f23096,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23095,f10917])).
fof(f23095,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22456,f426])).
fof(f22456,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f560])).
fof(f22463,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1135,f2484,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f22464,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1135,f2484,f562])).
fof(f22465,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1135,f2484,f562])).
fof(f22466,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1135,f2484,f562])).
fof(f22467,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1740,f2484,f562])).
fof(f22468,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1740,f2484,f562])).
fof(f22469,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1740,f2484,f562])).
fof(f22470,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1740,f2484,f562])).
fof(f22471,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f2484,f562])).
fof(f22472,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f2484,f562])).
fof(f22473,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f2484,f562])).
fof(f22474,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f2484,f562])).
fof(f22475,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f2484,f562])).
fof(f22476,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f2484,f562])).
fof(f22477,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f2484,f562])).
fof(f22478,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f2484,f562])).
fof(f23076,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f23075,f421])).
fof(f23075,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f23074,f11795])).
fof(f23074,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f23073,f10917])).
fof(f23073,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(forward_demodulation,[],[f22491,f426])).
fof(f22491,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1135,f591,f2484,f562])).
fof(f23072,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f23071,f421])).
fof(f23071,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f23070,f11795])).
fof(f23070,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f23069,f10917])).
fof(f23069,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(forward_demodulation,[],[f22492,f426])).
fof(f22492,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1740,f591,f2484,f562])).
fof(f23068,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23067,f421])).
fof(f23067,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23066,f11795])).
fof(f23066,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23065,f10917])).
fof(f23065,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22493,f426])).
fof(f22493,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f421,f591,f2484,f562])).
fof(f23064,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23063,f421])).
fof(f23063,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23062,f11795])).
fof(f23062,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23061,f10917])).
fof(f23061,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22494,f426])).
fof(f22494,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f422,f591,f2484,f562])).
fof(f22495,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f352,f2442,f2484,f564])).
fof(f564,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f22496,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f352,f2442,f2484,f564])).
fof(f22497,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f2484,f564])).
fof(f22498,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f2484,f564])).
fof(f22499,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f421,f2484,f564])).
fof(f22500,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f421,f2484,f564])).
fof(f22501,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f2484,f564])).
fof(f22502,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f2484,f564])).
fof(f22503,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f422,f2484,f564])).
fof(f22504,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f422,f2484,f564])).
fof(f22505,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f564])).
fof(f22506,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f591,f2484,f564])).
fof(f23048,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23047,f421])).
fof(f23047,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23046,f11795])).
fof(f23046,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f23045,f10917])).
fof(f23045,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(forward_demodulation,[],[f22519,f426])).
fof(f22519,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl304 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2442,f591,f2484,f564])).
fof(f23044,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23043,f421])).
fof(f23043,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23042,f11795])).
fof(f23042,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23041,f10917])).
fof(f23041,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22520,f426])).
fof(f22520,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f421,f591,f2484,f564])).
fof(f23040,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23039,f421])).
fof(f23039,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23038,f11795])).
fof(f23038,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23037,f10917])).
fof(f23037,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22521,f426])).
fof(f22521,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f422,f591,f2484,f564])).
fof(f23036,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23035,f421])).
fof(f23035,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23034,f11795])).
fof(f23034,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23033,f10917])).
fof(f23033,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22522,f426])).
fof(f22522,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f564])).
fof(f22524,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f421,f2484,f564])).
fof(f22526,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f421,f2484,f564])).
fof(f23027,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23026,f421])).
fof(f23026,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23025,f11795])).
fof(f23025,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23024,f10917])).
fof(f23024,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22532,f426])).
fof(f22532,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f421,f2484,f564])).
fof(f22540,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f422,f2484,f564])).
fof(f22542,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f422,f2484,f564])).
fof(f23012,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23011,f421])).
fof(f23011,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f23010,f11795])).
fof(f23010,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f23009,f10917])).
fof(f23009,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22548,f426])).
fof(f22548,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f422,f2484,f564])).
fof(f22556,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f345,f591,f2484,f564])).
fof(f22558,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f344,f591,f2484,f564])).
fof(f22997,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22996,f421])).
fof(f22996,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22995,f11795])).
fof(f22995,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22994,f10917])).
fof(f22994,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22564,f426])).
fof(f22564,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f591,f2484,f564])).
fof(f22571,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f345,f2484,f565])).
fof(f565,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f22572,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f345,f2484,f565])).
fof(f22573,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f344,f2484,f565])).
fof(f22574,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f344,f2484,f565])).
fof(f22979,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22978,f421])).
fof(f22978,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22977,f11795])).
fof(f22977,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22976,f10917])).
fof(f22976,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22583,f426])).
fof(f22583,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f591,f2484,f565])).
fof(f22975,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22974,f421])).
fof(f22974,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22973,f11795])).
fof(f22973,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22972,f10917])).
fof(f22972,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22584,f426])).
fof(f22584,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f591,f2484,f565])).
fof(f22587,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f345,f2484,f565])).
fof(f22588,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f345,f2484,f565])).
fof(f22591,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f344,f2484,f565])).
fof(f22592,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f344,f2484,f565])).
fof(f22952,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22951,f421])).
fof(f22951,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22950,f11795])).
fof(f22950,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22949,f10917])).
fof(f22949,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22612,f426])).
fof(f22612,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f591,f2484,f565])).
fof(f22613,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1135,f2484,f566])).
fof(f566,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f22614,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1135,f2484,f566])).
fof(f22615,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1135,f2484,f566])).
fof(f22616,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1135,f2484,f566])).
fof(f22617,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f345,f1740,f2484,f566])).
fof(f22618,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f345,f1740,f2484,f566])).
fof(f22619,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f344,f1740,f2484,f566])).
fof(f22620,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f344,f344,f1740,f2484,f566])).
fof(f22621,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f2484,f566])).
fof(f22622,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f2484,f566])).
fof(f22623,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f2484,f566])).
fof(f22624,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f2484,f566])).
fof(f22625,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f2484,f566])).
fof(f22626,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f2484,f566])).
fof(f22627,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f2484,f566])).
fof(f22628,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f2484,f566])).
fof(f22936,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22935,f421])).
fof(f22935,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22934,f11795])).
fof(f22934,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl102 | $spl307)),
  inference(forward_demodulation,[],[f22933,f10917])).
fof(f22933,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(forward_demodulation,[],[f22641,f426])).
fof(f22641,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl102 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1135,f591,f2484,f566])).
fof(f22932,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22931,f421])).
fof(f22931,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22930,f11795])).
fof(f22930,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl208 | $spl307)),
  inference(forward_demodulation,[],[f22929,f10917])).
fof(f22929,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(forward_demodulation,[],[f22642,f426])).
fof(f22642,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl208 | $spl307)),
  inference(unit_resulting_resolution,[],[f345,f1740,f591,f2484,f566])).
fof(f22928,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22927,f421])).
fof(f22927,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22926,f11795])).
fof(f22926,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22925,f10917])).
fof(f22925,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22643,f426])).
fof(f22643,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f421,f591,f2484,f566])).
fof(f22924,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22923,f421])).
fof(f22923,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22922,f11795])).
fof(f22922,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22921,f10917])).
fof(f22921,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22644,f426])).
fof(f22644,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f422,f591,f2484,f566])).
fof(f22920,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl777),
  inference(cnf_transformation,[],[f22920_D])).
fof(f22920_D,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') <=> ~$spl777),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl777])])).
fof(f22645,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f328])).
fof(f328,plain,(
  ! [X3,X1,X2] : ((~sP15(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP15(X3,X1,X2)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X3,X1,X2] : ((~sP15(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP15(X3,X1,X2)))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X3,X1,X2] : (sP15(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f22646,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f567])).
fof(f22647,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f567])).
fof(f22648,plain,(
  ~sP15('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'T_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f567])).
fof(f22916,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl775),
  inference(cnf_transformation,[],[f22916_D])).
fof(f22916_D,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl775),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl775])])).
fof(f22653,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X2)))),
  inference(rectify,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X3] : ((~sP14(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X3)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X3] : ((~sP14(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X3)))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X3] : (sP14(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f22654,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f571])).
fof(f22655,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f571])).
fof(f22656,plain,(
  ~sP14('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f571])).
fof(f22912,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl773),
  inference(cnf_transformation,[],[f22912_D])).
fof(f22912_D,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl773),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl773])])).
fof(f22661,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f345,f2484,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f334])).
fof(f334,plain,(
  ! [X2,X1,X3] : ((~sP13(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP13(X2,X1,X3)))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X2,X1,X3] : ((~sP13(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP13(X2,X1,X3)))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X2,X1,X3] : (sP13(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f22662,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f345,f2484,f575])).
fof(f22663,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f344,f2484,f575])).
fof(f22664,plain,(
  ~sP13('T_0','T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f344,f2484,f575])).
fof(f22908,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl771),
  inference(cnf_transformation,[],[f22908_D])).
fof(f22908_D,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') <=> ~$spl771),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl771])])).
fof(f22669,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f579])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f337])).
fof(f337,plain,(
  ! [X0,X2,X3] : ((~sP12(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP12(X0,X2,X3)))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0,X2,X3] : ((~sP12(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP12(X0,X2,X3)))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X2,X3] : (sP12(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f22670,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f579])).
fof(f22901,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl767),
  inference(cnf_transformation,[],[f22901_D])).
fof(f22901_D,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl767),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl767])])).
fof(f22671,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f579])).
fof(f22672,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f579])).
fof(f22903,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl769),
  inference(cnf_transformation,[],[f22903_D])).
fof(f22903_D,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl769),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl769])])).
fof(f22676,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f352,f2484,f579])).
fof(f22677,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f352,f2484,f579])).
fof(f22678,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f579])).
fof(f22679,plain,(
  ~sP12('T_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f579])).
fof(f22881,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22880,f422])).
fof(f22880,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl2 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22706,f2442])).
fof(f22706,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2150,f2159,f2144,f893,f2484,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X3,X1] : (~sP15(X3,X1,X2) | p(s(bool,X3)) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | ~sP14(X0,X1,X3) | ~sP13(X2,X1,X3) | ~sP12(X0,X2,X3)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & (~sP12(X0,X2,X3) | ~sP13(X2,X1,X3) | ~sP14(X0,X1,X3) | ~sP15(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & ((~sP12(X0,X2,X3) | ~sP13(X2,X1,X3) | ~sP14(X0,X1,X3) | ~sP15(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3)))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(nnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> (sP12(X0,X2,X3) & sP13(X2,X1,X3) & sP14(X0,X1,X3) & sP15(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(definition_folding,[],[f153,f212,f211,f210,f209])).
fof(f153,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X27,X24,X25,X26] : (s(bool,X26) = s(bool,'const.bool.COND_3'(s(bool,X25),s(bool,X24),s(bool,X27))) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | ~p(s(bool,X27))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X27))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X25)) | p(s(bool,X27)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.sat.sat_3|49|_prove_ltb')).
fof(f893,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP15(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f329])).
fof(f2144,plain,(
  ( ! [X0,X1] : (sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f2159,plain,(
  ( ! [X0,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f2150,plain,(
  ( ! [X0,X1] : (sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f22879,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22878,f422])).
fof(f22878,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl4 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22707,f2442])).
fof(f22707,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2150,f2159,f2144,f1498,f2484,f590])).
fof(f1498,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f569])).
fof(f22877,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22876,f422])).
fof(f22876,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) ) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22708,f2442])).
fof(f22708,plain,(
  ( ! [X0] : (s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,X0),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))))) ) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2144,f2159,f2150,f2137,f2484,f590])).
fof(f2137,plain,(
  ( ! [X0,X1] : (sP15(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f22875,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22874,f422])).
fof(f22874,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22709,f2442])).
fof(f22709,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2143,f2159,f2150,f2137,f2484,f590])).
fof(f2143,plain,(
  ( ! [X0,X1] : (sP14(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f332])).
fof(f22873,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22872,f422])).
fof(f22872,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22710,f2442])).
fof(f22710,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2159,f2143,f2149,f2137,f2484,f590])).
fof(f2149,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f335])).
fof(f22869,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22868,f422])).
fof(f22868,plain,(
  s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'F_0'),s(bool,'F_0'))) = s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl304 | $spl307)),
  inference(forward_demodulation,[],[f22711,f2442])).
fof(f22711,plain,(
  s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) = s(bool,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2159,f2144,f2149,f2137,f2484,f590])).
fof(f22714,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f345,f2484,f598])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f22715,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f344,f2484,f598])).
fof(f22863,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22862,f421])).
fof(f22862,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22861,f11795])).
fof(f22861,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22860,f10917])).
fof(f22860,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22720,f426])).
fof(f22720,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f352,f591,f2484,f598])).
fof(f22723,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f345,f2484,f2484,f598])).
fof(f22724,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl307),
  inference(unit_resulting_resolution,[],[f344,f2484,f2484,f598])).
fof(f22844,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22843,f421])).
fof(f22843,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl4 | $spl307)),
  inference(forward_demodulation,[],[f22842,f11795])).
fof(f22842,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,'T_0'),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | ($spl2 | $spl307)),
  inference(forward_demodulation,[],[f22841,f10917])).
fof(f22841,plain,(
  ~p(s(bool,'const.bool.COND_3'(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(forward_demodulation,[],[f22736,f426])).
fof(f22736,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))))),s(bool,happ(s(fun(sK17,bool),sK20),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))))) | $spl307),
  inference(unit_resulting_resolution,[],[f2484,f591,f2484,f598])).
fof(f22840,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl758),
  inference(cnf_transformation,[],[f22840_D])).
fof(f22840_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl758),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl758])])).
fof(f22737,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f931])).
fof(f931,plain,(
  ( ! [X6] : (sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X6) | p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f602,f454])).
fof(f22838,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl756),
  inference(cnf_transformation,[],[f22838_D])).
fof(f22838_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl756),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl756])])).
fof(f22738,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f932])).
fof(f932,plain,(
  ( ! [X7] : (sP1(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f602,f457])).
fof(f22836,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl754),
  inference(cnf_transformation,[],[f22836_D])).
fof(f22836_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl754),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl754])])).
fof(f22739,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f948])).
fof(f948,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f602,f553])).
fof(f22832,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl750),
  inference(cnf_transformation,[],[f22832_D])).
fof(f22832_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl750),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl750])])).
fof(f22740,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f948])).
fof(f22834,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl752),
  inference(cnf_transformation,[],[f22834_D])).
fof(f22834_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl752),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl752])])).
fof(f22741,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f948])).
fof(f22742,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f948])).
fof(f22830,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl748),
  inference(cnf_transformation,[],[f22830_D])).
fof(f22830_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl748),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl748])])).
fof(f22744,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1536])).
fof(f1536,plain,(
  ( ! [X6] : (sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X6) | p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f604,f454])).
fof(f22828,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl746),
  inference(cnf_transformation,[],[f22828_D])).
fof(f22828_D,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl746),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl746])])).
fof(f22745,plain,(
  sP1(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f1537])).
fof(f1537,plain,(
  ( ! [X7] : (sP1(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | p(s(bool,X7))) ) | $spl4),
  inference(resolution,[],[f604,f457])).
fof(f22826,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl744),
  inference(cnf_transformation,[],[f22826_D])).
fof(f22826_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl744),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl744])])).
fof(f22746,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f1553])).
fof(f1553,plain,(
  ( ! [X35,X36] : (sP11(X36,X35,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl4),
  inference(resolution,[],[f604,f553])).
fof(f22822,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl740),
  inference(cnf_transformation,[],[f22822_D])).
fof(f22822_D,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl740),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl740])])).
fof(f22747,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f1553])).
fof(f22824,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl742),
  inference(cnf_transformation,[],[f22824_D])).
fof(f22824_D,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl742),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl742])])).
fof(f22748,plain,(
  sP11('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f1553])).
fof(f22749,plain,(
  sP11(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f1553])).
fof(f22820,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl738),
  inference(cnf_transformation,[],[f22820_D])).
fof(f22820_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl738])])).
fof(f22755,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2219])).
fof(f2219,plain,(
  ( ! [X26,X27] : (sP9(X26,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X27) | p(s(bool,X27)) | p(s(bool,X26))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f22816,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl734),
  inference(cnf_transformation,[],[f22816_D])).
fof(f22816_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl734),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl734])])).
fof(f22756,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2219])).
fof(f22818,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | $spl736),
  inference(cnf_transformation,[],[f22818_D])).
fof(f22818_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) <=> ~$spl736),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl736])])).
fof(f22757,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2219])).
fof(f22758,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2219])).
fof(f22814,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl732),
  inference(cnf_transformation,[],[f22814_D])).
fof(f22814_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl732),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl732])])).
fof(f22759,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2220])).
fof(f2220,plain,(
  ( ! [X28,X29] : (sP9(X28,X29,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f22810,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl728),
  inference(cnf_transformation,[],[f22810_D])).
fof(f22810_D,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl728),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl728])])).
fof(f22760,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2220])).
fof(f22812,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl730),
  inference(cnf_transformation,[],[f22812_D])).
fof(f22812_D,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl730])])).
fof(f22761,plain,(
  sP9('F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2220])).
fof(f22762,plain,(
  sP9(happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2220])).
fof(f22763,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),'F_0') | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2221])).
fof(f2221,plain,(
  ( ! [X30,X31] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X30,X31) | p(s(bool,X30)) | p(s(bool,X31))) ) | $spl1),
  inference(resolution,[],[f600,f536])).
fof(f22764,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2221])).
fof(f22766,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0',happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f352,f2484,f2221])).
fof(f22767,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))),happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17)))))) | ($spl1 | $spl307)),
  inference(unit_resulting_resolution,[],[f2484,f2484,f2221])).
fof(f22006,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl721),
  inference(cnf_transformation,[],[f22006_D])).
fof(f22006_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl721),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl721])])).
fof(f22004,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl277),
  inference(unit_resulting_resolution,[],[f2414,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f21994,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl719),
  inference(cnf_transformation,[],[f21994_D])).
fof(f21994_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl719),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl719])])).
fof(f21992,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl273),
  inference(unit_resulting_resolution,[],[f2410,f485])).
fof(f21986,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl717),
  inference(cnf_transformation,[],[f21986_D])).
fof(f21986_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl717),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl717])])).
fof(f21984,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl271),
  inference(unit_resulting_resolution,[],[f2408,f485])).
fof(f21965,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl714),
  inference(cnf_transformation,[],[f21965_D])).
fof(f21965_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl714),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl714])])).
fof(f21962,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(resolution,[],[f2386,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f21956,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(unit_resulting_resolution,[],[f2386,f527])).
fof(f21919,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl712),
  inference(cnf_transformation,[],[f21919_D])).
fof(f21919_D,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl712),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl712])])).
fof(f21915,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f12725,f1864,f12738,f491])).
fof(f21916,plain,(
  sP5('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f1864,f12738,f491])).
fof(f21828,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl710),
  inference(cnf_transformation,[],[f21828_D])).
fof(f21828_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl710),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl710])])).
fof(f21820,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f12725,f12736,f12737,f489])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f21821,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f12725,f12736,f12737,f489])).
fof(f21826,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl708),
  inference(cnf_transformation,[],[f21826_D])).
fof(f21826_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl708),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl708])])).
fof(f21822,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f12736,f12737,f490])).
fof(f21823,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f1858,f1864,f12737,f490])).
fof(f21471,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl706),
  inference(cnf_transformation,[],[f21471_D])).
fof(f21471_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl706),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl706])])).
fof(f21458,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl211)),
  inference(unit_resulting_resolution,[],[f2241,f2211])).
fof(f2211,plain,(
  ( ! [X10,X11] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f600,f504])).
fof(f21469,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl704),
  inference(cnf_transformation,[],[f21469_D])).
fof(f21469_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl704),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl704])])).
fof(f21461,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | ($spl1 | $spl211)),
  inference(unit_resulting_resolution,[],[f2241,f2211])).
fof(f21378,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl702),
  inference(cnf_transformation,[],[f21378_D])).
fof(f21378_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl702),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl702])])).
fof(f21367,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f683,f682,f12725,f489])).
fof(f21368,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f683,f682,f12725,f489])).
fof(f21376,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl700),
  inference(cnf_transformation,[],[f21376_D])).
fof(f21376_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl700),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl700])])).
fof(f21369,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1286,f1285,f12725,f489])).
fof(f21370,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1286,f1285,f12725,f489])).
fof(f20761,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl698),
  inference(cnf_transformation,[],[f20761_D])).
fof(f20761_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl698),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl698])])).
fof(f16630,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl304),
  inference(unit_resulting_resolution,[],[f345,f2442,f505])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f16631,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl304),
  inference(unit_resulting_resolution,[],[f344,f2442,f505])).
fof(f19137,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl696),
  inference(cnf_transformation,[],[f19137_D])).
fof(f19137_D,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl696),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl696])])).
fof(f17444,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(unit_resulting_resolution,[],[f345,f2442,f505])).
fof(f17445,plain,(
  sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(unit_resulting_resolution,[],[f344,f2442,f505])).
fof(f16369,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl694),
  inference(cnf_transformation,[],[f16369_D])).
fof(f16369_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl694),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl694])])).
fof(f16354,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl211),
  inference(unit_resulting_resolution,[],[f352,f2241,f504])).
fof(f16366,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl692),
  inference(cnf_transformation,[],[f16366_D])).
fof(f16366_D,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl692),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl692])])).
fof(f16358,plain,(
  sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl211),
  inference(unit_resulting_resolution,[],[f352,f2241,f504])).
fof(f16347,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl691),
  inference(cnf_transformation,[],[f16347_D])).
fof(f16347_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl691),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl691])])).
fof(f16343,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl285),
  inference(unit_resulting_resolution,[],[f2422,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f16338,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl689),
  inference(cnf_transformation,[],[f16338_D])).
fof(f16338_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') <=> ~$spl689),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl689])])).
fof(f16334,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl283),
  inference(unit_resulting_resolution,[],[f2420,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f16329,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl687),
  inference(cnf_transformation,[],[f16329_D])).
fof(f16329_D,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl687),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl687])])).
fof(f16327,plain,(
  ~sP5('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl281),
  inference(unit_resulting_resolution,[],[f2418,f485])).
fof(f16322,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl685),
  inference(cnf_transformation,[],[f16322_D])).
fof(f16322_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl685),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl685])])).
fof(f16320,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl279),
  inference(unit_resulting_resolution,[],[f2416,f485])).
fof(f16315,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl683),
  inference(cnf_transformation,[],[f16315_D])).
fof(f16315_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl683),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl683])])).
fof(f16313,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl275),
  inference(unit_resulting_resolution,[],[f2412,f485])).
fof(f16273,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl680),
  inference(cnf_transformation,[],[f16273_D])).
fof(f16273_D,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl680),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl680])])).
fof(f16271,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(resolution,[],[f2390,f527])).
fof(f16265,plain,(
  sP8('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(unit_resulting_resolution,[],[f2390,f527])).
fof(f16256,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl678),
  inference(cnf_transformation,[],[f16256_D])).
fof(f16256_D,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl678),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl678])])).
fof(f16254,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(resolution,[],[f2388,f527])).
fof(f16248,plain,(
  sP8('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(unit_resulting_resolution,[],[f2388,f527])).
fof(f16201,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl676),
  inference(cnf_transformation,[],[f16201_D])).
fof(f16201_D,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl676),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl676])])).
fof(f16199,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(resolution,[],[f2368,f544])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f319])).
fof(f16193,plain,(
  sP10('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(unit_resulting_resolution,[],[f2368,f544])).
fof(f15945,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  introduced(theory_axiom,[])).
fof(f15655,plain,(
  p(s(bool,'F_0')) | $spl668),
  inference(cnf_transformation,[],[f15655_D])).
fof(f15655_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl668])])).
fof(f15943,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl674),
  inference(cnf_transformation,[],[f15943_D])).
fof(f15943_D,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) <=> ~$spl674),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl674])])).
fof(f15942,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) | p(s(bool,'F_0')) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f15941,f2442])).
fof(f15941,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f15935,f2442])).
fof(f15935,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X8))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP14(X8,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP13(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP12(X8,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(resolution,[],[f2138,f590])).
fof(f2138,plain,(
  ( ! [X0,X1] : (sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f15667,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl672),
  inference(cnf_transformation,[],[f15667_D])).
fof(f15667_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f15659,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f684,f675,f1864,f491])).
fof(f15665,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl670),
  inference(cnf_transformation,[],[f15665_D])).
fof(f15665_D,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl670])])).
fof(f15660,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1287,f1278,f1864,f491])).
fof(f15661,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f684,f1858,f1864,f491])).
fof(f15662,plain,(
  sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1287,f1858,f1864,f491])).
fof(f15653,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl666),
  inference(cnf_transformation,[],[f15653_D])).
fof(f15653_D,plain,(
  ( ! [X6,X7] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f15652,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | ($spl1 | $spl304)),
  inference(forward_demodulation,[],[f15649,f2442])).
fof(f15649,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X6,X7) | ~sP3(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ~sP2(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) ) | $spl1),
  inference(resolution,[],[f1859,f491])).
fof(f1859,plain,(
  ( ! [X0,X1] : (sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f15640,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl664),
  inference(cnf_transformation,[],[f15640_D])).
fof(f15640_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl664),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl664])])).
fof(f15631,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f352,f683,f682,f1858,f490])).
fof(f15638,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl662),
  inference(cnf_transformation,[],[f15638_D])).
fof(f15638_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl662),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl662])])).
fof(f15632,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f352,f1286,f1285,f1858,f490])).
fof(f15521,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl660),
  inference(cnf_transformation,[],[f15521_D])).
fof(f15521_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl660),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl660])])).
fof(f15606,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(global_subsumption,[],[f343,f342,f341,f12610,f15079,f15199,f15053,f15052,f15047,f15046,f15043,f15042,f15041,f15038,f15037,f15036,f15035,f15032,f15031,f15030,f15027,f15026,f15025,f15024,f15021,f15020,f15019,f15016,f15015,f15014,f14990,f14989,f14984,f14972,f14958,f14944,f14930,f14918,f14917,f14916,f14909,f14908,f14896,f14882,f14868,f14842,f14841,f14836,f14829,f14828,f14825,f14822,f14821,f14820,f14819,f14818,f14817,f14816,f14815,f14812,f14810,f14807,f14806,f14803,f14802,f14801,f14798,f14797,f14796,f14770,f14763,f14762,f14761,f14760,f14755,f14751,f14750,f14747,f14745,f14744,f14739,f14720,f14719,f14718,f14717,f14716,f14715,f14714,f14713,f14712,f14711,f14710,f14709,f14708,f14707,f14706,f14705,f14704,f14703,f14702,f14701,f14700,f14699,f14698,f14697,f14696,f14695,f14694,f14693,f14692,f14691,f14690,f14689,f14688,f14687,f14686,f14685,f14682,f14681,f14680,f14677,f14676,f14675,f14648,f14647,f14646,f14645,f14642,f14641,f14640,f14639,f14638,f14635,f14634,f14633,f14632,f14631,f14628,f14627,f14626,f14623,f14622,f14621,f14620,f14617,f14616,f14615,f14612,f14611,f14610,f14605,f14604,f14603,f14602,f14601,f14600,f14599,f14598,f14597,f14596,f14595,f14594,f14593,f14592,f14591,f14590,f14589,f14588,f14587,f14586,f14585,f14584,f14583,f14582,f14578,f14577,f14575,f14574,f15523,f15525,f15527,f14524,f14523,f14522,f14521,f15109,f15110,f15117,f15122,f15123,f15124,f15127,f15134,f15137,f15138,f15140,f15141,f15142,f15148])).
fof(f15148,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(superposition,[],[f12610,f354])).
fof(f354,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9',aHL_BOOL_CASES)).
fof(f15142,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X39,X40)) ) | $spl372),
  inference(resolution,[],[f12610,f555])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f15141,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X38,X37)) ) | $spl372),
  inference(resolution,[],[f12610,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f15140,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | sP11(X36,X35,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | p(s(bool,X36))) ) | $spl372),
  inference(resolution,[],[f12610,f553])).
fof(f15138,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X31,X32)) ) | $spl372),
  inference(resolution,[],[f12610,f538])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f15137,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X30,X29)) ) | $spl372),
  inference(resolution,[],[f12610,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f15134,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X23,X24)) ) | $spl372),
  inference(resolution,[],[f12610,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f15127,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X10)) ) | $spl372),
  inference(resolution,[],[f12610,f460])).
fof(f460,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f15124,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(resolution,[],[f12610,f457])).
fof(f15123,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X6)) ) | $spl372),
  inference(resolution,[],[f12610,f454])).
fof(f15122,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~sP1(X5,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(resolution,[],[f12610,f453])).
fof(f453,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f15117,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(resolution,[],[f12610,f434])).
fof(f15527,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X5)))) ) | $spl372),
  inference(forward_demodulation,[],[f15526,f422])).
fof(f15526,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,X1),s(X0,X5)))) ) | $spl372),
  inference(forward_demodulation,[],[f14571,f422])).
fof(f14571,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f422,f422,f12610,f428])).
fof(f428,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X11,X10,X13,X18,X16,X19] : ((s(bool,X10) = s(bool,X11) & (p(s(bool,X10)) => s(X6,X13) = s(X6,X18)) & (~p(s(bool,X10)) => s(X6,X16) = s(X6,X19))) => s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X18),s(X6,X19))) = s(X6,'const.bool.COND_3'(s(bool,X11),s(X6,X13),s(X6,X16))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.COND_CONG')).
fof(f15525,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl372),
  inference(forward_demodulation,[],[f15524,f422])).
fof(f15524,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl372),
  inference(forward_demodulation,[],[f14572,f426])).
fof(f14572,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f426,f422,f12610,f428])).
fof(f15523,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl372),
  inference(forward_demodulation,[],[f15522,f422])).
fof(f15522,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl372),
  inference(forward_demodulation,[],[f14573,f591])).
fof(f14573,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f591,f422,f12610,f428])).
fof(f14582,plain,(
  ( ! [X0] : (sP0(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f443])).
fof(f443,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f14583,plain,(
  ( ! [X0] : (sP0('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f444])).
fof(f444,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f14620,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f493])).
fof(f14631,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f497])).
fof(f14632,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f501])).
fof(f14633,plain,(
  ( ! [X0,X1] : (sP2(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f502])).
fof(f14634,plain,(
  ( ! [X0,X1] : (sP2('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f300])).
fof(f14635,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f418,f12610,f505])).
fof(f418,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X6,bool),X5))) = s(bool,happ(s(fun(X6,bool),X5),s(X6,'const.min.@_1'(s(fun(X6,bool),X5)))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.EXISTS_DEF')).
fof(f14638,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f421,f12610,f505])).
fof(f14639,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f422,f12610,f505])).
fof(f14640,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f426,f12610,f505])).
fof(f14641,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f591,f12610,f505])).
fof(f14642,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f418,f12610,f505])).
fof(f14645,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f421,f12610,f505])).
fof(f14646,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f422,f12610,f505])).
fof(f14647,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f426,f12610,f505])).
fof(f14648,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f591,f12610,f505])).
fof(f14685,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f306])).
fof(f14713,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f308])).
fof(f14714,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0',X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f308])).
fof(f14744,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f14745,plain,(
  ( ! [X0,X1] : (sP8(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f14747,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'T_0')) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f537])).
fof(f14750,plain,(
  ( ! [X0] : (sP9('T_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f537])).
fof(f14751,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f537])).
fof(f14755,plain,(
  ( ! [X0] : (sP9('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f538])).
fof(f14760,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f538])).
fof(f14761,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f538])).
fof(f14763,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f540])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f14770,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f14806,plain,(
  ( ! [X0,X1] : (sP10(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f551])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f322])).
fof(f14810,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f554])).
fof(f14812,plain,(
  ( ! [X0] : (sP11('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f555])).
fof(f14815,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f555])).
fof(f14816,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f555])).
fof(f14825,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),X0)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f15024,plain,(
  ( ! [X0,X1] : (sP15(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f569])).
fof(f15035,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP14(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f332])).
fof(f15046,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f576])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f335])).
fof(f15052,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f338])).
fof(f15053,plain,(
  ( ! [X0,X1] : (sP12(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK19)),X1)) ) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f581])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP12(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f338])).
fof(f15110,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(resolution,[],[f12610,f364])).
fof(f364,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f15109,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(resolution,[],[f12610,f362])).
fof(f362,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f220])).
fof(f14521,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f362])).
fof(f14522,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f364])).
fof(f15173,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl549),
  inference(cnf_transformation,[],[f15173_D])).
fof(f15173_D,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl549),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl549])])).
fof(f14523,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f365])).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f14524,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f367])).
fof(f367,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f14574,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f434])).
fof(f14575,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f434])).
fof(f14577,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f434])).
fof(f14578,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f434])).
fof(f15513,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl653),
  inference(cnf_transformation,[],[f15513_D])).
fof(f15513_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl653),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl653])])).
fof(f14584,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f453])).
fof(f14585,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f453])).
fof(f15511,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl651),
  inference(cnf_transformation,[],[f15511_D])).
fof(f15511_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl651),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl651])])).
fof(f14586,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f453])).
fof(f15515,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl655),
  inference(cnf_transformation,[],[f15515_D])).
fof(f15515_D,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl655),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl655])])).
fof(f14587,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f453])).
fof(f14588,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f453])).
fof(f14589,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f453])).
fof(f15517,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl656),
  inference(cnf_transformation,[],[f15517_D])).
fof(f15517_D,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f14590,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f454])).
fof(f15519,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl658),
  inference(cnf_transformation,[],[f15519_D])).
fof(f15519_D,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl658])])).
fof(f14591,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f457])).
fof(f14592,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f458])).
fof(f14593,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f459])).
fof(f14594,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f460])).
fof(f14595,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f460])).
fof(f14596,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f460])).
fof(f14597,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f460])).
fof(f14598,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f460])).
fof(f14599,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f460])).
fof(f15469,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl617),
  inference(cnf_transformation,[],[f15469_D])).
fof(f15469_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl617),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl617])])).
fof(f14600,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f488])).
fof(f14601,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f488])).
fof(f15507,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl647),
  inference(cnf_transformation,[],[f15507_D])).
fof(f15507_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl647),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl647])])).
fof(f14602,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f488])).
fof(f15509,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl649),
  inference(cnf_transformation,[],[f15509_D])).
fof(f15509_D,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl649),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl649])])).
fof(f14603,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f488])).
fof(f14604,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f488])).
fof(f14605,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f488])).
fof(f15505,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl644),
  inference(cnf_transformation,[],[f15505_D])).
fof(f15505_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f14606,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f684,f682,f675,f12610,f489])).
fof(f15503,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl642),
  inference(cnf_transformation,[],[f15503_D])).
fof(f15503_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f14607,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f684,f1285,f675,f12610,f489])).
fof(f15501,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl640),
  inference(cnf_transformation,[],[f15501_D])).
fof(f15501_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f14608,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f1287,f682,f1278,f12610,f489])).
fof(f15499,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl638),
  inference(cnf_transformation,[],[f15499_D])).
fof(f15499_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f14609,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f1287,f1285,f1278,f12610,f489])).
fof(f15497,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl637),
  inference(cnf_transformation,[],[f15497_D])).
fof(f15497_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl637),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl637])])).
fof(f14610,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f492])).
fof(f14611,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f492])).
fof(f15491,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl633),
  inference(cnf_transformation,[],[f15491_D])).
fof(f15491_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl633),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl633])])).
fof(f14612,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f492])).
fof(f15493,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl635),
  inference(cnf_transformation,[],[f15493_D])).
fof(f15493_D,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl635),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl635])])).
fof(f14615,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f492])).
fof(f14616,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f492])).
fof(f14617,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f492])).
fof(f15487,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl631),
  inference(cnf_transformation,[],[f15487_D])).
fof(f15487_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl631),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl631])])).
fof(f14621,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f496])).
fof(f14622,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f496])).
fof(f15481,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl627),
  inference(cnf_transformation,[],[f15481_D])).
fof(f15481_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl627),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl627])])).
fof(f14623,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f496])).
fof(f15483,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl629),
  inference(cnf_transformation,[],[f15483_D])).
fof(f15483_D,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl629),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl629])])).
fof(f14626,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f496])).
fof(f14627,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f496])).
fof(f14628,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f496])).
fof(f15477,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl624),
  inference(cnf_transformation,[],[f15477_D])).
fof(f15477_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f14636,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f1135,f12610,f505])).
fof(f15475,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl622),
  inference(cnf_transformation,[],[f15475_D])).
fof(f15475_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f14637,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f1740,f12610,f505])).
fof(f15473,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl620),
  inference(cnf_transformation,[],[f15473_D])).
fof(f15473_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f14643,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f1135,f12610,f505])).
fof(f15471,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl618),
  inference(cnf_transformation,[],[f15471_D])).
fof(f15471_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f14644,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f1740,f12610,f505])).
fof(f14649,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | ($spl29 | $spl372)),
  inference(unit_resulting_resolution,[],[f1003,f12610,f506])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f15467,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl615),
  inference(cnf_transformation,[],[f15467_D])).
fof(f15467_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl615),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl615])])).
fof(f14651,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | ($spl29 | $spl372)),
  inference(unit_resulting_resolution,[],[f1003,f12610,f506])).
fof(f15162,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl540),
  inference(cnf_transformation,[],[f15162_D])).
fof(f15162_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f14655,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f704,f12610,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f201])).
fof(f704,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f516])).
fof(f14656,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f704,f12610,f511])).
fof(f15160,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl538),
  inference(cnf_transformation,[],[f15160_D])).
fof(f15160_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f14657,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f704,f12610,f511])).
fof(f15152,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl530),
  inference(cnf_transformation,[],[f15152_D])).
fof(f15152_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl530),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl530])])).
fof(f14659,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f1309,f12610,f511])).
fof(f1309,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f516])).
fof(f14660,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f1309,f12610,f511])).
fof(f15150,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl528),
  inference(cnf_transformation,[],[f15150_D])).
fof(f15150_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f14661,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f1309,f12610,f511])).
fof(f15164,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl542),
  inference(cnf_transformation,[],[f15164_D])).
fof(f15164_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f14663,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f704,f12610,f511])).
fof(f14664,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f704,f12610,f511])).
fof(f14665,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f704,f12610,f511])).
fof(f15154,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl532),
  inference(cnf_transformation,[],[f15154_D])).
fof(f15154_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl532),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl532])])).
fof(f14667,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f1309,f12610,f511])).
fof(f14668,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f1309,f12610,f511])).
fof(f14669,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f1309,f12610,f511])).
fof(f15461,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl613),
  inference(cnf_transformation,[],[f15461_D])).
fof(f15461_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl613),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl613])])).
fof(f14675,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f515])).
fof(f14676,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f515])).
fof(f15455,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl609),
  inference(cnf_transformation,[],[f15455_D])).
fof(f15455_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl609),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl609])])).
fof(f14677,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f515])).
fof(f15457,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl611),
  inference(cnf_transformation,[],[f15457_D])).
fof(f15457_D,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl611),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl611])])).
fof(f14680,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f515])).
fof(f14681,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f515])).
fof(f14682,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f515])).
fof(f15451,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl606),
  inference(cnf_transformation,[],[f15451_D])).
fof(f15451_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f14686,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f345,f12610,f521])).
fof(f14687,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f345,f12610,f521])).
fof(f15447,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl602),
  inference(cnf_transformation,[],[f15447_D])).
fof(f15447_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f14688,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f345,f12610,f521])).
fof(f14689,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f344,f12610,f521])).
fof(f14690,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f344,f12610,f521])).
fof(f14691,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f344,f12610,f521])).
fof(f15443,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl598),
  inference(cnf_transformation,[],[f15443_D])).
fof(f15443_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f14692,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f12610,f521])).
fof(f14693,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f12610,f521])).
fof(f15439,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl594),
  inference(cnf_transformation,[],[f15439_D])).
fof(f15439_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f14694,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f12610,f521])).
fof(f15449,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl604),
  inference(cnf_transformation,[],[f15449_D])).
fof(f15449_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f14695,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f345,f12610,f521])).
fof(f14696,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f345,f12610,f521])).
fof(f14697,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f345,f12610,f521])).
fof(f14698,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f344,f12610,f521])).
fof(f14699,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f344,f12610,f521])).
fof(f14700,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f344,f12610,f521])).
fof(f15441,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl596),
  inference(cnf_transformation,[],[f15441_D])).
fof(f15441_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f14701,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f12610,f521])).
fof(f14702,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f12610,f521])).
fof(f14703,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f12610,f521])).
fof(f15445,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl600),
  inference(cnf_transformation,[],[f15445_D])).
fof(f15445_D,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f14704,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f345,f12610,f521])).
fof(f14705,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f345,f12610,f521])).
fof(f14706,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f345,f12610,f521])).
fof(f14707,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f344,f12610,f521])).
fof(f14708,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f344,f12610,f521])).
fof(f14709,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f344,f12610,f521])).
fof(f14710,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f12610,f12610,f521])).
fof(f14711,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f12610,f12610,f521])).
fof(f14712,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f12610,f12610,f521])).
fof(f15437,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl593),
  inference(cnf_transformation,[],[f15437_D])).
fof(f15437_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl593),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl593])])).
fof(f14715,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f524])).
fof(f14716,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f524])).
fof(f15433,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl589),
  inference(cnf_transformation,[],[f15433_D])).
fof(f15433_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl589),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl589])])).
fof(f14717,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f524])).
fof(f15435,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl591),
  inference(cnf_transformation,[],[f15435_D])).
fof(f15435_D,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl591),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl591])])).
fof(f14718,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f524])).
fof(f14719,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f524])).
fof(f14720,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f524])).
fof(f15423,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl587),
  inference(cnf_transformation,[],[f15423_D])).
fof(f15423_D,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl587),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl587])])).
fof(f14739,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f532])).
fof(f15417,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl585),
  inference(cnf_transformation,[],[f15417_D])).
fof(f15417_D,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl585),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl585])])).
fof(f14762,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f539])).
fof(f15403,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl583),
  inference(cnf_transformation,[],[f15403_D])).
fof(f15403_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl583),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl583])])).
fof(f14796,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f549])).
fof(f14797,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f549])).
fof(f15397,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl579),
  inference(cnf_transformation,[],[f15397_D])).
fof(f15397_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl579),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl579])])).
fof(f14798,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f549])).
fof(f15399,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl581),
  inference(cnf_transformation,[],[f15399_D])).
fof(f15399_D,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl581),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl581])])).
fof(f14801,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f549])).
fof(f14802,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f549])).
fof(f14803,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f549])).
fof(f15393,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl576),
  inference(cnf_transformation,[],[f15393_D])).
fof(f15393_D,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f14807,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f553])).
fof(f15389,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl575),
  inference(cnf_transformation,[],[f15389_D])).
fof(f15389_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl575),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl575])])).
fof(f14817,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f556])).
fof(f14818,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f556])).
fof(f15385,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl571),
  inference(cnf_transformation,[],[f15385_D])).
fof(f15385_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl571),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl571])])).
fof(f14819,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f556])).
fof(f15387,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl573),
  inference(cnf_transformation,[],[f15387_D])).
fof(f15387_D,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl573),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl573])])).
fof(f14820,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f556])).
fof(f14821,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f556])).
fof(f14822,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f556])).
fof(f14828,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f12610,f560])).
fof(f14829,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f12610,f560])).
fof(f14836,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f418,f12610,f562])).
fof(f14837,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f12610,f562])).
fof(f14838,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f12610,f562])).
fof(f14839,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f12610,f562])).
fof(f14840,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f12610,f562])).
fof(f14841,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f12610,f562])).
fof(f14842,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f12610,f562])).
fof(f14868,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f418,f12610,f562])).
fof(f14882,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f421,f12610,f562])).
fof(f14896,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f422,f12610,f562])).
fof(f14908,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f12610,f564])).
fof(f14909,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f12610,f564])).
fof(f14916,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f12610,f565])).
fof(f14917,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f12610,f565])).
fof(f14918,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f12610,f565])).
fof(f14930,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f12610,f565])).
fof(f14944,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f418,f12610,f566])).
fof(f14958,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f421,f12610,f566])).
fof(f14972,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f422,f12610,f566])).
fof(f14984,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f418,f12610,f566])).
fof(f14985,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f12610,f566])).
fof(f14986,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl102 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f12610,f566])).
fof(f14987,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f12610,f566])).
fof(f14988,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl208 | $spl372)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f12610,f566])).
fof(f14989,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f12610,f566])).
fof(f14990,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f12610,f566])).
fof(f15235,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl569),
  inference(cnf_transformation,[],[f15235_D])).
fof(f15235_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl569),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl569])])).
fof(f15014,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f345,f12610,f567])).
fof(f15015,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f344,f12610,f567])).
fof(f15229,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl565),
  inference(cnf_transformation,[],[f15229_D])).
fof(f15229_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl565),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl565])])).
fof(f15016,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f12610,f12610,f567])).
fof(f15231,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl567),
  inference(cnf_transformation,[],[f15231_D])).
fof(f15231_D,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl567),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl567])])).
fof(f15019,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f567])).
fof(f15020,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f567])).
fof(f15021,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f567])).
fof(f15225,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl563),
  inference(cnf_transformation,[],[f15225_D])).
fof(f15225_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl563),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl563])])).
fof(f15025,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f571])).
fof(f15026,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f571])).
fof(f15219,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl559),
  inference(cnf_transformation,[],[f15219_D])).
fof(f15219_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl559),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl559])])).
fof(f15027,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f571])).
fof(f15221,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl561),
  inference(cnf_transformation,[],[f15221_D])).
fof(f15221_D,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl561),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl561])])).
fof(f15030,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f571])).
fof(f15031,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f571])).
fof(f15032,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f571])).
fof(f15215,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl557),
  inference(cnf_transformation,[],[f15215_D])).
fof(f15215_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl557),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl557])])).
fof(f15036,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f575])).
fof(f15037,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f575])).
fof(f15209,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl553),
  inference(cnf_transformation,[],[f15209_D])).
fof(f15209_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl553),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl553])])).
fof(f15038,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f575])).
fof(f15211,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl555),
  inference(cnf_transformation,[],[f15211_D])).
fof(f15211_D,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl555),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl555])])).
fof(f15041,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f345,f352,f12610,f575])).
fof(f15042,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f344,f352,f12610,f575])).
fof(f15043,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f12610,f352,f12610,f575])).
fof(f15205,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl551),
  inference(cnf_transformation,[],[f15205_D])).
fof(f15205_D,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl551),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl551])])).
fof(f15047,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f579])).
fof(f15199,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(forward_demodulation,[],[f15054,f422])).
fof(f15054,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')))) ) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f587])).
fof(f587,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | p(s(bool,X0)) | p(s(bool,X2)) | ~p(s(bool,X3))) )),
  inference(cnf_transformation,[],[f340])).
fof(f15079,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl372),
  inference(unit_resulting_resolution,[],[f352,f352,f12610,f598])).
fof(f15168,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl547),
  inference(cnf_transformation,[],[f15168_D])).
fof(f15168_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl547),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl547])])).
fof(f15083,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f930])).
fof(f930,plain,(
  ( ! [X5] : (~sP1(X5,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ~p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f602,f453])).
fof(f15166,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl545),
  inference(cnf_transformation,[],[f15166_D])).
fof(f15166_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl545),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl545])])).
fof(f15084,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f935])).
fof(f935,plain,(
  ( ! [X10] : (~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X10) | ~p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f602,f460])).
fof(f15085,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f12610,f942])).
fof(f942,plain,(
  ( ! [X24,X23] : (sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f602,f521])).
fof(f15086,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f12610,f942])).
fof(f15087,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f12610,f942])).
fof(f15089,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f12610,f942])).
fof(f15090,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f12610,f942])).
fof(f15091,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl2 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f12610,f942])).
fof(f15158,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl537),
  inference(cnf_transformation,[],[f15158_D])).
fof(f15158_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl537),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl537])])).
fof(f15096,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f1535])).
fof(f1535,plain,(
  ( ! [X5] : (~sP1(X5,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ~p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f604,f453])).
fof(f15156,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | $spl535),
  inference(cnf_transformation,[],[f15156_D])).
fof(f15156_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) <=> ~$spl535),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl535])])).
fof(f15097,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f1540])).
fof(f1540,plain,(
  ( ! [X10] : (~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X10) | ~p(s(bool,X10))) ) | $spl4),
  inference(resolution,[],[f604,f460])).
fof(f15098,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f12610,f1547])).
fof(f1547,plain,(
  ( ! [X24,X23] : (sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X23,X24) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl4),
  inference(resolution,[],[f604,f521])).
fof(f15099,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'T_0') | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f12610,f1547])).
fof(f15100,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f12610,f1547])).
fof(f15102,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f345,f12610,f1547])).
fof(f15103,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f344,f12610,f1547])).
fof(f15104,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK19))) | ($spl4 | $spl372)),
  inference(unit_resulting_resolution,[],[f12610,f12610,f1547])).
fof(f14402,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl527),
  inference(cnf_transformation,[],[f14402_D])).
fof(f14402_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl527),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl527])])).
fof(f14398,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl189),
  inference(unit_resulting_resolution,[],[f1720,f487])).
fof(f14390,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl525),
  inference(cnf_transformation,[],[f14390_D])).
fof(f14390_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl525),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl525])])).
fof(f14386,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl183),
  inference(unit_resulting_resolution,[],[f1714,f486])).
fof(f14359,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl522),
  inference(cnf_transformation,[],[f14359_D])).
fof(f14359_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f14352,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl162),
  inference(resolution,[],[f1694,f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f14346,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl162),
  inference(unit_resulting_resolution,[],[f1694,f508])).
fof(f14211,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl521),
  inference(cnf_transformation,[],[f14211_D])).
fof(f14211_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl521),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl521])])).
fof(f14207,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl81),
  inference(unit_resulting_resolution,[],[f1113,f487])).
fof(f14199,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl519),
  inference(cnf_transformation,[],[f14199_D])).
fof(f14199_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl519),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl519])])).
fof(f14195,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl75),
  inference(unit_resulting_resolution,[],[f1107,f486])).
fof(f14168,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl516),
  inference(cnf_transformation,[],[f14168_D])).
fof(f14168_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f14161,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl58),
  inference(resolution,[],[f1091,f508])).
fof(f14155,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl58),
  inference(unit_resulting_resolution,[],[f1091,f508])).
fof(f13594,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl514),
  inference(cnf_transformation,[],[f13594_D])).
fof(f13594_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f13679,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(global_subsumption,[],[f343,f342,f341,f11732,f13176,f13268,f13150,f13149,f13144,f13143,f13140,f13139,f13138,f13135,f13134,f13133,f13132,f13129,f13128,f13127,f13124,f13123,f13122,f13121,f13118,f13117,f13116,f13113,f13112,f13111,f13087,f13086,f13081,f13069,f13055,f13041,f13027,f13015,f13014,f13013,f13006,f13005,f12993,f12979,f12965,f12939,f12938,f12933,f12926,f12925,f12922,f12919,f12918,f12917,f12916,f12915,f12914,f12913,f12912,f12909,f12907,f12904,f12903,f12900,f12899,f12898,f12895,f12894,f12893,f12869,f12864,f12863,f12862,f12861,f12856,f12852,f12851,f12848,f12846,f12845,f12840,f12825,f12824,f12823,f12822,f12821,f12820,f12819,f12818,f12817,f12816,f12815,f12814,f12813,f12812,f12811,f12810,f12809,f12808,f12807,f12806,f12805,f12804,f12803,f12802,f12801,f12800,f12799,f12798,f12797,f12796,f12795,f12794,f12793,f12792,f12791,f12790,f12787,f12786,f12785,f12782,f12781,f12780,f12753,f12752,f12751,f12750,f12747,f12746,f12745,f12744,f12743,f12740,f12739,f12738,f12737,f12736,f12733,f12732,f12731,f12728,f12727,f12726,f12725,f12722,f12721,f12720,f12717,f12716,f12715,f12710,f12709,f12708,f12707,f12706,f12705,f12704,f12703,f12702,f12701,f12700,f12699,f12698,f12697,f12696,f12695,f12694,f12693,f12692,f12691,f12690,f12689,f12688,f12687,f12683,f12682,f12680,f12679,f13596,f13598,f13600,f12629,f12628,f12627,f12626,f13190,f13191,f13198,f13203,f13204,f13205,f13208,f13215,f13218,f13219,f13221,f13222,f13223,f13229])).
fof(f13229,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(superposition,[],[f11732,f354])).
fof(f13223,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X39,X40)) ) | $spl362),
  inference(resolution,[],[f11732,f555])).
fof(f13222,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X38,X37)) ) | $spl362),
  inference(resolution,[],[f11732,f554])).
fof(f13221,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | sP11(X36,X35,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | p(s(bool,X36))) ) | $spl362),
  inference(resolution,[],[f11732,f553])).
fof(f13219,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X31,X32)) ) | $spl362),
  inference(resolution,[],[f11732,f538])).
fof(f13218,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X30,X29)) ) | $spl362),
  inference(resolution,[],[f11732,f537])).
fof(f13215,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X23,X24)) ) | $spl362),
  inference(resolution,[],[f11732,f521])).
fof(f13208,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X10)) ) | $spl362),
  inference(resolution,[],[f11732,f460])).
fof(f13205,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(resolution,[],[f11732,f457])).
fof(f13204,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X6)) ) | $spl362),
  inference(resolution,[],[f11732,f454])).
fof(f13203,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~sP1(X5,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(resolution,[],[f11732,f453])).
fof(f13198,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(resolution,[],[f11732,f434])).
fof(f13600,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X5)))) ) | $spl362),
  inference(forward_demodulation,[],[f13599,f422])).
fof(f13599,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,X1),s(X0,X5)))) ) | $spl362),
  inference(forward_demodulation,[],[f12676,f422])).
fof(f12676,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f422,f11732,f428])).
fof(f13598,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(forward_demodulation,[],[f13597,f422])).
fof(f13597,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(forward_demodulation,[],[f12677,f426])).
fof(f12677,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f422,f11732,f428])).
fof(f13596,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl362),
  inference(forward_demodulation,[],[f13595,f422])).
fof(f13595,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl362),
  inference(forward_demodulation,[],[f12678,f591])).
fof(f12678,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f422,f11732,f428])).
fof(f12687,plain,(
  ( ! [X0] : (sP0(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f443])).
fof(f12688,plain,(
  ( ! [X0] : (sP0('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f444])).
fof(f12739,plain,(
  ( ! [X0,X1] : (sP2('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f503])).
fof(f12740,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f418,f11732,f505])).
fof(f12743,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f421,f11732,f505])).
fof(f12744,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f11732,f505])).
fof(f12745,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f11732,f505])).
fof(f12746,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f11732,f505])).
fof(f12747,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f418,f11732,f505])).
fof(f12750,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f421,f11732,f505])).
fof(f12751,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f422,f11732,f505])).
fof(f12752,plain,(
  ( ! [X2,X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f426,f11732,f505])).
fof(f12753,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f591,f11732,f505])).
fof(f12790,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f516])).
fof(f12818,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'F_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f522])).
fof(f12819,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f523])).
fof(f12845,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f533])).
fof(f12846,plain,(
  ( ! [X0,X1] : (sP8(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f534])).
fof(f12848,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'T_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f537])).
fof(f12851,plain,(
  ( ! [X0] : (sP9('T_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f537])).
fof(f12852,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f537])).
fof(f12856,plain,(
  ( ! [X0] : (sP9('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f538])).
fof(f12861,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f538])).
fof(f12862,plain,(
  ( ! [X0] : (sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f538])).
fof(f12864,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f540])).
fof(f12869,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f541])).
fof(f12903,plain,(
  ( ! [X0,X1] : (sP10(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f551])).
fof(f12907,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0,'F_0')) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f554])).
fof(f12909,plain,(
  ( ! [X0] : (sP11('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f555])).
fof(f12912,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0',X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f555])).
fof(f12913,plain,(
  ( ! [X0] : (sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f555])).
fof(f12922,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),X0)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f558])).
fof(f13121,plain,(
  ( ! [X0,X1] : (sP15(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f569])).
fof(f13132,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f572])).
fof(f13143,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f576])).
fof(f13149,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f580])).
fof(f13150,plain,(
  ( ! [X0,X1] : (sP12(X0,'|dollar|exists_1'(s(fun(sK17,bool),sK21)),X1)) ) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f581])).
fof(f13191,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(resolution,[],[f11732,f364])).
fof(f13190,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(resolution,[],[f11732,f362])).
fof(f12626,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f362])).
fof(f12627,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f364])).
fof(f13242,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl391),
  inference(cnf_transformation,[],[f13242_D])).
fof(f13242_D,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl391),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl391])])).
fof(f12628,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f365])).
fof(f12629,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f367])).
fof(f12679,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f434])).
fof(f12680,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f434])).
fof(f12682,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f434])).
fof(f12683,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f434])).
fof(f13586,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl507),
  inference(cnf_transformation,[],[f13586_D])).
fof(f13586_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl507),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl507])])).
fof(f12689,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f453])).
fof(f12690,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f453])).
fof(f13584,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl505),
  inference(cnf_transformation,[],[f13584_D])).
fof(f13584_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl505),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl505])])).
fof(f12691,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f453])).
fof(f13588,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl509),
  inference(cnf_transformation,[],[f13588_D])).
fof(f13588_D,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl509),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl509])])).
fof(f12692,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f453])).
fof(f12693,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f453])).
fof(f12694,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f453])).
fof(f13590,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl510),
  inference(cnf_transformation,[],[f13590_D])).
fof(f13590_D,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f12695,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f454])).
fof(f13592,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl512),
  inference(cnf_transformation,[],[f13592_D])).
fof(f13592_D,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f12696,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f457])).
fof(f12697,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f458])).
fof(f12698,plain,(
  sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f459])).
fof(f12699,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f460])).
fof(f12700,plain,(
  ~sP1('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f460])).
fof(f12701,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f460])).
fof(f12702,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f460])).
fof(f12703,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f460])).
fof(f12704,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f460])).
fof(f13542,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl471),
  inference(cnf_transformation,[],[f13542_D])).
fof(f13542_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl471),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl471])])).
fof(f12705,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f488])).
fof(f12706,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f488])).
fof(f13580,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl501),
  inference(cnf_transformation,[],[f13580_D])).
fof(f13580_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl501),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl501])])).
fof(f12707,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f488])).
fof(f13582,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl503),
  inference(cnf_transformation,[],[f13582_D])).
fof(f13582_D,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl503),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl503])])).
fof(f12708,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f488])).
fof(f12709,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f488])).
fof(f12710,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f488])).
fof(f13578,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl498),
  inference(cnf_transformation,[],[f13578_D])).
fof(f13578_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f12711,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f684,f682,f675,f11732,f489])).
fof(f13576,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl496),
  inference(cnf_transformation,[],[f13576_D])).
fof(f13576_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f12712,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f684,f1285,f675,f11732,f489])).
fof(f13574,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl494),
  inference(cnf_transformation,[],[f13574_D])).
fof(f13574_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f12713,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f1287,f682,f1278,f11732,f489])).
fof(f13572,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl492),
  inference(cnf_transformation,[],[f13572_D])).
fof(f13572_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f12714,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f1287,f1285,f1278,f11732,f489])).
fof(f13570,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl491),
  inference(cnf_transformation,[],[f13570_D])).
fof(f13570_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl491),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl491])])).
fof(f12715,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f492])).
fof(f12716,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f492])).
fof(f13564,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl487),
  inference(cnf_transformation,[],[f13564_D])).
fof(f13564_D,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl487),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl487])])).
fof(f12717,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f492])).
fof(f13566,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl489),
  inference(cnf_transformation,[],[f13566_D])).
fof(f13566_D,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl489),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl489])])).
fof(f12720,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f492])).
fof(f12721,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f492])).
fof(f12722,plain,(
  ~sP4('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f492])).
fof(f13560,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl485),
  inference(cnf_transformation,[],[f13560_D])).
fof(f13560_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl485),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl485])])).
fof(f12726,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f496])).
fof(f12727,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f496])).
fof(f13554,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl481),
  inference(cnf_transformation,[],[f13554_D])).
fof(f13554_D,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl481),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl481])])).
fof(f12728,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f496])).
fof(f13556,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl483),
  inference(cnf_transformation,[],[f13556_D])).
fof(f13556_D,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl483),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl483])])).
fof(f12731,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f496])).
fof(f12732,plain,(
  ~sP3('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f496])).
fof(f12733,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f496])).
fof(f13550,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl478),
  inference(cnf_transformation,[],[f13550_D])).
fof(f13550_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f12741,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f1135,f11732,f505])).
fof(f13548,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl476),
  inference(cnf_transformation,[],[f13548_D])).
fof(f13548_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f12742,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f1740,f11732,f505])).
fof(f13546,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl474),
  inference(cnf_transformation,[],[f13546_D])).
fof(f13546_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f12748,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f1135,f11732,f505])).
fof(f13544,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl472),
  inference(cnf_transformation,[],[f13544_D])).
fof(f13544_D,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f12749,plain,(
  sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f1740,f11732,f505])).
fof(f12754,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | ($spl29 | $spl362)),
  inference(unit_resulting_resolution,[],[f1003,f11732,f506])).
fof(f13540,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl469),
  inference(cnf_transformation,[],[f13540_D])).
fof(f13540_D,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl469),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl469])])).
fof(f12756,plain,(
  ~sP5('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | ($spl29 | $spl362)),
  inference(unit_resulting_resolution,[],[f1003,f11732,f506])).
fof(f13538,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl466),
  inference(cnf_transformation,[],[f13538_D])).
fof(f13538_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f12760,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f704,f11732,f511])).
fof(f12761,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f704,f11732,f511])).
fof(f13532,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl460),
  inference(cnf_transformation,[],[f13532_D])).
fof(f13532_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f12762,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f704,f11732,f511])).
fof(f13536,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl464),
  inference(cnf_transformation,[],[f13536_D])).
fof(f13536_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f12764,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1309,f11732,f511])).
fof(f12765,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1309,f11732,f511])).
fof(f13528,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl456),
  inference(cnf_transformation,[],[f13528_D])).
fof(f13528_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f12766,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1309,f11732,f511])).
fof(f13534,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl462),
  inference(cnf_transformation,[],[f13534_D])).
fof(f13534_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f12768,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f704,f11732,f511])).
fof(f12769,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f704,f11732,f511])).
fof(f12770,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f704,f11732,f511])).
fof(f13530,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl458),
  inference(cnf_transformation,[],[f13530_D])).
fof(f13530_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f12772,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f345,f1309,f11732,f511])).
fof(f12773,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f344,f1309,f11732,f511])).
fof(f12774,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1309,f11732,f511])).
fof(f13522,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl455),
  inference(cnf_transformation,[],[f13522_D])).
fof(f13522_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl455),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl455])])).
fof(f12780,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f515])).
fof(f12781,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f515])).
fof(f13516,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl451),
  inference(cnf_transformation,[],[f13516_D])).
fof(f13516_D,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl451),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl451])])).
fof(f12782,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f515])).
fof(f13518,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl453),
  inference(cnf_transformation,[],[f13518_D])).
fof(f13518_D,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl453),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl453])])).
fof(f12785,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f515])).
fof(f12786,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f515])).
fof(f12787,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f515])).
fof(f13512,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl448),
  inference(cnf_transformation,[],[f13512_D])).
fof(f13512_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f12791,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12792,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f13508,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl444),
  inference(cnf_transformation,[],[f13508_D])).
fof(f13508_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f12793,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12794,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12795,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12796,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f13504,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl440),
  inference(cnf_transformation,[],[f13504_D])).
fof(f13504_D,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f12797,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12798,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f13500,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl436),
  inference(cnf_transformation,[],[f13500_D])).
fof(f13500_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f12799,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13510,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl446),
  inference(cnf_transformation,[],[f13510_D])).
fof(f13510_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f12800,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12801,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f12802,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12803,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12804,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12805,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f13502,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl438),
  inference(cnf_transformation,[],[f13502_D])).
fof(f13502_D,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f12806,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12807,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f12808,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13506,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl442),
  inference(cnf_transformation,[],[f13506_D])).
fof(f13506_D,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f12809,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f345,f11732,f521])).
fof(f12810,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f345,f11732,f521])).
fof(f12811,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f345,f11732,f521])).
fof(f12812,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f344,f11732,f521])).
fof(f12813,plain,(
  sP7('T_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f344,f11732,f521])).
fof(f12814,plain,(
  sP7('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f344,f11732,f521])).
fof(f12815,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f11732,f11732,f521])).
fof(f12816,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f11732,f11732,f521])).
fof(f12817,plain,(
  sP7('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f11732,f11732,f521])).
fof(f13498,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl435),
  inference(cnf_transformation,[],[f13498_D])).
fof(f13498_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl435),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl435])])).
fof(f12820,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f524])).
fof(f12821,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f524])).
fof(f13494,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl431),
  inference(cnf_transformation,[],[f13494_D])).
fof(f13494_D,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl431),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl431])])).
fof(f12822,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f524])).
fof(f13496,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl433),
  inference(cnf_transformation,[],[f13496_D])).
fof(f13496_D,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f12823,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f524])).
fof(f12824,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f524])).
fof(f12825,plain,(
  ~sP7('F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f524])).
fof(f13488,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl429),
  inference(cnf_transformation,[],[f13488_D])).
fof(f13488_D,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl429),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl429])])).
fof(f12840,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f532])).
fof(f13482,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl427),
  inference(cnf_transformation,[],[f13482_D])).
fof(f13482_D,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl427),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl427])])).
fof(f12863,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f539])).
fof(f13472,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl425),
  inference(cnf_transformation,[],[f13472_D])).
fof(f13472_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl425),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl425])])).
fof(f12893,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f549])).
fof(f12894,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f549])).
fof(f13466,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl421),
  inference(cnf_transformation,[],[f13466_D])).
fof(f13466_D,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl421),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl421])])).
fof(f12895,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f549])).
fof(f13468,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl423),
  inference(cnf_transformation,[],[f13468_D])).
fof(f13468_D,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl423),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl423])])).
fof(f12898,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f549])).
fof(f12899,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f549])).
fof(f12900,plain,(
  ~sP10('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f549])).
fof(f13462,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl418),
  inference(cnf_transformation,[],[f13462_D])).
fof(f13462_D,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f12904,plain,(
  sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f553])).
fof(f13458,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl417),
  inference(cnf_transformation,[],[f13458_D])).
fof(f13458_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl417),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl417])])).
fof(f12914,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f556])).
fof(f12915,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f556])).
fof(f13454,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl413),
  inference(cnf_transformation,[],[f13454_D])).
fof(f13454_D,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f12916,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f556])).
fof(f13456,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl415),
  inference(cnf_transformation,[],[f13456_D])).
fof(f13456_D,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl415),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl415])])).
fof(f12917,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f556])).
fof(f12918,plain,(
  ~sP11('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f556])).
fof(f12919,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f556])).
fof(f12925,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f11732,f560])).
fof(f12926,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f11732,f560])).
fof(f12933,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f418,f11732,f562])).
fof(f12934,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f11732,f562])).
fof(f12935,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f11732,f562])).
fof(f12936,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f11732,f562])).
fof(f12937,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f11732,f562])).
fof(f12938,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f11732,f562])).
fof(f12939,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f11732,f562])).
fof(f12965,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f418,f11732,f562])).
fof(f12979,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f421,f11732,f562])).
fof(f12993,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f422,f11732,f562])).
fof(f13005,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f11732,f564])).
fof(f13006,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f11732,f564])).
fof(f13013,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f11732,f565])).
fof(f13014,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f11732,f565])).
fof(f13015,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f11732,f565])).
fof(f13027,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f11732,f565])).
fof(f13041,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f418,f11732,f566])).
fof(f13055,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f421,f11732,f566])).
fof(f13069,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f422,f11732,f566])).
fof(f13081,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f418,f11732,f566])).
fof(f13082,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1135,f11732,f566])).
fof(f13083,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl102 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1135,f11732,f566])).
fof(f13084,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f345,f1740,f11732,f566])).
fof(f13085,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl208 | $spl362)),
  inference(unit_resulting_resolution,[],[f352,f344,f1740,f11732,f566])).
fof(f13086,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f11732,f566])).
fof(f13087,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f11732,f566])).
fof(f13304,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl411),
  inference(cnf_transformation,[],[f13304_D])).
fof(f13304_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl411),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl411])])).
fof(f13111,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f345,f11732,f567])).
fof(f13112,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f344,f11732,f567])).
fof(f13298,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl407),
  inference(cnf_transformation,[],[f13298_D])).
fof(f13298_D,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl407),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl407])])).
fof(f13113,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f11732,f11732,f567])).
fof(f13300,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl409),
  inference(cnf_transformation,[],[f13300_D])).
fof(f13300_D,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl409),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl409])])).
fof(f13116,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f567])).
fof(f13117,plain,(
  ~sP15('T_0','F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f567])).
fof(f13118,plain,(
  ~sP15('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f567])).
fof(f13294,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl405),
  inference(cnf_transformation,[],[f13294_D])).
fof(f13294_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f13122,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f571])).
fof(f13123,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f571])).
fof(f13288,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl401),
  inference(cnf_transformation,[],[f13288_D])).
fof(f13288_D,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl401),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl401])])).
fof(f13124,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f571])).
fof(f13290,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl403),
  inference(cnf_transformation,[],[f13290_D])).
fof(f13290_D,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl403),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl403])])).
fof(f13127,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f571])).
fof(f13128,plain,(
  ~sP14('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f571])).
fof(f13129,plain,(
  ~sP14('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f571])).
fof(f13284,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl399),
  inference(cnf_transformation,[],[f13284_D])).
fof(f13284_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f13133,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f575])).
fof(f13134,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f575])).
fof(f13278,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl395),
  inference(cnf_transformation,[],[f13278_D])).
fof(f13278_D,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f13135,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f575])).
fof(f13280,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl397),
  inference(cnf_transformation,[],[f13280_D])).
fof(f13280_D,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl397),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl397])])).
fof(f13138,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f345,f352,f11732,f575])).
fof(f13139,plain,(
  ~sP13('T_0','|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f344,f352,f11732,f575])).
fof(f13140,plain,(
  ~sP13('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f11732,f352,f11732,f575])).
fof(f13274,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl393),
  inference(cnf_transformation,[],[f13274_D])).
fof(f13274_D,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl393),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl393])])).
fof(f13144,plain,(
  ~sP12('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f579])).
fof(f13268,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(forward_demodulation,[],[f13151,f422])).
fof(f13151,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')))) ) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f587])).
fof(f13176,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl362),
  inference(unit_resulting_resolution,[],[f352,f352,f11732,f598])).
fof(f13237,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl389),
  inference(cnf_transformation,[],[f13237_D])).
fof(f13237_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f13180,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f930])).
fof(f13235,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl387),
  inference(cnf_transformation,[],[f13235_D])).
fof(f13235_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl387),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl387])])).
fof(f13181,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl2 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f935])).
fof(f13233,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl385),
  inference(cnf_transformation,[],[f13233_D])).
fof(f13233_D,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f13185,plain,(
  ~sP1('|dollar|exists_1'(s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1535])).
fof(f13231,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | $spl383),
  inference(cnf_transformation,[],[f13231_D])).
fof(f13231_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f13186,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(sK17,bool),sK21))) | ($spl4 | $spl362)),
  inference(unit_resulting_resolution,[],[f11732,f1540])).
fof(f12610,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) | $spl372),
  inference(cnf_transformation,[],[f12610_D])).
fof(f12610_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f12620,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) | $spl4),
  inference(forward_demodulation,[],[f12261,f418])).
fof(f12261,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK19)))))) | $spl4),
  inference(resolution,[],[f1208,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X13] : (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X13)))) => p(s(bool,happ(s(fun(X6,bool),X11),s(X6,'const.min.@_1'(s(fun(X6,bool),X11)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.bool.SELECT_AX')).
fof(f12618,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl379),
  inference(cnf_transformation,[],[f12618_D])).
fof(f12618_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl379),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl379])])).
fof(f12619,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0))))) ) | $spl380),
  inference(cnf_transformation,[],[f12619_D])).
fof(f12619_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0))))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f12260,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))) ) | $spl4),
  inference(resolution,[],[f1208,f474])).
fof(f474,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f12612,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl374),
  inference(cnf_transformation,[],[f12612_D])).
fof(f12612_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f12259,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(resolution,[],[f1208,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7,X6] : ! [X8,X11] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11)))) <=> ! [X20] : (p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X20)))) => ? [X21] : p(s(bool,happ(s(fun(X7,bool),X11),s(X7,happ(s(fun(X6,X7),X8),s(X6,X21))))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f12614,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl377),
  inference(cnf_transformation,[],[f12614_D])).
fof(f12614_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f12258,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(resolution,[],[f1208,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33])],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X7,X6] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X22),s(fun(X7,bool),X23)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X7,bool),X23),s(X7,happ(s(fun(X6,X7),X22),s(X6,X21)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f11749,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(unit_resulting_resolution,[],[f1208,f477])).
fof(f11750,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(unit_resulting_resolution,[],[f1208,f470])).
fof(f12608,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK19)))) | $spl4),
  inference(forward_demodulation,[],[f11751,f418])).
fof(f11751,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK19),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK19)))))) | $spl4),
  inference(unit_resulting_resolution,[],[f1208,f423])).
fof(f11732,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) | $spl362),
  inference(cnf_transformation,[],[f11732_D])).
fof(f11732_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f11742,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) | $spl2),
  inference(forward_demodulation,[],[f11383,f418])).
fof(f11383,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK21)))))) | $spl2),
  inference(resolution,[],[f605,f423])).
fof(f11740,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl369),
  inference(cnf_transformation,[],[f11740_D])).
fof(f11740_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f11741,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,X0))))) ) | $spl370),
  inference(cnf_transformation,[],[f11741_D])).
fof(f11741_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,X0))))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f11382,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,X0)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))) ) | $spl2),
  inference(resolution,[],[f605,f474])).
fof(f11734,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl364),
  inference(cnf_transformation,[],[f11734_D])).
fof(f11734_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f11381,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl2),
  inference(resolution,[],[f605,f470])).
fof(f11736,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl367),
  inference(cnf_transformation,[],[f11736_D])).
fof(f11736_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f11380,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl2),
  inference(resolution,[],[f605,f477])).
fof(f10871,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f477])).
fof(f10872,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f470])).
fof(f11730,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK17,bool),sK21)))) | $spl2),
  inference(forward_demodulation,[],[f10873,f418])).
fof(f10873,plain,(
  p(s(bool,happ(s(fun(sK17,bool),sK21),s(sK17,'const.min.@_1'(s(fun(sK17,bool),sK21)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f423])).
fof(f10283,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl360),
  inference(cnf_transformation,[],[f10283_D])).
fof(f10283_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f6929,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl208),
  inference(unit_resulting_resolution,[],[f345,f1740,f505])).
fof(f6930,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl208),
  inference(unit_resulting_resolution,[],[f344,f1740,f505])).
fof(f8983,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl358),
  inference(cnf_transformation,[],[f8983_D])).
fof(f8983_D,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f7571,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(unit_resulting_resolution,[],[f345,f1740,f505])).
fof(f7572,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(unit_resulting_resolution,[],[f344,f1740,f505])).
fof(f6381,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl356),
  inference(cnf_transformation,[],[f6381_D])).
fof(f6381_D,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f3970,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl102),
  inference(unit_resulting_resolution,[],[f345,f1135,f505])).
fof(f3971,plain,(
  sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl102),
  inference(unit_resulting_resolution,[],[f344,f1135,f505])).
fof(f5484,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl354),
  inference(cnf_transformation,[],[f5484_D])).
fof(f5484_D,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f4451,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl102),
  inference(unit_resulting_resolution,[],[f345,f1135,f505])).
fof(f4452,plain,(
  sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl102),
  inference(unit_resulting_resolution,[],[f344,f1135,f505])).
fof(f3780,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl353),
  inference(cnf_transformation,[],[f3780_D])).
fof(f3780_D,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f3776,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl191),
  inference(unit_resulting_resolution,[],[f1722,f487])).
fof(f3770,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl351),
  inference(cnf_transformation,[],[f3770_D])).
fof(f3770_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f3766,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl187),
  inference(unit_resulting_resolution,[],[f1718,f486])).
fof(f3760,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl349),
  inference(cnf_transformation,[],[f3760_D])).
fof(f3760_D,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f3756,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl185),
  inference(unit_resulting_resolution,[],[f1716,f486])).
fof(f3733,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl346),
  inference(cnf_transformation,[],[f3733_D])).
fof(f3733_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f3728,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl170),
  inference(resolution,[],[f1702,f508])).
fof(f3722,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1702,f508])).
fof(f3705,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl344),
  inference(cnf_transformation,[],[f3705_D])).
fof(f3705_D,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f3700,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl164),
  inference(resolution,[],[f1696,f508])).
fof(f3694,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl164),
  inference(unit_resulting_resolution,[],[f1696,f508])).
fof(f3642,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl342),
  inference(cnf_transformation,[],[f3642_D])).
fof(f3642_D,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f3640,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(resolution,[],[f1672,f544])).
fof(f3634,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(unit_resulting_resolution,[],[f1672,f544])).
fof(f3209,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl340),
  inference(cnf_transformation,[],[f3209_D])).
fof(f3209_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f3205,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f1278,f1285,f1286,f489])).
fof(f3206,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f1278,f1285,f1286,f489])).
fof(f3198,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl338),
  inference(cnf_transformation,[],[f3198_D])).
fof(f3198_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f3194,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f345,f684,f675,f1285,f489])).
fof(f3195,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f344,f684,f675,f1285,f489])).
fof(f3186,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl336),
  inference(cnf_transformation,[],[f3186_D])).
fof(f3186_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f3179,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f345,f683,f682,f1278,f489])).
fof(f3180,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f344,f683,f682,f1278,f489])).
fof(f3167,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl335),
  inference(cnf_transformation,[],[f3167_D])).
fof(f3167_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f3163,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl85),
  inference(unit_resulting_resolution,[],[f1117,f487])).
fof(f3152,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl333),
  inference(cnf_transformation,[],[f3152_D])).
fof(f3152_D,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f3148,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl83),
  inference(unit_resulting_resolution,[],[f1115,f487])).
fof(f3142,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl331),
  inference(cnf_transformation,[],[f3142_D])).
fof(f3142_D,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f3138,plain,(
  ~sP5('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl79),
  inference(unit_resulting_resolution,[],[f1111,f486])).
fof(f3132,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl329),
  inference(cnf_transformation,[],[f3132_D])).
fof(f3132_D,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f3128,plain,(
  ~sP5('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl77),
  inference(unit_resulting_resolution,[],[f1109,f486])).
fof(f3105,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl326),
  inference(cnf_transformation,[],[f3105_D])).
fof(f3105_D,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f3100,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl66),
  inference(resolution,[],[f1099,f508])).
fof(f3094,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl66),
  inference(unit_resulting_resolution,[],[f1099,f508])).
fof(f3077,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl324),
  inference(cnf_transformation,[],[f3077_D])).
fof(f3077_D,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3072,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl60),
  inference(resolution,[],[f1093,f508])).
fof(f3066,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl60),
  inference(unit_resulting_resolution,[],[f1093,f508])).
fof(f3010,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl322),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f3008,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl36),
  inference(resolution,[],[f1069,f544])).
fof(f3002,plain,(
  sP10('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl36),
  inference(unit_resulting_resolution,[],[f1069,f544])).
fof(f2831,plain,(
  sP4('F_0','T_0','F_0') | $spl316),
  inference(cnf_transformation,[],[f2831_D])).
fof(f2831_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f2827,plain,(
  sP4('F_0','T_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f487])).
fof(f2833,plain,(
  sP3('F_0','T_0','F_0') | $spl318),
  inference(cnf_transformation,[],[f2833_D])).
fof(f2833_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f2826,plain,(
  sP3('F_0','T_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f486])).
fof(f2835,plain,(
  sP2('T_0','F_0','F_0') | $spl320),
  inference(cnf_transformation,[],[f2835_D])).
fof(f2835_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f2825,plain,(
  sP2('T_0','F_0','F_0') | $spl110),
  inference(resolution,[],[f1207,f485])).
fof(f2821,plain,(
  sP2('T_0','F_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f485])).
fof(f2822,plain,(
  sP3('F_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f486])).
fof(f2823,plain,(
  sP4('F_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f1207,f487])).
fof(f2618,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl314),
  inference(cnf_transformation,[],[f2618_D])).
fof(f2618_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f2614,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f675,f682,f683,f489])).
fof(f2615,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f675,f682,f683,f489])).
fof(f2608,plain,(
  sP4('F_0','F_0','T_0') | $spl308),
  inference(cnf_transformation,[],[f2608_D])).
fof(f2608_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f2605,plain,(
  sP4('F_0','F_0','T_0') | $spl106),
  inference(resolution,[],[f1203,f487])).
fof(f2610,plain,(
  sP3('T_0','F_0','F_0') | $spl310),
  inference(cnf_transformation,[],[f2610_D])).
fof(f2610_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f2604,plain,(
  sP3('T_0','F_0','F_0') | $spl106),
  inference(resolution,[],[f1203,f486])).
fof(f2612,plain,(
  sP2('F_0','T_0','F_0') | $spl312),
  inference(cnf_transformation,[],[f2612_D])).
fof(f2612_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f2603,plain,(
  sP2('F_0','T_0','F_0') | $spl106),
  inference(resolution,[],[f1203,f485])).
fof(f2598,plain,(
  sP2('F_0','T_0','F_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f485])).
fof(f2599,plain,(
  sP3('T_0','F_0','F_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f486])).
fof(f2600,plain,(
  sP4('F_0','F_0','T_0') | $spl106),
  inference(unit_resulting_resolution,[],[f1203,f487])).
fof(f2442,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl304),
  inference(cnf_transformation,[],[f2442_D])).
fof(f2442_D,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f2201,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(resolution,[],[f600,f368])).
fof(f2200,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(resolution,[],[f600,f366])).
fof(f2484,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) | $spl307),
  inference(cnf_transformation,[],[f2484_D])).
fof(f2484_D,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f1787,plain,(
  ~p(s(bool,happ(s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))),s(sK17,happ(s(fun(sK16,sK17),sK18),s(sK16,sK24('_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)),sK18,sK16,sK17))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK24(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f2241,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl211),
  inference(cnf_transformation,[],[f2241_D])).
fof(f2241_D,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f1790,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f361])).
fof(f1791,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f363])).
fof(f1792,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f366])).
fof(f1793,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f368])).
fof(f1822,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f431])).
fof(f1824,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f431])).
fof(f2440,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl303),
  inference(cnf_transformation,[],[f2440_D])).
fof(f2440_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f1826,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f437])).
fof(f2436,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl299),
  inference(cnf_transformation,[],[f2436_D])).
fof(f2436_D,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f1827,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f437])).
fof(f2438,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl301),
  inference(cnf_transformation,[],[f2438_D])).
fof(f2438_D,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f1828,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f437])).
fof(f1829,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f437])).
fof(f1830,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f441])).
fof(f1831,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f441])).
fof(f1832,plain,(
  ~sP0('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f441])).
fof(f1833,plain,(
  ~sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f441])).
fof(f2428,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl291),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f1834,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f452])).
fof(f2426,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl289),
  inference(cnf_transformation,[],[f2426_D])).
fof(f2426_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f1835,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f452])).
fof(f2430,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl293),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f1836,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f452])).
fof(f1837,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f452])).
fof(f2432,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl294),
  inference(cnf_transformation,[],[f2432_D])).
fof(f2432_D,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f1838,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f454])).
fof(f1839,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f454])).
fof(f2434,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl296),
  inference(cnf_transformation,[],[f2434_D])).
fof(f2434_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f1840,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f457])).
fof(f1841,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f457])).
fof(f1842,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f458])).
fof(f1843,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f458])).
fof(f1844,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f459])).
fof(f1845,plain,(
  sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f459])).
fof(f1846,plain,(
  ~sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f461])).
fof(f1847,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f461])).
fof(f1848,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f461])).
fof(f1849,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f461])).
fof(f2424,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl287),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f1850,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f488])).
fof(f1851,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f488])).
fof(f1852,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f488])).
fof(f1853,plain,(
  ~sP5('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f488])).
fof(f2422,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl285),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f1854,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f492])).
fof(f1855,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f492])).
fof(f1856,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f492])).
fof(f1857,plain,(
  ~sP4('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f492])).
fof(f2420,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl283),
  inference(cnf_transformation,[],[f2420_D])).
fof(f2420_D,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f1860,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f496])).
fof(f1861,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f496])).
fof(f1862,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f496])).
fof(f1863,plain,(
  ~sP3('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f496])).
fof(f2418,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl281),
  inference(cnf_transformation,[],[f2418_D])).
fof(f2418_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f1866,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f2414,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl277),
  inference(cnf_transformation,[],[f2414_D])).
fof(f2414_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f1867,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f2410,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl273),
  inference(cnf_transformation,[],[f2410_D])).
fof(f2410_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f1868,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f2406,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl269),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f1869,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2416,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl279),
  inference(cnf_transformation,[],[f2416_D])).
fof(f2416_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f1870,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f1871,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f2408,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl271),
  inference(cnf_transformation,[],[f2408_D])).
fof(f2408_D,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl271),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl271])])).
fof(f1872,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f1873,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2412,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl275),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f1874,plain,(
  ~sP2('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f500])).
fof(f1875,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f500])).
fof(f1876,plain,(
  ~sP2('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f500])).
fof(f1877,plain,(
  ~sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f500])).
fof(f2404,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') | $spl266),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f1878,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','F_0') | ($spl1 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f600,f504])).
fof(f2402,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl264),
  inference(cnf_transformation,[],[f2402_D])).
fof(f2402_D,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f1879,plain,(
  sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | ($spl1 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f600,f504])).
fof(f2400,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl263),
  inference(cnf_transformation,[],[f2400_D])).
fof(f2400_D,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl263),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl263])])).
fof(f1894,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f515])).
fof(f1895,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f515])).
fof(f1896,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f515])).
fof(f1897,plain,(
  ~sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f515])).
fof(f2398,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl261),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f1912,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f524])).
fof(f1913,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f524])).
fof(f1914,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f524])).
fof(f1915,plain,(
  ~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f524])).
fof(f2396,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl259),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f1920,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f532])).
fof(f1921,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f532])).
fof(f2392,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl255),
  inference(cnf_transformation,[],[f2392_D])).
fof(f2392_D,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f1922,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f532])).
fof(f1923,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f532])).
fof(f2394,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl257),
  inference(cnf_transformation,[],[f2394_D])).
fof(f2394_D,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f1924,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f532])).
fof(f1925,plain,(
  ~sP8('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f532])).
fof(f1926,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f532])).
fof(f1927,plain,(
  ~sP8('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f532])).
fof(f2390,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl252),
  inference(cnf_transformation,[],[f2390_D])).
fof(f2390_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f1929,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f2382,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl244),
  inference(cnf_transformation,[],[f2382_D])).
fof(f2382_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f1930,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f2386,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl248),
  inference(cnf_transformation,[],[f2386_D])).
fof(f2386_D,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f1931,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f2378,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl240),
  inference(cnf_transformation,[],[f2378_D])).
fof(f2378_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f1932,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2388,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl250),
  inference(cnf_transformation,[],[f2388_D])).
fof(f2388_D,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f1933,plain,(
  sP9('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f2380,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl242),
  inference(cnf_transformation,[],[f2380_D])).
fof(f2380_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f1934,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f1935,plain,(
  sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f1936,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2384,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl246),
  inference(cnf_transformation,[],[f2384_D])).
fof(f2384_D,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f1937,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f352,f600,f536])).
fof(f1938,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f352,f600,f536])).
fof(f1939,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f600,f600,f536])).
fof(f1940,plain,(
  sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f600,f600,f536])).
fof(f2376,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl239),
  inference(cnf_transformation,[],[f2376_D])).
fof(f2376_D,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f1941,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f539])).
fof(f2372,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl235),
  inference(cnf_transformation,[],[f2372_D])).
fof(f2372_D,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f1942,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f539])).
fof(f1943,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f539])).
fof(f1944,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f539])).
fof(f2374,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl237),
  inference(cnf_transformation,[],[f2374_D])).
fof(f2374_D,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f1945,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f539])).
fof(f1946,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f539])).
fof(f1947,plain,(
  ~sP9('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f539])).
fof(f1948,plain,(
  ~sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f539])).
fof(f2370,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl233),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f1959,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f549])).
fof(f1960,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f549])).
fof(f1961,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f549])).
fof(f1962,plain,(
  ~sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f549])).
fof(f2368,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl230),
  inference(cnf_transformation,[],[f2368_D])).
fof(f2368_D,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f1965,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f553])).
fof(f2364,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl226),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f1966,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f600,f553])).
fof(f1967,plain,(
  sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f553])).
fof(f1968,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f600,f553])).
fof(f2366,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl228),
  inference(cnf_transformation,[],[f2366_D])).
fof(f2366_D,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f1969,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f553])).
fof(f1970,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f553])).
fof(f1971,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f553])).
fof(f1972,plain,(
  sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f553])).
fof(f2362,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl225),
  inference(cnf_transformation,[],[f2362_D])).
fof(f2362_D,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f1975,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f556])).
fof(f1976,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f556])).
fof(f1977,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f556])).
fof(f1978,plain,(
  ~sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f556])).
fof(f1985,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f600,f560])).
fof(f1986,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f600,f560])).
fof(f1987,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f421,f600,f560])).
fof(f1988,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f421,f600,f560])).
fof(f1989,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f600,f560])).
fof(f1990,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f600,f560])).
fof(f1991,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f422,f600,f560])).
fof(f1992,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f422,f600,f560])).
fof(f2002,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f421,f600,f560])).
fof(f2004,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f421,f600,f560])).
fof(f2014,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f422,f600,f560])).
fof(f2016,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f422,f600,f560])).
fof(f2025,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f600,f562])).
fof(f2026,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f600,f562])).
fof(f2027,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f600,f562])).
fof(f2028,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f600,f562])).
fof(f2029,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f600,f562])).
fof(f2030,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f600,f562])).
fof(f2031,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f600,f562])).
fof(f2032,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f600,f562])).
fof(f2041,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f600,f564])).
fof(f2042,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f600,f564])).
fof(f2043,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f421,f600,f564])).
fof(f2044,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f421,f600,f564])).
fof(f2045,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f600,f564])).
fof(f2046,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f600,f564])).
fof(f2047,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f422,f600,f564])).
fof(f2048,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f422,f600,f564])).
fof(f2058,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f421,f600,f564])).
fof(f2060,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f421,f600,f564])).
fof(f2070,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f345,f422,f600,f564])).
fof(f2072,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f600,f344,f422,f600,f564])).
fof(f2081,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f345,f600,f565])).
fof(f2082,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f345,f600,f565])).
fof(f2083,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f344,f600,f565])).
fof(f2084,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f344,f600,f565])).
fof(f2095,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f345,f600,f565])).
fof(f2096,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f345,f600,f565])).
fof(f2099,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f344,f600,f565])).
fof(f2100,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f344,f600,f565])).
fof(f2117,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f421,f600,f566])).
fof(f2118,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f421,f600,f566])).
fof(f2119,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f421,f600,f566])).
fof(f2120,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f421,f600,f566])).
fof(f2121,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f422,f600,f566])).
fof(f2122,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f422,f600,f566])).
fof(f2123,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f422,f600,f566])).
fof(f2124,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f422,f600,f566])).
fof(f2272,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl223),
  inference(cnf_transformation,[],[f2272_D])).
fof(f2272_D,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f2133,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f567])).
fof(f2134,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f567])).
fof(f2135,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f567])).
fof(f2136,plain,(
  ~sP15('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f567])).
fof(f2270,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl221),
  inference(cnf_transformation,[],[f2270_D])).
fof(f2270_D,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f2139,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f571])).
fof(f2140,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f571])).
fof(f2141,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f571])).
fof(f2142,plain,(
  ~sP14('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f571])).
fof(f2268,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl219),
  inference(cnf_transformation,[],[f2268_D])).
fof(f2268_D,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f2145,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f345,f600,f575])).
fof(f2146,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f345,f600,f575])).
fof(f2147,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f344,f600,f575])).
fof(f2148,plain,(
  ~sP13('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f344,f600,f575])).
fof(f2266,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl217),
  inference(cnf_transformation,[],[f2266_D])).
fof(f2266_D,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2151,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f579])).
fof(f2152,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f579])).
fof(f2262,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl213),
  inference(cnf_transformation,[],[f2262_D])).
fof(f2262_D,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f2153,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f579])).
fof(f2154,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f579])).
fof(f2264,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl215),
  inference(cnf_transformation,[],[f2264_D])).
fof(f2264_D,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f2155,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f352,f600,f579])).
fof(f2156,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f352,f600,f579])).
fof(f2157,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f579])).
fof(f2158,plain,(
  ~sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f579])).
fof(f2182,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f345,f600,f598])).
fof(f2183,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f352,f344,f600,f598])).
fof(f2190,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f345,f600,f600,f598])).
fof(f2191,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))) | $spl1),
  inference(unit_resulting_resolution,[],[f344,f600,f600,f598])).
fof(f1740,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl208),
  inference(cnf_transformation,[],[f1740_D])).
fof(f1740_D,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f1781,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(global_subsumption,[],[f343,f342,f341,f604,f1519,f1566,f1515,f1514,f1513,f1512,f1511,f1510,f1509,f1508,f1507,f1506,f1505,f1504,f1503,f1502,f1501,f1500,f1499,f1498,f1497,f1496,f1495,f1494,f1493,f1492,f1479,f1478,f1469,f1458,f1447,f1438,f1437,f1436,f1431,f1430,f1421,f1410,f1395,f1394,f1389,f1388,f1387,f1386,f1385,f1384,f1383,f1382,f1381,f1380,f1379,f1376,f1374,f1373,f1372,f1371,f1370,f1369,f1368,f1367,f1366,f1364,f1363,f1362,f1361,f1360,f1357,f1355,f1354,f1351,f1349,f1348,f1347,f1344,f1343,f1342,f1341,f1340,f1339,f1338,f1337,f1336,f1335,f1334,f1333,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1325,f1324,f1323,f1322,f1321,f1320,f1319,f1318,f1317,f1316,f1315,f1314,f1313,f1312,f1311,f1310,f1309,f1308,f1307,f1306,f1305,f1304,f1303,f1298,f1297,f1296,f1295,f1294,f1293,f1292,f1291,f1290,f1289,f1288,f1287,f1286,f1285,f1284,f1283,f1282,f1281,f1280,f1279,f1278,f1277,f1276,f1275,f1274,f1273,f1272,f1271,f1270,f1269,f1268,f1267,f1266,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1256,f1255,f1254,f1253,f1252,f1251,f1250,f1249,f1248,f1244,f1243,f1241,f1240,f1742,f1744,f1746,f1212,f1211,f1210,f1209,f1208,f1522,f1523,f1530,f1535,f1536,f1537,f1540,f1547,f1550,f1551,f1553,f1554,f1555,f1561])).
fof(f1561,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(superposition,[],[f604,f354])).
fof(f1555,plain,(
  ( ! [X39,X40] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X39,X40) | ~p(s(bool,X39))) ) | $spl4),
  inference(resolution,[],[f604,f555])).
fof(f1554,plain,(
  ( ! [X37,X38] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X38,X37) | p(s(bool,X37))) ) | $spl4),
  inference(resolution,[],[f604,f554])).
fof(f1551,plain,(
  ( ! [X31,X32] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X31,X32) | ~p(s(bool,X31))) ) | $spl4),
  inference(resolution,[],[f604,f538])).
fof(f1550,plain,(
  ( ! [X30,X29] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X30,X29) | ~p(s(bool,X29))) ) | $spl4),
  inference(resolution,[],[f604,f537])).
fof(f1530,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(resolution,[],[f604,f434])).
fof(f1746,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X5)))) ) | $spl4),
  inference(forward_demodulation,[],[f1745,f422])).
fof(f1745,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,X1),s(X0,X5)))) ) | $spl4),
  inference(forward_demodulation,[],[f1237,f422])).
fof(f1237,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f422,f604,f428])).
fof(f1744,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(forward_demodulation,[],[f1743,f422])).
fof(f1743,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(forward_demodulation,[],[f1238,f426])).
fof(f1238,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f422,f604,f428])).
fof(f1742,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl4),
  inference(forward_demodulation,[],[f1741,f422])).
fof(f1741,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl4),
  inference(forward_demodulation,[],[f1239,f591])).
fof(f1239,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f422,f604,f428])).
fof(f1248,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f443])).
fof(f1249,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f444])).
fof(f1288,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f503])).
fof(f1289,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f604,f505])).
fof(f1290,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f421,f604,f505])).
fof(f1291,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f604,f505])).
fof(f1292,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f604,f505])).
fof(f1293,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f604,f505])).
fof(f1294,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f418,f604,f505])).
fof(f1295,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f421,f604,f505])).
fof(f1296,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f422,f604,f505])).
fof(f1297,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f426,f604,f505])).
fof(f1298,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f591,f604,f505])).
fof(f1337,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f522])).
fof(f1338,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f523])).
fof(f1348,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f533])).
fof(f1349,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f534])).
fof(f1351,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f537])).
fof(f1354,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f537])).
fof(f1355,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f537])).
fof(f1357,plain,(
  ( ! [X0] : (sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f538])).
fof(f1360,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f538])).
fof(f1361,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f538])).
fof(f1363,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f540])).
fof(f1364,plain,(
  ( ! [X0] : (~sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f541])).
fof(f1372,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f551])).
fof(f1374,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f554])).
fof(f1376,plain,(
  ( ! [X0] : (sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f555])).
fof(f1379,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f555])).
fof(f1380,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f555])).
fof(f1387,plain,(
  ( ! [X0] : (~sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f558])).
fof(f1505,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f572])).
fof(f1512,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f576])).
fof(f1514,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f580])).
fof(f1515,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f581])).
fof(f1523,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(resolution,[],[f604,f364])).
fof(f1522,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(resolution,[],[f604,f362])).
fof(f1209,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f362])).
fof(f1210,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f364])).
fof(f1565,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl113),
  inference(cnf_transformation,[],[f1565_D])).
fof(f1565_D,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1211,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f365])).
fof(f1212,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f367])).
fof(f1240,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f434])).
fof(f1241,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f434])).
fof(f1243,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f434])).
fof(f1244,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f434])).
fof(f1732,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl201),
  inference(cnf_transformation,[],[f1732_D])).
fof(f1732_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f1250,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f453])).
fof(f1251,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f453])).
fof(f1730,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl199),
  inference(cnf_transformation,[],[f1730_D])).
fof(f1730_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f1252,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f453])).
fof(f1734,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl203),
  inference(cnf_transformation,[],[f1734_D])).
fof(f1734_D,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f1253,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f453])).
fof(f1254,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f453])).
fof(f1255,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f453])).
fof(f1736,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl204),
  inference(cnf_transformation,[],[f1736_D])).
fof(f1736_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f1256,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f454])).
fof(f1738,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl206),
  inference(cnf_transformation,[],[f1738_D])).
fof(f1738_D,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f1257,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f457])).
fof(f1258,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f458])).
fof(f1259,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f459])).
fof(f1260,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f460])).
fof(f1261,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f460])).
fof(f1262,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f460])).
fof(f1263,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f460])).
fof(f1264,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f460])).
fof(f1265,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f460])).
fof(f1712,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl181),
  inference(cnf_transformation,[],[f1712_D])).
fof(f1712_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f1266,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f488])).
fof(f1267,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f488])).
fof(f1726,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl195),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f1268,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f488])).
fof(f1728,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl197),
  inference(cnf_transformation,[],[f1728_D])).
fof(f1728_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f1269,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f488])).
fof(f1270,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f488])).
fof(f1271,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f488])).
fof(f1724,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl193),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f1272,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f492])).
fof(f1273,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f492])).
fof(f1720,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl189),
  inference(cnf_transformation,[],[f1720_D])).
fof(f1720_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f1274,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f492])).
fof(f1722,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl191),
  inference(cnf_transformation,[],[f1722_D])).
fof(f1722_D,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f1275,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f492])).
fof(f1276,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f492])).
fof(f1277,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f492])).
fof(f1718,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl187),
  inference(cnf_transformation,[],[f1718_D])).
fof(f1718_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f1279,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f496])).
fof(f1280,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f496])).
fof(f1714,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl183),
  inference(cnf_transformation,[],[f1714_D])).
fof(f1714_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f1281,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f496])).
fof(f1716,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl185),
  inference(cnf_transformation,[],[f1716_D])).
fof(f1716_D,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f1282,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f496])).
fof(f1283,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f496])).
fof(f1284,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f496])).
fof(f1299,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | ($spl4 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f604,f506])).
fof(f1710,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl179),
  inference(cnf_transformation,[],[f1710_D])).
fof(f1710_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f1300,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | ($spl4 | $spl29)),
  inference(unit_resulting_resolution,[],[f1003,f604,f506])).
fof(f1708,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl177),
  inference(cnf_transformation,[],[f1708_D])).
fof(f1708_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f1303,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f515])).
fof(f1304,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f515])).
fof(f1704,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl173),
  inference(cnf_transformation,[],[f1704_D])).
fof(f1704_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f1305,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f515])).
fof(f1706,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl175),
  inference(cnf_transformation,[],[f1706_D])).
fof(f1706_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f1306,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f515])).
fof(f1307,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f515])).
fof(f1308,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f515])).
fof(f1702,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl170),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1310,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1311,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1698,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl166),
  inference(cnf_transformation,[],[f1698_D])).
fof(f1698_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1312,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1313,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1314,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1315,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1694,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl162),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1316,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1317,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1690,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl158),
  inference(cnf_transformation,[],[f1690_D])).
fof(f1690_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f1318,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1700,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl168),
  inference(cnf_transformation,[],[f1700_D])).
fof(f1700_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1319,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1320,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1321,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1322,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1323,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1324,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1692,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl160),
  inference(cnf_transformation,[],[f1692_D])).
fof(f1692_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1325,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1326,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1327,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1696,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl164),
  inference(cnf_transformation,[],[f1696_D])).
fof(f1696_D,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f1328,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f345,f604,f521])).
fof(f1329,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f345,f604,f521])).
fof(f1330,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f345,f604,f521])).
fof(f1331,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f344,f604,f521])).
fof(f1332,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f344,f604,f521])).
fof(f1333,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f344,f604,f521])).
fof(f1334,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f604,f604,f521])).
fof(f1335,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f604,f604,f521])).
fof(f1336,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f604,f604,f521])).
fof(f1688,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl157),
  inference(cnf_transformation,[],[f1688_D])).
fof(f1688_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f1339,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f524])).
fof(f1340,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f524])).
fof(f1684,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl153),
  inference(cnf_transformation,[],[f1684_D])).
fof(f1684_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f1341,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f524])).
fof(f1686,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl155),
  inference(cnf_transformation,[],[f1686_D])).
fof(f1686_D,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f1342,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f524])).
fof(f1343,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f524])).
fof(f1344,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f524])).
fof(f1682,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl151),
  inference(cnf_transformation,[],[f1682_D])).
fof(f1682_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1347,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f532])).
fof(f1680,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl149),
  inference(cnf_transformation,[],[f1680_D])).
fof(f1680_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1362,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f539])).
fof(f1678,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl147),
  inference(cnf_transformation,[],[f1678_D])).
fof(f1678_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f1366,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f549])).
fof(f1367,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f549])).
fof(f1674,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl143),
  inference(cnf_transformation,[],[f1674_D])).
fof(f1674_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1368,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f549])).
fof(f1676,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl145),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f1369,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f549])).
fof(f1370,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f549])).
fof(f1371,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f549])).
fof(f1672,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl140),
  inference(cnf_transformation,[],[f1672_D])).
fof(f1672_D,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f1373,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f553])).
fof(f1670,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl139),
  inference(cnf_transformation,[],[f1670_D])).
fof(f1670_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1381,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f556])).
fof(f1382,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f556])).
fof(f1666,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl135),
  inference(cnf_transformation,[],[f1666_D])).
fof(f1666_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1383,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f556])).
fof(f1668,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl137),
  inference(cnf_transformation,[],[f1668_D])).
fof(f1668_D,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f1384,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f556])).
fof(f1385,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f556])).
fof(f1386,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f556])).
fof(f1388,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f604,f560])).
fof(f1389,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f604,f560])).
fof(f1394,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f604,f562])).
fof(f1395,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f604,f562])).
fof(f1410,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f421,f604,f562])).
fof(f1421,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f422,f604,f562])).
fof(f1430,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f604,f564])).
fof(f1431,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f604,f564])).
fof(f1436,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f604,f565])).
fof(f1437,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f604,f565])).
fof(f1438,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f604,f565])).
fof(f1447,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f604,f565])).
fof(f1458,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f421,f604,f566])).
fof(f1469,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f422,f604,f566])).
fof(f1478,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f604,f566])).
fof(f1479,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f604,f566])).
fof(f1586,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl133),
  inference(cnf_transformation,[],[f1586_D])).
fof(f1586_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f1492,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f345,f604,f567])).
fof(f1493,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f344,f604,f567])).
fof(f1582,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl129),
  inference(cnf_transformation,[],[f1582_D])).
fof(f1582_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f1494,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f604,f604,f567])).
fof(f1584,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl131),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f1495,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f567])).
fof(f1496,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f567])).
fof(f1497,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f567])).
fof(f1580,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl127),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f1499,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f571])).
fof(f1500,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f571])).
fof(f1576,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl123),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f1501,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f571])).
fof(f1578,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl125),
  inference(cnf_transformation,[],[f1578_D])).
fof(f1578_D,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f1502,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f571])).
fof(f1503,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f571])).
fof(f1504,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f571])).
fof(f1574,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl121),
  inference(cnf_transformation,[],[f1574_D])).
fof(f1574_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f1506,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f575])).
fof(f1507,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f575])).
fof(f1570,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl117),
  inference(cnf_transformation,[],[f1570_D])).
fof(f1570_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1508,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f575])).
fof(f1572,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl119),
  inference(cnf_transformation,[],[f1572_D])).
fof(f1572_D,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1509,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f345,f352,f604,f575])).
fof(f1510,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f344,f352,f604,f575])).
fof(f1511,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f604,f352,f604,f575])).
fof(f1568,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl115),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1513,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f579])).
fof(f1566,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(forward_demodulation,[],[f1516,f422])).
fof(f1516,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f587])).
fof(f1519,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))) | $spl4),
  inference(unit_resulting_resolution,[],[f352,f352,f604,f598])).
fof(f1207,plain,(
  sP5('F_0','F_0','T_0') | $spl110),
  inference(cnf_transformation,[],[f1207_D])).
fof(f1207_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1185,plain,(
  sP5('F_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f352,f1003,f504])).
fof(f1205,plain,(
  ~sP5('T_0','F_0','T_0') | $spl109),
  inference(cnf_transformation,[],[f1205_D])).
fof(f1205_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1186,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f345,f1003,f506])).
fof(f1187,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f344,f1003,f506])).
fof(f1203,plain,(
  sP5('F_0','T_0','F_0') | $spl106),
  inference(cnf_transformation,[],[f1203_D])).
fof(f1203_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1193,plain,(
  sP5('F_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f352,f1003,f504])).
fof(f1201,plain,(
  ~sP5('T_0','T_0','F_0') | $spl105),
  inference(cnf_transformation,[],[f1201_D])).
fof(f1201_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1194,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f345,f1003,f506])).
fof(f1195,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f344,f1003,f506])).
fof(f1135,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl102),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1176,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(global_subsumption,[],[f343,f342,f341,f602,f914,f961,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f899,f898,f897,f896,f895,f894,f893,f892,f891,f890,f889,f888,f887,f874,f873,f864,f863,f862,f853,f852,f851,f842,f841,f840,f833,f832,f831,f826,f825,f816,f815,f814,f805,f804,f803,f790,f789,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f771,f769,f768,f767,f766,f765,f764,f763,f762,f761,f759,f758,f757,f756,f755,f752,f750,f749,f746,f744,f743,f742,f739,f738,f737,f736,f735,f734,f733,f732,f731,f730,f729,f728,f727,f726,f725,f724,f723,f722,f721,f720,f719,f718,f717,f716,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f695,f694,f693,f692,f691,f690,f689,f688,f687,f686,f685,f684,f683,f682,f681,f680,f679,f678,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f665,f664,f663,f662,f661,f660,f659,f658,f657,f656,f655,f654,f653,f652,f651,f650,f649,f648,f647,f646,f645,f641,f640,f638,f637,f1137,f1139,f1141,f609,f608,f607,f606,f605,f917,f918,f925,f930,f931,f932,f935,f942,f945,f946,f948,f949,f950,f956])).
fof(f956,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(superposition,[],[f602,f354])).
fof(f950,plain,(
  ( ! [X39,X40] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X39,X40) | ~p(s(bool,X39))) ) | $spl2),
  inference(resolution,[],[f602,f555])).
fof(f949,plain,(
  ( ! [X37,X38] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X38,X37) | p(s(bool,X37))) ) | $spl2),
  inference(resolution,[],[f602,f554])).
fof(f946,plain,(
  ( ! [X31,X32] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X31,X32) | ~p(s(bool,X31))) ) | $spl2),
  inference(resolution,[],[f602,f538])).
fof(f945,plain,(
  ( ! [X30,X29] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X30,X29) | ~p(s(bool,X29))) ) | $spl2),
  inference(resolution,[],[f602,f537])).
fof(f925,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(resolution,[],[f602,f434])).
fof(f1141,plain,(
  ( ! [X2,X0,X5,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X5)))) ) | $spl2),
  inference(forward_demodulation,[],[f1140,f422])).
fof(f1140,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,X1),s(X0,X5)))) ) | $spl2),
  inference(forward_demodulation,[],[f634,f422])).
fof(f634,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X1))),s(X0,X5)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f422,f602,f428])).
fof(f1139,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(forward_demodulation,[],[f1138,f422])).
fof(f1138,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(forward_demodulation,[],[f635,f426])).
fof(f635,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),happ(s(fun(bool,fun(X0,fun(X0,X0))),'const.bool.COND_0'),s(bool,X1))),s(X0,X2))),s(X0,X3))),s(X0,X6)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f422,f602,f428])).
fof(f1137,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl2),
  inference(forward_demodulation,[],[f1136,f422])).
fof(f1136,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X12)))) ) | $spl2),
  inference(forward_demodulation,[],[f636,f591])).
fof(f636,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bool.COND_3'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(fun(X1,X2),X6),s(fun(X1,X3),X7),s(fun(X1,X4),X8))),s(X1,X9))),s(X0,X10))) = s(X0,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),s(X0,happ(s(fun(X4,X0),happ(s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),X5),s(X2,happ(s(fun(X1,X2),X6),s(X1,X9))))),s(X3,happ(s(fun(X1,X3),X7),s(X1,X9))))),s(X4,happ(s(fun(X1,X4),X8),s(X1,X9))))),s(X0,X12)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f422,f602,f428])).
fof(f645,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f443])).
fof(f646,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f444])).
fof(f685,plain,(
  ( ! [X0,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f503])).
fof(f686,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f418,f602,f505])).
fof(f687,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f421,f602,f505])).
fof(f688,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f602,f505])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f602,f505])).
fof(f690,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),'_dst_x0x1x1x1_4'(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(fun(X0,X1),X5),s(fun(X0,X2),X6),s(fun(X0,X3),X7))),s(X0,X8)),happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),happ(s(fun(X1,fun(X2,fun(X3,bool))),X4),s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))))),s(X2,happ(s(fun(X0,X2),X6),s(X0,X8))))),s(X3,happ(s(fun(X0,X3),X7),s(X0,X8)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f602,f505])).
fof(f691,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f418,f602,f505])).
fof(f692,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f421,f602,f505])).
fof(f693,plain,(
  ( ! [X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f422,f602,f505])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),happ(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(bool,X0))),s(bool,X1))),s(bool,X2)),'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f426,f602,f505])).
fof(f695,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),happ(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),X8),s(X4,X6)))),happ(s(fun(X4,bool),'_dst_x0x1x1x1_4'(s(fun(X2,fun(X1,fun(X0,bool))),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7),s(fun(X4,X0),X8))),s(X4,X6)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f591,f602,f505])).
fof(f732,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f522])).
fof(f733,plain,(
  ( ! [X0] : (~sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f523])).
fof(f743,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f533])).
fof(f744,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f534])).
fof(f746,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f537])).
fof(f749,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f537])).
fof(f750,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f537])).
fof(f752,plain,(
  ( ! [X0] : (sP9('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f538])).
fof(f755,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f538])).
fof(f756,plain,(
  ( ! [X0] : (sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f538])).
fof(f758,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f540])).
fof(f759,plain,(
  ( ! [X0] : (~sP9('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f541])).
fof(f767,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f551])).
fof(f769,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f554])).
fof(f771,plain,(
  ( ! [X0] : (sP11('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f555])).
fof(f774,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f555])).
fof(f775,plain,(
  ( ! [X0] : (sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f555])).
fof(f782,plain,(
  ( ! [X0] : (~sP11('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f558])).
fof(f900,plain,(
  ( ! [X0,X1] : (sP14(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f572])).
fof(f907,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f576])).
fof(f909,plain,(
  ( ! [X0,X1] : (sP12(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f580])).
fof(f910,plain,(
  ( ! [X0,X1] : (sP12(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f581])).
fof(f918,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(resolution,[],[f602,f364])).
fof(f917,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(resolution,[],[f602,f362])).
fof(f606,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f362])).
fof(f607,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f364])).
fof(f960,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl7),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f608,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f365])).
fof(f609,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f367])).
fof(f637,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f434])).
fof(f638,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f434])).
fof(f640,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f434])).
fof(f641,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f434])).
fof(f1127,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl95),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f647,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f453])).
fof(f648,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f453])).
fof(f1125,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl93),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f649,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f453])).
fof(f1129,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl97),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f650,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f453])).
fof(f651,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f453])).
fof(f652,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f453])).
fof(f1131,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl98),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f653,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f454])).
fof(f1133,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl100),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f654,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f457])).
fof(f655,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f458])).
fof(f656,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f459])).
fof(f657,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f460])).
fof(f658,plain,(
  ~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f460])).
fof(f659,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f460])).
fof(f660,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f460])).
fof(f661,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f460])).
fof(f662,plain,(
  ~sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f460])).
fof(f1123,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl91),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f663,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f488])).
fof(f664,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f488])).
fof(f1119,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl87),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f665,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f488])).
fof(f1121,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl89),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f666,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f488])).
fof(f667,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f488])).
fof(f668,plain,(
  ~sP5('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f488])).
fof(f1117,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl85),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f669,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f492])).
fof(f670,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f492])).
fof(f1113,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl81),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f671,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f492])).
fof(f1115,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl83),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f672,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f492])).
fof(f673,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f492])).
fof(f674,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f492])).
fof(f1111,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl79),
  inference(cnf_transformation,[],[f1111_D])).
fof(f1111_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f676,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f496])).
fof(f677,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f496])).
fof(f1107,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl75),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f678,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f496])).
fof(f1109,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl77),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f679,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f496])).
fof(f680,plain,(
  ~sP3('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f496])).
fof(f681,plain,(
  ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f496])).
fof(f1105,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl73),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f698,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f515])).
fof(f699,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f515])).
fof(f1101,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl69),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f700,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f515])).
fof(f1103,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl71),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f701,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f515])).
fof(f702,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f515])).
fof(f703,plain,(
  ~sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f515])).
fof(f1099,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl66),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f705,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f706,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f1095,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl62),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f707,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f708,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f709,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f710,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f1091,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl58),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f711,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f712,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f1087,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl54),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f713,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1097,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl64),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f714,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f715,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f716,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f717,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f718,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f719,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f1089,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl56),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f720,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f721,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f722,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1093,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl60),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f723,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f345,f602,f521])).
fof(f724,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f345,f602,f521])).
fof(f725,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f345,f602,f521])).
fof(f726,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f344,f602,f521])).
fof(f727,plain,(
  sP7('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f344,f602,f521])).
fof(f728,plain,(
  sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f344,f602,f521])).
fof(f729,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f602,f602,f521])).
fof(f730,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f602,f602,f521])).
fof(f731,plain,(
  sP7('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f602,f602,f521])).
fof(f1085,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl53),
  inference(cnf_transformation,[],[f1085_D])).
fof(f1085_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f734,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f524])).
fof(f735,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f524])).
fof(f1081,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl49),
  inference(cnf_transformation,[],[f1081_D])).
fof(f1081_D,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f736,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f524])).
fof(f1083,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl51),
  inference(cnf_transformation,[],[f1083_D])).
fof(f1083_D,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f737,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f524])).
fof(f738,plain,(
  ~sP7('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f524])).
fof(f739,plain,(
  ~sP7('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f524])).
fof(f1079,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl47),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f742,plain,(
  ~sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f532])).
fof(f1077,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl45),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f757,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f539])).
fof(f1075,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl43),
  inference(cnf_transformation,[],[f1075_D])).
fof(f1075_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f761,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f549])).
fof(f762,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f549])).
fof(f1071,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl39),
  inference(cnf_transformation,[],[f1071_D])).
fof(f1071_D,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f763,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f549])).
fof(f1073,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl41),
  inference(cnf_transformation,[],[f1073_D])).
fof(f1073_D,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f764,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f549])).
fof(f765,plain,(
  ~sP10('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f549])).
fof(f766,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f549])).
fof(f1069,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl36),
  inference(cnf_transformation,[],[f1069_D])).
fof(f1069_D,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f768,plain,(
  sP11('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f553])).
fof(f1067,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl35),
  inference(cnf_transformation,[],[f1067_D])).
fof(f1067_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f776,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f556])).
fof(f777,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f556])).
fof(f1063,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl31),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f778,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f556])).
fof(f1065,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl33),
  inference(cnf_transformation,[],[f1065_D])).
fof(f1065_D,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f779,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f556])).
fof(f780,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f556])).
fof(f781,plain,(
  ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f556])).
fof(f783,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f602,f560])).
fof(f784,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f602,f560])).
fof(f789,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f562])).
fof(f790,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f562])).
fof(f1003,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f803,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f562])).
fof(f804,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f602,f562])).
fof(f805,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f421,f602,f562])).
fof(f814,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f562])).
fof(f815,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f602,f562])).
fof(f816,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f422,f602,f562])).
fof(f825,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f421,f602,f564])).
fof(f826,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f422,f602,f564])).
fof(f831,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f602,f565])).
fof(f832,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f602,f565])).
fof(f833,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f602,f565])).
fof(f840,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f345,f602,f565])).
fof(f841,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f344,f602,f565])).
fof(f842,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f602,f565])).
fof(f851,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f566])).
fof(f852,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f421,f602,f566])).
fof(f853,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f421,f602,f566])).
fof(f862,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f566])).
fof(f863,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f422,f602,f566])).
fof(f864,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f422,f602,f566])).
fof(f873,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f421,f602,f566])).
fof(f874,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f422,f602,f566])).
fof(f981,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl27),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f887,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f345,f602,f567])).
fof(f888,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f344,f602,f567])).
fof(f977,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl23),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f889,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f602,f602,f567])).
fof(f979,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl25),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f890,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f567])).
fof(f891,plain,(
  ~sP15('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f567])).
fof(f892,plain,(
  ~sP15('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f567])).
fof(f975,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl21),
  inference(cnf_transformation,[],[f975_D])).
fof(f975_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f894,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f571])).
fof(f895,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f571])).
fof(f971,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl17),
  inference(cnf_transformation,[],[f971_D])).
fof(f971_D,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f896,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f571])).
fof(f973,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl19),
  inference(cnf_transformation,[],[f973_D])).
fof(f973_D,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f897,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f571])).
fof(f898,plain,(
  ~sP14('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f571])).
fof(f899,plain,(
  ~sP14('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f571])).
fof(f969,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl15),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f901,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f575])).
fof(f902,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f575])).
fof(f965,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl11),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f903,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f575])).
fof(f967,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl13),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f904,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f345,f352,f602,f575])).
fof(f905,plain,(
  ~sP13('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f344,f352,f602,f575])).
fof(f906,plain,(
  ~sP13('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f602,f352,f602,f575])).
fof(f963,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl9),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f908,plain,(
  ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f579])).
fof(f961,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(forward_demodulation,[],[f911,f422])).
fof(f911,plain,(
  ( ! [X0] : (s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0'))) != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f587])).
fof(f914,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))) | $spl2),
  inference(unit_resulting_resolution,[],[f352,f352,f602,f598])).
fof(f604,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) | $spl4),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f341,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19))))),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK19)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18,sK19,sK20,sK21])],[f157])).
fof(f157,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X5)))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X5))))) & ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X5))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X5))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X1,bool),X3),s(fun(X1,bool),X5),s(fun(X1,bool),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X7,bool),X9),s(fun(X7,bool),X11),s(fun(X7,bool),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))) & p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X11))))) => p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(X7,bool),X9),s(fun(X7,bool),X11),s(fun(X7,bool),X10)))))))),
  file('thm.quantHeuristics.GUESS_RULES_COND|split|9','thm.quantHeuristics.GUESS_RULES_COND|split|9')).
fof(f602,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) | $spl2),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f342,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),sK21))))),
  inference(cnf_transformation,[],[f214])).
fof(f600,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) | $spl1),
  inference(cnf_transformation,[],[f600_D])).
fof(f600_D,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f343,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK16,sK17),sK18),s(fun(sK17,bool),'_dst_x0x1x1x1_4'(s(fun(bool,fun(bool,fun(bool,bool))),'const.bool.COND_0'),s(fun(sK17,bool),sK19),s(fun(sK17,bool),sK21),s(fun(sK17,bool),sK20))))))),
  inference(cnf_transformation,[],[f214])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 502
% Passive clauses: 1224
% Generated clauses: 22386
% Final active clauses: 495
% Final passive clauses: 859
% Input formulas: 62
% Initial clauses: 251
% 
% Duplicate literals: 85
% Trivial inequalities: 6
% Fw demodulations: 10818
% Bw demodulations: 34
% Global subsumptions: 8
% 
% Simple tautologies: 285
% Equational tautologies: 14
% Forward subsumptions: 6179
% Backward subsumptions: 19
% Fw demodulations to eq. taut.: 2600
% 
% Binary resolution: 1252
% Unit resulting resolution: 8997
% Forward superposition: 420
% Backward superposition: 413
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 3720
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2957
% 
% Sat splits: 4
% 
% Memory used [KB]: 14072
% Time elapsed: 0.595 s
% ------------------------------
% ------------------------------
% Success in time 18.655 s
