Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (99766)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 1759
% Generated clauses: 21657
% Final active clauses: 112
% Final passive clauses: 1477
% Input formulas: 26
% Initial clauses: 34
% 
% Duplicate literals: 742
% Trivial inequalities: 207
% Fw subsumption resolutions: 1900
% Fw demodulations: 3017
% Bw demodulations: 2
% 
% Simple tautologies: 198
% Equational tautologies: 26
% Forward subsumptions: 13159
% 
% Binary resolution: 6
% Forward superposition: 8508
% Backward superposition: 6842
% Self superposition: 239
% Equality factoring: 53
% Equality resolution: 17
% 
% Split clauses: 252
% Split components: 515
% SAT solver clauses: 616
% SAT solver unit clauses: 39
% SAT solver binary clauses: 123
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 252
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 615
% TWLsolver calls for satisfiability: 109
% 
% Memory used [KB]: 15735
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (99789)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 93
% Passive clauses: 546
% Generated clauses: 15315
% Final active clauses: 81
% Final passive clauses: 445
% Input formulas: 26
% Initial clauses: 34
% 
% Duplicate literals: 667
% Trivial inequalities: 173
% Fw subsumption resolutions: 870
% Fw demodulations: 1538
% Bw demodulations: 7
% Condensations: 1942
% 
% Simple tautologies: 166
% Equational tautologies: 26
% Forward subsumptions: 8427
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 4
% Forward superposition: 5300
% Backward superposition: 3884
% Self superposition: 213
% Equality factoring: 64
% Equality resolution: 18
% 
% Split clauses: 41
% Split components: 85
% SAT solver clauses: 966
% SAT solver unit clauses: 52
% SAT solver binary clauses: 45
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 7
% 
% Sat splits: 41
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 963
% TWLsolver calls for satisfiability: 230
% 
% Memory used [KB]: 15351
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f13083,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f87,f126,f124,f122,f120,f118,f86,f125,f85,f123,f84,f121,f83,f119,f117,f132,f155,f158,f156,f708,f709,f862,f710,f863,f861,f2661,f2663,f2662,f10697,f10719,f10727,f11326,f11349,f11357,f12012,f12038,f12046,f12085,f12971,f13001,f13041,f13079,f13080,f13081,f13082])).
fof(f13082,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f122,f12977])).
fof(f12977,plain,(
  s(sK0,sK2) = s(sK0,sK6) | $spl0),
  inference(forward_demodulation,[],[f12976,f127])).
fof(f127,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f115,f114])).
fof(f114,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15,X6] : ! [X16,X17] : s(X6,X16) = s(X6,'const.combin.K_2'(s(X6,X16),s(X15,X17)))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.combin.K_THM')).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X15] : ! [X22,X23] : s(X6,'const.combin.K_2'(s(X6,X22),s(X15,X23))) = s(X6,happ(s(fun(X15,X6),'const.combin.K_1'(s(X6,X22))),s(X15,X23)))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','ARITY_1_2_const.combin.K')).
fof(f12976,plain,(
  s(sK0,sK6) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12975,f127])).
fof(f12975,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12974,f93])).
fof(f93,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,happ(s(fun(X6,X6),X19),s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X18))))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_29_prove_ltb')).
fof(f12974,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12973,f93])).
fof(f12973,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12972,f92])).
fof(f92,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_|49|9_prove_ltb')).
fof(f12972,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12895,f92])).
fof(f12895,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) & ((s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))) | s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) <=> (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X20,X21] : (s('type.semi_ring.semi_ring'(X6),X20) = s('type.semi_ring.semi_ring'(X6),X21) <=> (s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X20))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X21))) & s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X20))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X21))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X21))) & s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X21)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_component_equality')).
fof(f13081,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f87,f86,f85,f84,f83,f88,f89,f90,f91,f92,f93,f94,f95,f96,f97,f98,f99,f100,f101,f102,f103,f104,f105,f106,f111,f110,f109,f108,f107,f112,f113,f114,f127,f116,f144,f146,f147,f155,f184,f194,f195,f227,f228,f240,f250,f251,f294,f297,f464,f466,f468,f479,f518,f520,f522,f533,f601,f602,f603,f604,f614,f670,f672,f674,f708,f709,f710,f733,f747,f777,f900,f901,f992,f994,f996,f1497,f1083,f1084,f1085,f1086,f1087,f1088,f1089,f1107,f1108,f1109,f1110,f1111,f1131,f1133,f1135,f1150,f1151,f1263,f1264,f1265,f1343,f1331,f1637,f1650,f1651,f1652,f1653,f1566,f1567,f1568,f1692,f1693,f1694,f1698,f1699,f1716,f1718,f1721,f1726,f1746,f1854,f1855,f1997,f1998,f1999,f2003,f2004,f2022,f2024,f2027,f2032,f2056,f2347,f2082,f2083,f2084,f2085,f2086,f2087,f2088,f2344,f2110,f2111,f2112,f2114,f2116,f2374,f2637,f2421,f2574,f2702,f2703,f2704,f2708,f2709,f2710,f2731,f2733,f2736,f2740,f2743,f2771,f2927,f2928,f2929,f2931,f2933,f2934,f2954,f2955,f3057,f3058,f3059,f3063,f3064,f3065,f3088,f3090,f3093,f3097,f3100,f3131,f3105,f3275,f3276,f4151,f4147,f4016,f4017,f4018,f4019,f4020,f4021,f4022,f4023,f4024,f4025,f4051,f4052,f4053,f4054,f4056,f4057,f4059,f4060,f4354,f4258,f4259,f4260,f4261,f4263,f4264,f4266,f4267,f4328,f4399,f4430,f4431,f4432,f4433,f4434,f4435,f4436,f4437,f4438,f4439,f4616,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f5000,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f5864,f5518,f5519,f6014,f6015,f6016,f6018,f6019,f6021,f6213,f6214,f6215,f6218,f6219,f6220,f6237,f6224,f6225,f6226,f6229,f6230,f6231,f6232,f6239,f6488,f6489,f6361,f6362,f6363,f6367,f6368,f6369,f6403,f6405,f6410,f6414,f6417,f6418,f6419,f6501,f6424,f6429,f6431,f6432,f6435,f6773,f6774,f6643,f6644,f6645,f6649,f6650,f6651,f6687,f6689,f6694,f6698,f6701,f6702,f6703,f6786,f6708,f6713,f6715,f6716,f6719,f6863,f6931,f6932,f7616,f7622,f7624,f7628,f7633,f7635,f7474,f7475,f7718,f7858,f7851,f7852,f7853,f7854,f7938,f7939,f7940,f7899,f7900,f7901,f7972,f7965,f7966,f7967,f7968,f8097,f8193,f8194,f8755,f8778,f8890,f8891,f8892,f8849,f8850,f8851,f8965,f8966,f9187,f9188,f9189,f9190,f9191,f9192,f9193,f9194,f9196,f9197,f9198,f9199,f9200,f9234,f9235,f9236,f9237,f9238,f9239,f9240,f9241,f9243,f9244,f9286,f9287,f9288,f9290,f9291,f9293,f9294,f9295,f9296,f9324,f9325,f9326,f9509,f9522,f9658,f9671,f9676,f9678,f9679,f10096,f9722,f9723,f9724,f9727,f9728,f9730,f9738,f9739,f9742,f9746,f9872,f9885,f9987,f9975,f10750,f10749,f10933,f10931,f10804,f10805,f10806,f11007,f11012,f11013,f11014,f11386,f11391,f11392,f11393,f11665,f11666,f11667,f117,f12719,f12717,f12130,f12131,f12132,f12133,f12134,f12135,f12136,f12137,f12139,f12140,f12712,f12186,f12187,f12188,f12190,f12191,f12193,f12195,f12196,f12637,f12797,f12630,f12413,f12977,f12978,f12986,f12987,f12995,f12996,f13040,f13000,f12970])).
fof(f12970,plain,(
  s(sK0,sK1) = s(sK0,sK5) | $spl0),
  inference(forward_demodulation,[],[f12969,f127])).
fof(f12969,plain,(
  s(sK0,sK5) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12968,f127])).
fof(f12968,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12967,f94])).
fof(f12967,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12896,f94])).
fof(f12896,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f107])).
fof(f13000,plain,(
  s(sK0,sK1) = s(sK0,sK5) | $spl0),
  inference(forward_demodulation,[],[f12999,f127])).
fof(f12999,plain,(
  s(sK0,sK5) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12998,f127])).
fof(f12998,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12997,f94])).
fof(f12997,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12892,f94])).
fof(f12892,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f107])).
fof(f13040,plain,(
  s(sK0,sK1) = s(sK0,sK5) | $spl0),
  inference(forward_demodulation,[],[f13039,f127])).
fof(f13039,plain,(
  s(sK0,sK5) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f13038,f127])).
fof(f13038,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12899,f94])).
fof(f12899,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(superposition,[],[f94,f117])).
fof(f12996,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(backward_demodulation,[],[f12995,f12987])).
fof(f12995,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl0),
  inference(forward_demodulation,[],[f12994,f7899])).
fof(f12994,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | $spl0),
  inference(forward_demodulation,[],[f12993,f7899])).
fof(f12993,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | $spl0),
  inference(forward_demodulation,[],[f12992,f103])).
fof(f12992,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | $spl0),
  inference(forward_demodulation,[],[f12991,f103])).
fof(f12991,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | $spl0),
  inference(forward_demodulation,[],[f12990,f101])).
fof(f12990,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12989,f101])).
fof(f12989,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12988,f102])).
fof(f12988,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12893,f102])).
fof(f12893,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f110])).
fof(f12987,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(backward_demodulation,[],[f12986,f12978])).
fof(f12986,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | $spl0),
  inference(forward_demodulation,[],[f12985,f8849])).
fof(f12985,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK7) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | $spl0),
  inference(forward_demodulation,[],[f12984,f8849])).
fof(f12984,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | $spl0),
  inference(forward_demodulation,[],[f12983,f8892])).
fof(f12983,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | $spl0),
  inference(forward_demodulation,[],[f12982,f8892])).
fof(f12982,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | $spl0),
  inference(forward_demodulation,[],[f12981,f100])).
fof(f12981,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12980,f100])).
fof(f12980,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(forward_demodulation,[],[f12979,f99])).
fof(f12979,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f12894,f99])).
fof(f12894,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f109])).
fof(f12978,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(backward_demodulation,[],[f12977,f117])).
fof(f12413,plain,(
  ( ! [X947,X945,X946,X944,X948,X943] : (s(bool,f(s(X943,sK9(X944,X945,bool,X943)))) != s(bool,happ(s(fun(X946,bool),'const.combin.K_1'(s(bool,X947))),s(X946,X948))) | s(fun(X943,bool),X944) = s(fun(X943,bool),X945) | s(bool,'T_0') = s(bool,X947) | s(bool,'T_0') = s(bool,g(s(X943,sK9(X944,X945,bool,X943))))) )),
  inference(superposition,[],[f116,f251])).
fof(f12630,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,X3))),s(X2,X4))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f12629])).
fof(f12629,plain,(
  ( ! [X292,X288,X290,X293,X289,X291] : (s(bool,X292) != s(bool,happ(s(fun(X289,bool),'const.combin.K_1'(s(bool,X290))),s(X289,X291))) | s(bool,X292) = s(bool,X293) | s(bool,'T_0') = s(bool,X293) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,X290)) )),
  inference(duplicate_literal_removal,[],[f12328])).
fof(f12328,plain,(
  ( ! [X292,X288,X290,X293,X289,X291] : (s(bool,X292) != s(bool,happ(s(fun(X289,bool),'const.combin.K_1'(s(bool,X290))),s(X289,X291))) | s(bool,X292) = s(bool,X293) | s(bool,'T_0') = s(bool,X293) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,X288)) )),
  inference(superposition,[],[f4399,f251])).
fof(f12797,plain,(
  ( ! [X4,X0,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X0))),s(X3,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f12636,f114])).
fof(f12636,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X0))),s(X3,X4)))) )),
  inference(condensation,[],[f12635])).
fof(f12635,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X187,X188))) != s(bool,happ(s(fun(X184,bool),'const.combin.K_1'(s(bool,X185))),s(X184,X186))) | s(bool,'T_0') = s(bool,X183) | s(bool,'T_0') = s(bool,X185)) )),
  inference(duplicate_literal_removal,[],[f12313])).
fof(f12313,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X187,X188))) != s(bool,happ(s(fun(X184,bool),'const.combin.K_1'(s(bool,X185))),s(X184,X186))) | s(bool,'T_0') = s(bool,X183) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f1343,f251])).
fof(f12637,plain,(
  ( ! [X175,X173,X171,X169,X174,X172,X170] : (s(bool,'const.combin.K_2'(s(bool,X173),s(X174,X175))) != s(bool,happ(s(fun(X170,bool),'const.combin.K_1'(s(bool,X171))),s(X170,X172))) | s(bool,X169) = s(bool,X173) | s(bool,'T_0') = s(bool,X169) | s(bool,'T_0') = s(bool,X171)) )),
  inference(duplicate_literal_removal,[],[f12311])).
fof(f12311,plain,(
  ( ! [X175,X173,X171,X169,X174,X172,X170] : (s(bool,'const.combin.K_2'(s(bool,X173),s(X174,X175))) != s(bool,happ(s(fun(X170,bool),'const.combin.K_1'(s(bool,X171))),s(X170,X172))) | s(bool,X169) = s(bool,X173) | s(bool,'T_0') = s(bool,X169) | s(bool,'T_0') = s(bool,X171) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f1331,f251])).
fof(f12196,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (s(bool,X339) = s(bool,happ(s(fun(X340,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X336),s(bool,X337),s(bool,X338))))),s(X340,X341))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X336))),s(bool,X337))),s(bool,X338))) | s(bool,'T_0') = s(bool,X339)) )),
  inference(superposition,[],[f251,f113])).
fof(f12195,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,X333) = s(bool,happ(s(fun(X334,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X330),s(bool,X331),s(bool,X332))))),s(X334,X335))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X330))),s(bool,X331))),s(bool,X332))) | s(bool,'T_0') = s(bool,X333)) )),
  inference(superposition,[],[f251,f112])).
fof(f12193,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,X321) = s(bool,happ(s(fun(X322,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X320))))),s(X322,X323))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X319),s('type.semi_ring.semi_ring'(bool),X320))))) | s(bool,'T_0') = s(bool,X321)) )),
  inference(superposition,[],[f251,f92])).
fof(f12191,plain,(
  ( ! [X313,X310,X312,X309,X311] : (s(bool,X311) = s(bool,happ(s(fun(X312,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X310))))),s(X312,X313))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X309),s('type.semi_ring.semi_ring'(bool),X310))))) | s(bool,'T_0') = s(bool,X311)) )),
  inference(superposition,[],[f251,f97])).
fof(f12190,plain,(
  ( ! [X308,X304,X306,X305,X307] : (s(bool,X306) = s(bool,happ(s(fun(X307,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X305))))),s(X307,X308))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X304),s('type.semi_ring.semi_ring'(bool),X305))))) | s(bool,'T_0') = s(bool,X306)) )),
  inference(superposition,[],[f251,f98])).
fof(f12188,plain,(
  ( ! [X294,X296,X298,X295,X297] : (s(bool,X296) = s(bool,happ(s(fun(X297,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X295))))),s(X297,X298))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X294),s('type.semi_ring.semi_ring'(bool),X295))))) | s(bool,'T_0') = s(bool,X296)) )),
  inference(superposition,[],[f251,f91])).
fof(f12187,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,X291) = s(bool,happ(s(fun(X292,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X290))))),s(X292,X293))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X289),s('type.semi_ring.semi_ring'(bool),X290))))) | s(bool,'T_0') = s(bool,X291)) )),
  inference(superposition,[],[f251,f96])).
fof(f12186,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X286) = s(bool,happ(s(fun(X287,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X285))))),s(X287,X288))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X284),s('type.semi_ring.semi_ring'(bool),X285))))) | s(bool,'T_0') = s(bool,X286)) )),
  inference(superposition,[],[f251,f95])).
fof(f12712,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X1,X2)))) )),
  inference(condensation,[],[f12711])).
fof(f12711,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f12141])).
fof(f12141,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f251,f90])).
fof(f12140,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X336),s(bool,X337),s(bool,X338))) = s(bool,happ(s(fun(X339,bool),'const.combin.K_1'(s(bool,X340))),s(X339,X341))) | s(bool,'T_0') = s(bool,X340) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X336))),s(bool,X337))),s(bool,X338)))) )),
  inference(superposition,[],[f251,f113])).
fof(f12139,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X330),s(bool,X331),s(bool,X332))) = s(bool,happ(s(fun(X333,bool),'const.combin.K_1'(s(bool,X334))),s(X333,X335))) | s(bool,'T_0') = s(bool,X334) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X330))),s(bool,X331))),s(bool,X332)))) )),
  inference(superposition,[],[f251,f112])).
fof(f12137,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X320))) = s(bool,happ(s(fun(X321,bool),'const.combin.K_1'(s(bool,X322))),s(X321,X323))) | s(bool,'T_0') = s(bool,X322) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X319),s('type.semi_ring.semi_ring'(bool),X320)))))) )),
  inference(superposition,[],[f251,f92])).
fof(f12136,plain,(
  ( ! [X315,X316,X318,X314,X317] : (s(bool,happ(s(fun(X316,bool),'const.combin.K_1'(s(bool,X317))),s(X316,X318))) = s(bool,happ(s(fun(bool,bool),X314),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X315))))) | s(bool,'T_0') = s(bool,X317) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X314),s('type.semi_ring.semi_ring'(bool),X315)))))) )),
  inference(superposition,[],[f251,f93])).
fof(f12135,plain,(
  ( ! [X313,X310,X312,X309,X311] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X310))) = s(bool,happ(s(fun(X311,bool),'const.combin.K_1'(s(bool,X312))),s(X311,X313))) | s(bool,'T_0') = s(bool,X312) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X309),s('type.semi_ring.semi_ring'(bool),X310)))))) )),
  inference(superposition,[],[f251,f97])).
fof(f12134,plain,(
  ( ! [X308,X304,X306,X305,X307] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X305))) = s(bool,happ(s(fun(X306,bool),'const.combin.K_1'(s(bool,X307))),s(X306,X308))) | s(bool,'T_0') = s(bool,X307) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X304),s('type.semi_ring.semi_ring'(bool),X305)))))) )),
  inference(superposition,[],[f251,f98])).
fof(f12133,plain,(
  ( ! [X300,X302,X301,X303,X299] : (s(bool,happ(s(fun(X301,bool),'const.combin.K_1'(s(bool,X302))),s(X301,X303))) = s(bool,happ(s(fun(bool,bool),X299),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X300))))) | s(bool,'T_0') = s(bool,X302) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X299),s('type.semi_ring.semi_ring'(bool),X300)))))) )),
  inference(superposition,[],[f251,f94])).
fof(f12132,plain,(
  ( ! [X294,X296,X298,X295,X297] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X295))) = s(bool,happ(s(fun(X296,bool),'const.combin.K_1'(s(bool,X297))),s(X296,X298))) | s(bool,'T_0') = s(bool,X297) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X294),s('type.semi_ring.semi_ring'(bool),X295)))))) )),
  inference(superposition,[],[f251,f91])).
fof(f12131,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X290))) = s(bool,happ(s(fun(X291,bool),'const.combin.K_1'(s(bool,X292))),s(X291,X293))) | s(bool,'T_0') = s(bool,X292) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X289),s('type.semi_ring.semi_ring'(bool),X290)))))) )),
  inference(superposition,[],[f251,f96])).
fof(f12130,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X285))) = s(bool,happ(s(fun(X286,bool),'const.combin.K_1'(s(bool,X287))),s(X286,X288))) | s(bool,'T_0') = s(bool,X287) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X284),s('type.semi_ring.semi_ring'(bool),X285)))))) )),
  inference(superposition,[],[f251,f95])).
fof(f12717,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,X2))),s(X1,X3))) = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X0))),s(X4,X5))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f12716])).
fof(f12716,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,happ(s(fun(X74,bool),'const.combin.K_1'(s(bool,X75))),s(X74,X76))) = s(bool,happ(s(fun(X77,bool),'const.combin.K_1'(s(bool,X78))),s(X77,X79))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X75)) )),
  inference(duplicate_literal_removal,[],[f12102])).
fof(f12102,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,happ(s(fun(X74,bool),'const.combin.K_1'(s(bool,X75))),s(X74,X76))) = s(bool,happ(s(fun(X77,bool),'const.combin.K_1'(s(bool,X78))),s(X77,X79))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f251,f251])).
fof(f12719,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X0))),s(X4,X5))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f12718])).
fof(f12718,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.combin.K_2'(s(bool,X67),s(X68,X69))) = s(bool,happ(s(fun(X70,bool),'const.combin.K_1'(s(bool,X71))),s(X70,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X67)) )),
  inference(duplicate_literal_removal,[],[f12101])).
fof(f12101,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.combin.K_2'(s(bool,X67),s(X68,X69))) = s(bool,happ(s(fun(X70,bool),'const.combin.K_1'(s(bool,X71))),s(X70,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f251,f250])).
fof(f11667,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (s(bool,X19) != s(bool,happ(s(fun(X20,bool),'const.combin.K_1'(s(bool,X18))),s(X20,X21))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X22,X23))) | s(bool,X18) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f10933,f144])).
fof(f11666,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (s(bool,X10) != s(bool,happ(s(fun(X11,bool),'const.combin.K_1'(s(bool,X9))),s(X11,X12))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X10),s(X13,X14))) | s(bool,X9) = s(bool,X15) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X15),s(X16,X17)))) )),
  inference(superposition,[],[f10933,f227])).
fof(f11665,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,X0))),s(X2,X3))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X1),s(X4,X5))) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),'const.combin.K_1'(s(bool,X6))),s(X7,X8)))) )),
  inference(superposition,[],[f10933,f228])).
fof(f11393,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X10))),s('type.semi_ring.semi_ring'(bool),X11))))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f10750,f144])).
fof(f11392,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X5))),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X7),s(X8,X9)))) )),
  inference(superposition,[],[f10750,f227])).
fof(f11391,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X0))),s('type.semi_ring.semi_ring'(bool),X1))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(superposition,[],[f10750,f228])).
fof(f11386,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'T_0'))),s('type.semi_ring.semi_ring'(bool),X1)))))))) )),
  inference(unit_resulting_resolution,[],[f10750,f1650])).
fof(f11014,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X10))),s('type.semi_ring.semi_ring'(bool),X11))))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f10749,f144])).
fof(f11013,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X5))),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X7),s(X8,X9)))) )),
  inference(superposition,[],[f10749,f227])).
fof(f11012,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,X0))),s('type.semi_ring.semi_ring'(bool),X1))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(superposition,[],[f10749,f228])).
fof(f11007,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'T_0'))),s('type.semi_ring.semi_ring'(bool),X1)))))))) )),
  inference(unit_resulting_resolution,[],[f10749,f1652])).
fof(f10806,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) != s(bool,happ(s(fun(X16,bool),'const.combin.K_1'(s(bool,X14))),s(X16,X17))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f9987,f144])).
fof(f10805,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,X8) != s(bool,happ(s(fun(X9,bool),'const.combin.K_1'(s(bool,X7))),s(X9,X10))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X11),s(X12,X13)))) )),
  inference(superposition,[],[f9987,f227])).
fof(f10804,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,X0))),s(X2,X3))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X5,bool),'const.combin.K_1'(s(bool,X4))),s(X5,X6)))) )),
  inference(superposition,[],[f9987,f228])).
fof(f10931,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'const.combin.K_1'(s(bool,X1))),s(X0,X2))) | s(bool,X1) != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X3,X4)))) )),
  inference(condensation,[],[f10930])).
fof(f10930,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,X44) != s(bool,happ(s(fun(X45,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X45,X46))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),'const.combin.K_1'(s(bool,X44))),s(X47,X48)))) )),
  inference(duplicate_literal_removal,[],[f10762])).
fof(f10762,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,X44) != s(bool,happ(s(fun(X45,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X45,X46))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),'const.combin.K_1'(s(bool,X44))),s(X47,X48)))) )),
  inference(superposition,[],[f9987,f228])).
fof(f10933,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X3,X4))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) )),
  inference(condensation,[],[f10932])).
fof(f10932,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,happ(s(fun(X27,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(duplicate_literal_removal,[],[f10759])).
fof(f10759,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,happ(s(fun(X27,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(superposition,[],[f9987,f227])).
fof(f10749,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'T_0'))),s('type.semi_ring.semi_ring'(bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f93,f9987])).
fof(f10750,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'T_0'))),s('type.semi_ring.semi_ring'(bool),X0)))))) )),
  inference(unit_resulting_resolution,[],[f94,f9987])).
fof(f9975,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,X2))),s(X1,X3))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f228])).
fof(f9987,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f9969])).
fof(f9969,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f228])).
fof(f9885,plain,(
  ( ! [X905,X903,X901,X904,X902,X900] : (s(X902,f(s(bool,'T_0'))) != s(X902,g(s(bool,'T_0'))) | s(fun(bool,X902),X900) = s(fun(bool,X902),X901) | s(bool,X903) = s(bool,sK9(X900,X901,X902,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X904,bool),'const.combin.K_1'(s(bool,X903))),s(X904,X905)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9872,plain,(
  ( ! [X831,X835,X833,X830,X834,X832] : (s(bool,'T_0') != s(bool,f(s(X830,sK9(X831,X832,bool,X830)))) | s(fun(X830,bool),X831) = s(fun(X830,bool),X832) | s(bool,X833) = s(bool,g(s(X830,sK9(X831,X832,bool,X830)))) | s(bool,'T_0') = s(bool,happ(s(fun(X834,bool),'const.combin.K_1'(s(bool,X833))),s(X834,X835)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9746,plain,(
  ( ! [X499,X500,X502,X498,X504,X501,X503] : (s(bool,X498) != s(bool,'const.combin.K_2'(s(bool,X499),s(bool,X498))) | s(bool,'T_0') = s(bool,happ(s(fun(X500,bool),'const.combin.K_1'(s(bool,X499))),s(X500,X501))) | s(bool,X498) = s(bool,X502) | s(bool,'T_0') = s(bool,happ(s(fun(X503,bool),'const.combin.K_1'(s(bool,X502))),s(X503,X504)))) )),
  inference(superposition,[],[f2934,f228])).
fof(f9742,plain,(
  ( ! [X476,X478,X474,X475,X477,X479,X473] : (s(bool,X473) = s(bool,'const.combin.K_2'(s(bool,X474),s(bool,X473))) | s(bool,X474) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X475,X476))) | s(bool,X473) = s(bool,X477) | s(bool,'T_0') = s(bool,happ(s(fun(X478,bool),'const.combin.K_1'(s(bool,X477))),s(X478,X479)))) )),
  inference(superposition,[],[f2637,f228])).
fof(f9739,plain,(
  ( ! [X455,X457,X459,X454,X460,X456,X458] : (s(bool,X454) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X455))),s(bool,X454))) | s(bool,X455) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X456,X457))) | s(bool,X454) = s(bool,X458) | s(bool,'T_0') = s(bool,happ(s(fun(X459,bool),'const.combin.K_1'(s(bool,X458))),s(X459,X460)))) )),
  inference(superposition,[],[f2421,f228])).
fof(f9738,plain,(
  ( ! [X453,X449,X451,X452,X450] : (s(bool,X449) != s(bool,'const.combin.K_2'(s(bool,X450),s(bool,X449))) | s(bool,'T_0') = s(bool,X450) | s(bool,X449) = s(bool,X451) | s(bool,'T_0') = s(bool,happ(s(fun(X452,bool),'const.combin.K_1'(s(bool,X451))),s(X452,X453)))) )),
  inference(superposition,[],[f2374,f228])).
fof(f9730,plain,(
  ( ! [X405,X401,X403,X404,X406,X400,X402] : (s(bool,X400) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X401))),s(bool,X400))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X401),s(X402,X403))) | s(bool,X400) = s(bool,X404) | s(bool,'T_0') = s(bool,happ(s(fun(X405,bool),'const.combin.K_1'(s(bool,X404))),s(X405,X406)))) )),
  inference(superposition,[],[f601,f228])).
fof(f9728,plain,(
  ( ! [X389,X391,X393,X388,X390,X392] : (s(bool,X388) = s(bool,'const.combin.K_2'(s(bool,X389),s(bool,X388))) | s(bool,X389) = s(bool,X390) | s(bool,'T_0') = s(bool,X390) | s(bool,X388) = s(bool,X391) | s(bool,'T_0') = s(bool,happ(s(fun(X392,bool),'const.combin.K_1'(s(bool,X391))),s(X392,X393)))) )),
  inference(superposition,[],[f227,f228])).
fof(f9727,plain,(
  ( ! [X383,X385,X387,X382,X384,X386,X381] : (s(bool,X381) = s(bool,'const.combin.K_2'(s(bool,X382),s(bool,X381))) | s(bool,'F_0') = s(bool,happ(s(fun(X383,bool),'const.combin.K_1'(s(bool,X382))),s(X383,X384))) | s(bool,X381) = s(bool,X385) | s(bool,'T_0') = s(bool,happ(s(fun(X386,bool),'const.combin.K_1'(s(bool,X385))),s(X386,X387)))) )),
  inference(superposition,[],[f195,f228])).
fof(f9724,plain,(
  ( ! [X364,X366,X360,X362,X365,X361,X363] : (s(bool,X360) = s(bool,'const.combin.K_2'(s(bool,X361),s(bool,X360))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X361),s(X362,X363))) | s(bool,X360) = s(bool,X364) | s(bool,'T_0') = s(bool,happ(s(fun(X365,bool),'const.combin.K_1'(s(bool,X364))),s(X365,X366)))) )),
  inference(superposition,[],[f194,f228])).
fof(f9723,plain,(
  ( ! [X356,X358,X357,X359,X355] : (s(bool,X355) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X356))),s(bool,X355))) | s(bool,'F_0') = s(bool,X356) | s(bool,X355) = s(bool,X357) | s(bool,'T_0') = s(bool,happ(s(fun(X358,bool),'const.combin.K_1'(s(bool,X357))),s(X358,X359)))) )),
  inference(superposition,[],[f184,f228])).
fof(f9722,plain,(
  ( ! [X350,X352,X354,X351,X353] : (s(bool,X350) = s(bool,'const.combin.K_2'(s(bool,X351),s(bool,X350))) | s(bool,'F_0') = s(bool,X351) | s(bool,X350) = s(bool,X352) | s(bool,'T_0') = s(bool,happ(s(fun(X353,bool),'const.combin.K_1'(s(bool,X352))),s(X353,X354)))) )),
  inference(superposition,[],[f146,f228])).
fof(f10096,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(condensation,[],[f10095])).
fof(f10095,plain,(
  ( ! [X288,X285,X287,X283,X284,X286] : (s(bool,X283) != s(bool,X284) | s(bool,'T_0') = s(bool,X285) | s(bool,'T_0') = s(bool,X284) | s(bool,X283) = s(bool,X286) | s(bool,'T_0') = s(bool,happ(s(fun(X287,bool),'const.combin.K_1'(s(bool,X286))),s(X287,X288)))) )),
  inference(duplicate_literal_removal,[],[f9711])).
fof(f9711,plain,(
  ( ! [X288,X285,X287,X283,X284,X286] : (s(bool,X283) != s(bool,X284) | s(bool,'T_0') = s(bool,X285) | s(bool,'T_0') = s(bool,X285) | s(bool,'T_0') = s(bool,X284) | s(bool,X283) = s(bool,X286) | s(bool,'T_0') = s(bool,happ(s(fun(X287,bool),'const.combin.K_1'(s(bool,X286))),s(X287,X288)))) )),
  inference(superposition,[],[f4399,f228])).
fof(f9679,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(bool,X49) != s(bool,'const.combin.K_2'(s(bool,X48),s(X50,X51))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X49),s(X52,X53))) | s(bool,X48) = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun(X55,bool),'const.combin.K_1'(s(bool,X54))),s(X55,X56)))) )),
  inference(superposition,[],[f1637,f228])).
fof(f9678,plain,(
  ( ! [X39,X47,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X40) != s(bool,'const.combin.K_2'(s(bool,X39),s(X41,X42))) | s(bool,'T_0') = s(bool,happ(s(fun(X43,bool),'const.combin.K_1'(s(bool,X40))),s(X43,X44))) | s(bool,X39) = s(bool,X45) | s(bool,'T_0') = s(bool,happ(s(fun(X46,bool),'const.combin.K_1'(s(bool,X45))),s(X46,X47)))) )),
  inference(superposition,[],[f1566,f228])).
fof(f9676,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (s(bool,X26) != s(bool,'const.combin.K_2'(s(bool,X25),s(X27,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X29) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),'const.combin.K_1'(s(bool,X29))),s(X30,X31)))) )),
  inference(superposition,[],[f1343,f228])).
fof(f9671,plain,(
  ( ! [X947,X951,X949,X952,X950,X948] : (s(X949,f(s(bool,X950))) != s(X949,g(s(bool,X950))) | s(fun(bool,X949),X947) = s(fun(bool,X949),X948) | s(bool,'T_0') = s(bool,sK9(X947,X948,X949,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X951,bool),'const.combin.K_1'(s(bool,X950))),s(X951,X952)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9658,plain,(
  ( ! [X881,X878,X882,X880,X879,X877] : (s(bool,X880) != s(bool,f(s(X877,sK9(X878,X879,bool,X877)))) | s(fun(X877,bool),X878) = s(fun(X877,bool),X879) | s(bool,'T_0') = s(bool,g(s(X877,sK9(X878,X879,bool,X877)))) | s(bool,'T_0') = s(bool,happ(s(fun(X881,bool),'const.combin.K_1'(s(bool,X880))),s(X881,X882)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9522,plain,(
  ( ! [X969,X967,X965,X970,X968,X966] : (s(X967,f(s(bool,X968))) != s(X967,g(s(bool,X968))) | s(fun(bool,X967),X965) = s(fun(bool,X967),X966) | s(bool,'T_0') = s(bool,X968) | s(bool,'T_0') = s(bool,happ(s(fun(X969,bool),'const.combin.K_1'(s(bool,sK9(X965,X966,X967,bool)))),s(X969,X970)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9509,plain,(
  ( ! [X895,X899,X897,X898,X896,X900] : (s(bool,X898) != s(bool,f(s(X895,sK9(X896,X897,bool,X895)))) | s(fun(X895,bool),X896) = s(fun(X895,bool),X897) | s(bool,'T_0') = s(bool,X898) | s(bool,'T_0') = s(bool,happ(s(fun(X899,bool),'const.combin.K_1'(s(bool,g(s(X895,sK9(X896,X897,bool,X895)))))),s(X899,X900)))) )),
  inference(superposition,[],[f116,f228])).
fof(f9326,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X156) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X157))),s(bool,X156))) | s(bool,X157) = s(bool,X158) | s(bool,'T_0') = s(bool,X158) | s(bool,X156) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f228,f144])).
fof(f9325,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (s(bool,X150) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X151))),s(bool,X150))) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(bool,X150) = s(bool,X153) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X153),s(X154,X155)))) )),
  inference(superposition,[],[f228,f227])).
fof(f9324,plain,(
  ( ! [X146,X144,X149,X147,X145,X148] : (s(bool,X144) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X145))),s(bool,X144))) | s(bool,X145) = s(bool,X146) | s(bool,'T_0') = s(bool,X146) | s(bool,X144) = s(bool,X147) | s(bool,'T_0') = s(bool,happ(s(fun(X148,bool),'const.combin.K_1'(s(bool,X147))),s(X148,X149)))) )),
  inference(superposition,[],[f228,f228])).
fof(f9296,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (s(bool,'T_0') = s(bool,happ(s(fun(X325,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X322),s(bool,X323),s(bool,X324))))),s(X325,X326))) | s(bool,X327) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X322))),s(bool,X323))),s(bool,X324))) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f228,f113])).
fof(f9295,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (s(bool,'T_0') = s(bool,happ(s(fun(X319,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X316),s(bool,X317),s(bool,X318))))),s(X319,X320))) | s(bool,X321) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X316))),s(bool,X317))),s(bool,X318))) | s(bool,'T_0') = s(bool,X321)) )),
  inference(superposition,[],[f228,f112])).
fof(f9294,plain,(
  ( ! [X313,X315,X310,X312,X314,X311] : (s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),'const.combin.K_1'(s(bool,X311))),s(X313,X314))) | s(bool,X315) = s(bool,happ(s(fun(X310,bool),'const.combin.K_1'(s(bool,X311))),s(X310,X312))) | s(bool,'T_0') = s(bool,X315)) )),
  inference(superposition,[],[f228,f127])).
fof(f9293,plain,(
  ( ! [X308,X306,X309,X305,X307] : (s(bool,'T_0') = s(bool,happ(s(fun(X307,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X306))))),s(X307,X308))) | s(bool,X309) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X305),s('type.semi_ring.semi_ring'(bool),X306))))) | s(bool,'T_0') = s(bool,X309)) )),
  inference(superposition,[],[f228,f92])).
fof(f9291,plain,(
  ( ! [X296,X298,X295,X297,X299] : (s(bool,'T_0') = s(bool,happ(s(fun(X297,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X296))))),s(X297,X298))) | s(bool,X299) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X295),s('type.semi_ring.semi_ring'(bool),X296))))) | s(bool,'T_0') = s(bool,X299)) )),
  inference(superposition,[],[f228,f97])).
fof(f9290,plain,(
  ( ! [X292,X294,X290,X293,X291] : (s(bool,'T_0') = s(bool,happ(s(fun(X292,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X291))))),s(X292,X293))) | s(bool,X294) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X290),s('type.semi_ring.semi_ring'(bool),X291))))) | s(bool,'T_0') = s(bool,X294)) )),
  inference(superposition,[],[f228,f98])).
fof(f9288,plain,(
  ( ! [X280,X282,X281,X283,X284] : (s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X281))))),s(X282,X283))) | s(bool,X284) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X280),s('type.semi_ring.semi_ring'(bool),X281))))) | s(bool,'T_0') = s(bool,X284)) )),
  inference(superposition,[],[f228,f91])).
fof(f9287,plain,(
  ( ! [X277,X279,X275,X276,X278] : (s(bool,'T_0') = s(bool,happ(s(fun(X277,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X276))))),s(X277,X278))) | s(bool,X279) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X275),s('type.semi_ring.semi_ring'(bool),X276))))) | s(bool,'T_0') = s(bool,X279)) )),
  inference(superposition,[],[f228,f96])).
fof(f9286,plain,(
  ( ! [X271,X273,X270,X272,X274] : (s(bool,'T_0') = s(bool,happ(s(fun(X272,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X271))))),s(X272,X273))) | s(bool,X274) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X270),s('type.semi_ring.semi_ring'(bool),X271))))) | s(bool,'T_0') = s(bool,X274)) )),
  inference(superposition,[],[f228,f95])).
fof(f9244,plain,(
  ( ! [X241,X239,X237,X242,X240,X238] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X237),s(bool,X238),s(bool,X239))) | s(bool,X240) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X237))),s(bool,X238))),s(bool,X239))) | s(bool,'T_0') = s(bool,happ(s(fun(X241,bool),'const.combin.K_1'(s(bool,X240))),s(X241,X242)))) )),
  inference(superposition,[],[f228,f113])).
fof(f9243,plain,(
  ( ! [X231,X235,X233,X236,X234,X232] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X231),s(bool,X232),s(bool,X233))) | s(bool,X234) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X231))),s(bool,X232))),s(bool,X233))) | s(bool,'T_0') = s(bool,happ(s(fun(X235,bool),'const.combin.K_1'(s(bool,X234))),s(X235,X236)))) )),
  inference(superposition,[],[f228,f112])).
fof(f9241,plain,(
  ( ! [X222,X220,X223,X221,X224] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X221))) | s(bool,X222) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X220),s('type.semi_ring.semi_ring'(bool),X221))))) | s(bool,'T_0') = s(bool,happ(s(fun(X223,bool),'const.combin.K_1'(s(bool,X222))),s(X223,X224)))) )),
  inference(superposition,[],[f228,f92])).
fof(f9240,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X215),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X216))))) | s(bool,X217) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X215),s('type.semi_ring.semi_ring'(bool),X216))))) | s(bool,'T_0') = s(bool,happ(s(fun(X218,bool),'const.combin.K_1'(s(bool,X217))),s(X218,X219)))) )),
  inference(superposition,[],[f228,f93])).
fof(f9239,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X211))) | s(bool,X212) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X210),s('type.semi_ring.semi_ring'(bool),X211))))) | s(bool,'T_0') = s(bool,happ(s(fun(X213,bool),'const.combin.K_1'(s(bool,X212))),s(X213,X214)))) )),
  inference(superposition,[],[f228,f97])).
fof(f9238,plain,(
  ( ! [X208,X206,X209,X207,X205] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X206))) | s(bool,X207) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X205),s('type.semi_ring.semi_ring'(bool),X206))))) | s(bool,'T_0') = s(bool,happ(s(fun(X208,bool),'const.combin.K_1'(s(bool,X207))),s(X208,X209)))) )),
  inference(superposition,[],[f228,f98])).
fof(f9237,plain,(
  ( ! [X204,X202,X200,X203,X201] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X200),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X201))))) | s(bool,X202) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X200),s('type.semi_ring.semi_ring'(bool),X201))))) | s(bool,'T_0') = s(bool,happ(s(fun(X203,bool),'const.combin.K_1'(s(bool,X202))),s(X203,X204)))) )),
  inference(superposition,[],[f228,f94])).
fof(f9236,plain,(
  ( ! [X198,X196,X199,X197,X195] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X196))) | s(bool,X197) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X195),s('type.semi_ring.semi_ring'(bool),X196))))) | s(bool,'T_0') = s(bool,happ(s(fun(X198,bool),'const.combin.K_1'(s(bool,X197))),s(X198,X199)))) )),
  inference(superposition,[],[f228,f91])).
fof(f9235,plain,(
  ( ! [X191,X194,X192,X190,X193] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X191))) | s(bool,X192) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X190),s('type.semi_ring.semi_ring'(bool),X191))))) | s(bool,'T_0') = s(bool,happ(s(fun(X193,bool),'const.combin.K_1'(s(bool,X192))),s(X193,X194)))) )),
  inference(superposition,[],[f228,f96])).
fof(f9234,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X186))) | s(bool,X187) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X185),s('type.semi_ring.semi_ring'(bool),X186))))) | s(bool,'T_0') = s(bool,happ(s(fun(X188,bool),'const.combin.K_1'(s(bool,X187))),s(X188,X189)))) )),
  inference(superposition,[],[f228,f95])).
fof(f9200,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),'const.combin.K_1'(s(bool,X18))),s(X19,X20))) | s(bool,X16) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f228,f144])).
fof(f9199,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),'const.combin.K_1'(s(bool,X10))),s(X11,X12))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15)))) )),
  inference(superposition,[],[f228,f227])).
fof(f9198,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X5))),s(X6,X7)))) )),
  inference(superposition,[],[f228,f228])).
fof(f9197,plain,(
  ( ! [X308,X310,X306,X309,X311,X307] : (s(bool,X309) = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X306),s(bool,X307),s(bool,X308))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X306))),s(bool,X307))),s(bool,X308))) | s(bool,'T_0') = s(bool,happ(s(fun(X310,bool),'const.combin.K_1'(s(bool,X309))),s(X310,X311)))) )),
  inference(superposition,[],[f228,f113])).
fof(f9196,plain,(
  ( ! [X300,X302,X304,X301,X303,X305] : (s(bool,X303) = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X300),s(bool,X301),s(bool,X302))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X300))),s(bool,X301))),s(bool,X302))) | s(bool,'T_0') = s(bool,happ(s(fun(X304,bool),'const.combin.K_1'(s(bool,X303))),s(X304,X305)))) )),
  inference(superposition,[],[f228,f112])).
fof(f9194,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,X291) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X290))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X289),s('type.semi_ring.semi_ring'(bool),X290))))) | s(bool,'T_0') = s(bool,happ(s(fun(X292,bool),'const.combin.K_1'(s(bool,X291))),s(X292,X293)))) )),
  inference(superposition,[],[f228,f92])).
fof(f9193,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,X286) = s(bool,happ(s(fun(bool,bool),X284),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X285))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X284),s('type.semi_ring.semi_ring'(bool),X285))))) | s(bool,'T_0') = s(bool,happ(s(fun(X287,bool),'const.combin.K_1'(s(bool,X286))),s(X287,X288)))) )),
  inference(superposition,[],[f228,f93])).
fof(f9192,plain,(
  ( ! [X280,X282,X279,X281,X283] : (s(bool,X281) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X280))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X279),s('type.semi_ring.semi_ring'(bool),X280))))) | s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'const.combin.K_1'(s(bool,X281))),s(X282,X283)))) )),
  inference(superposition,[],[f228,f97])).
fof(f9191,plain,(
  ( ! [X277,X275,X276,X278,X274] : (s(bool,X276) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X275))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X274),s('type.semi_ring.semi_ring'(bool),X275))))) | s(bool,'T_0') = s(bool,happ(s(fun(X277,bool),'const.combin.K_1'(s(bool,X276))),s(X277,X278)))) )),
  inference(superposition,[],[f228,f98])).
fof(f9190,plain,(
  ( ! [X269,X271,X273,X270,X272] : (s(bool,X271) = s(bool,happ(s(fun(bool,bool),X269),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X270))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X269),s('type.semi_ring.semi_ring'(bool),X270))))) | s(bool,'T_0') = s(bool,happ(s(fun(X272,bool),'const.combin.K_1'(s(bool,X271))),s(X272,X273)))) )),
  inference(superposition,[],[f228,f94])).
fof(f9189,plain,(
  ( ! [X265,X267,X268,X264,X266] : (s(bool,X266) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X265))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X264),s('type.semi_ring.semi_ring'(bool),X265))))) | s(bool,'T_0') = s(bool,happ(s(fun(X267,bool),'const.combin.K_1'(s(bool,X266))),s(X267,X268)))) )),
  inference(superposition,[],[f228,f91])).
fof(f9188,plain,(
  ( ! [X261,X263,X259,X260,X262] : (s(bool,X261) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X260))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X259),s('type.semi_ring.semi_ring'(bool),X260))))) | s(bool,'T_0') = s(bool,happ(s(fun(X262,bool),'const.combin.K_1'(s(bool,X261))),s(X262,X263)))) )),
  inference(superposition,[],[f228,f96])).
fof(f9187,plain,(
  ( ! [X255,X257,X254,X256,X258] : (s(bool,X256) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X255))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X254),s('type.semi_ring.semi_ring'(bool),X255))))) | s(bool,'T_0') = s(bool,happ(s(fun(X257,bool),'const.combin.K_1'(s(bool,X256))),s(X257,X258)))) )),
  inference(superposition,[],[f228,f95])).
fof(f8966,plain,(
  ( ! [X158,X161,X159,X157,X160] : (s(bool,X157) != s(bool,'const.combin.K_2'(s(bool,X158),s(bool,X157))) | s(bool,'T_0') = s(bool,happ(s(fun(X159,bool),'const.combin.K_1'(s(bool,X158))),s(X159,X160))) | s(bool,X157) = s(bool,X161) | s(bool,'T_0') = s(bool,X161)) )),
  inference(superposition,[],[f2934,f144])).
fof(f8965,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (s(bool,X150) != s(bool,'const.combin.K_2'(s(bool,X151),s(bool,X150))) | s(bool,'T_0') = s(bool,happ(s(fun(X152,bool),'const.combin.K_1'(s(bool,X151))),s(X152,X153))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X154),s(X155,X156)))) )),
  inference(superposition,[],[f2934,f227])).
fof(f8851,plain,(
  ( ! [X12,X10,X13,X11] : (s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(fun(X10,fun(X10,X10))),X11),s(fun(X10,fun(X10,X10)),X12),s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X10),X13))))) = s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X10),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X10,fun(X10,X10)),fun(X10,fun(X10,X10))),happ(s(fun(fun(X10,fun(X10,X10)),fun(fun(X10,fun(X10,X10)),fun(X10,fun(X10,X10)))),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(fun(X10,fun(X10,X10))),X11))),s(fun(X10,fun(X10,X10)),X12))),s('type.semi_ring.semi_ring'(X10),X13)))))) )),
  inference(superposition,[],[f106,f113])).
fof(f8850,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(fun(X6,fun(X6,X6))),X7),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X9))))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),happ(s(fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6)))),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(fun(X6,fun(X6,X6))),X7))),s(fun(X6,fun(X6,X6)),X8))),s('type.semi_ring.semi_ring'(X6),X9)))))) )),
  inference(superposition,[],[f106,f112])).
fof(f8849,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,fun(X3,X3)),X4) = s(fun(X3,fun(X3,X3)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X3),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X3,fun(X3,X3)),fun(X3,fun(X3,X3))),'const.combin.K_1'(s(fun(X3,fun(X3,X3)),X4))),s('type.semi_ring.semi_ring'(X3),X5)))))) )),
  inference(superposition,[],[f106,f127])).
fof(f8892,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),X10))))) = s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s('type.semi_ring.semi_ring'(X8),X10)))))))) )),
  inference(forward_demodulation,[],[f8848,f106])).
fof(f8848,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,fun(X8,X8)),happ(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X8),X10))))) = s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s('type.semi_ring.semi_ring'(X8),X10)))))))) )),
  inference(superposition,[],[f106,f104])).
fof(f8891,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X4,X4),X5),s('type.semi_ring.semi_ring'(X4),X6)))))))) )),
  inference(forward_demodulation,[],[f8847,f106])).
fof(f8847,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),happ(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X4,X4),X5),s('type.semi_ring.semi_ring'(X4),X6)))))))) )),
  inference(superposition,[],[f106,f100])).
fof(f8890,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2)))))))) )),
  inference(forward_demodulation,[],[f8846,f106])).
fof(f8846,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2)))))))) )),
  inference(superposition,[],[f106,f99])).
fof(f8778,plain,(
  ( ! [X924,X928,X923,X925,X929] : (s(X925,f(s(bool,'F_0'))) != s(X925,g(s(bool,'F_0'))) | s(fun(bool,X925),X923) = s(fun(bool,X925),X924) | s(bool,'T_0') = s(bool,happ(s(fun(X928,bool),'const.combin.K_1'(s(bool,sK9(X923,X924,X925,bool)))),s(X928,X929)))) )),
  inference(forward_demodulation,[],[f8356,f114])).
fof(f8356,plain,(
  ( ! [X926,X924,X928,X923,X927,X925,X929] : (s(X925,f(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X926,X927))))) != s(X925,g(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X926,X927))))) | s(fun(bool,X925),X923) = s(fun(bool,X925),X924) | s(bool,'T_0') = s(bool,happ(s(fun(X928,bool),'const.combin.K_1'(s(bool,sK9(X923,X924,X925,bool)))),s(X928,X929)))) )),
  inference(superposition,[],[f116,f2421])).
fof(f8755,plain,(
  ( ! [X841,X845,X842,X840,X846] : (s(bool,'F_0') != s(bool,f(s(X840,sK9(X841,X842,bool,X840)))) | s(fun(X840,bool),X841) = s(fun(X840,bool),X842) | s(bool,'T_0') = s(bool,happ(s(fun(X845,bool),'const.combin.K_1'(s(bool,g(s(X840,sK9(X841,X842,bool,X840)))))),s(X845,X846)))) )),
  inference(forward_demodulation,[],[f8343,f114])).
fof(f8343,plain,(
  ( ! [X843,X841,X845,X842,X840,X846,X844] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X843,X844))) != s(bool,f(s(X840,sK9(X841,X842,bool,X840)))) | s(fun(X840,bool),X841) = s(fun(X840,bool),X842) | s(bool,'T_0') = s(bool,happ(s(fun(X845,bool),'const.combin.K_1'(s(bool,g(s(X840,sK9(X841,X842,bool,X840)))))),s(X845,X846)))) )),
  inference(superposition,[],[f116,f2421])).
fof(f8194,plain,(
  ( ! [X158,X161,X159,X157,X160] : (s(bool,X157) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X158))),s(bool,X157))) | s(bool,X158) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X159,X160))) | s(bool,X157) = s(bool,X161) | s(bool,'T_0') = s(bool,X161)) )),
  inference(superposition,[],[f2421,f144])).
fof(f8193,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (s(bool,X150) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X151))),s(bool,X150))) | s(bool,X151) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X152,X153))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X154),s(X155,X156)))) )),
  inference(superposition,[],[f2421,f227])).
fof(f8097,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),'const.combin.K_1'(s(bool,X2))),s(X6,X7)))) )),
  inference(superposition,[],[f2421,f2421])).
fof(f7968,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X57))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X58),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X56),s('type.semi_ring.semi_ring'(bool),X57)))))))) )),
  inference(superposition,[],[f1652,f92])).
fof(f7967,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X53),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X54))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X55),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X53),s('type.semi_ring.semi_ring'(bool),X54)))))))) )),
  inference(superposition,[],[f1652,f93])).
fof(f7966,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X51))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X52),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X50),s('type.semi_ring.semi_ring'(bool),X51)))))))) )),
  inference(superposition,[],[f1652,f97])).
fof(f7965,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X48))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X49),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X47),s('type.semi_ring.semi_ring'(bool),X48)))))))) )),
  inference(superposition,[],[f1652,f98])).
fof(f7972,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X2),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),X1))))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X1)))) )),
  inference(trivial_inequality_removal,[],[f7953])).
fof(f7953,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X2),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),X1))))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X1)))) )),
  inference(superposition,[],[f1652,f522])).
fof(f7901,plain,(
  ( ! [X12,X10,X13,X11] : (s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(fun(X10,fun(X10,X10))),X11),s(fun(X10,fun(X10,X10)),X12),s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X10),X13))))) = s(fun(X10,fun(X10,X10)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X10),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X10,fun(X10,X10)),fun(X10,fun(X10,X10))),happ(s(fun(fun(X10,fun(X10,X10)),fun(fun(X10,fun(X10,X10)),fun(X10,fun(X10,X10)))),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(fun(X10,fun(X10,X10))),X11))),s(fun(X10,fun(X10,X10)),X12))),s('type.semi_ring.semi_ring'(X10),X13)))))) )),
  inference(superposition,[],[f105,f113])).
fof(f7900,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(fun(X6,fun(X6,X6))),X7),s(fun(X6,fun(X6,X6)),X8),s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X9))))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),happ(s(fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6)))),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(fun(X6,fun(X6,X6))),X7))),s(fun(X6,fun(X6,X6)),X8))),s('type.semi_ring.semi_ring'(X6),X9)))))) )),
  inference(superposition,[],[f105,f112])).
fof(f7899,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,fun(X3,X3)),X4) = s(fun(X3,fun(X3,X3)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X3),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X3,fun(X3,X3)),fun(X3,fun(X3,X3))),'const.combin.K_1'(s(fun(X3,fun(X3,X3)),X4))),s('type.semi_ring.semi_ring'(X3),X5)))))) )),
  inference(superposition,[],[f105,f127])).
fof(f7940,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),X10))))) = s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s('type.semi_ring.semi_ring'(X8),X10)))))))) )),
  inference(forward_demodulation,[],[f7898,f105])).
fof(f7898,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,fun(X8,X8)),happ(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X8),X10))))) = s(fun(X8,fun(X8,X8)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X11),s('type.semi_ring.semi_ring'(X8),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X8,fun(X8,X8)),fun(X8,fun(X8,X8))),X9),s('type.semi_ring.semi_ring'(X8),X10)))))))) )),
  inference(superposition,[],[f105,f103])).
fof(f7939,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X4,X4),X5),s('type.semi_ring.semi_ring'(X4),X6)))))))) )),
  inference(forward_demodulation,[],[f7897,f105])).
fof(f7897,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),happ(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.semi_ring.semi_ring'(X4),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X4,X4),X5),s('type.semi_ring.semi_ring'(X4),X6)))))))) )),
  inference(superposition,[],[f105,f102])).
fof(f7938,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2)))))))) )),
  inference(forward_demodulation,[],[f7896,f105])).
fof(f7896,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2)))))))) )),
  inference(superposition,[],[f105,f101])).
fof(f7854,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X56),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X57))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X58),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X56),s('type.semi_ring.semi_ring'(bool),X57)))))))) )),
  inference(superposition,[],[f1650,f94])).
fof(f7853,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X54))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X55),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X53),s('type.semi_ring.semi_ring'(bool),X54)))))))) )),
  inference(superposition,[],[f1650,f91])).
fof(f7852,plain,(
  ( ! [X52,X50,X51] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X51))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X52),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X50),s('type.semi_ring.semi_ring'(bool),X51)))))))) )),
  inference(superposition,[],[f1650,f96])).
fof(f7851,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X48))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X49),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X47),s('type.semi_ring.semi_ring'(bool),X48)))))))) )),
  inference(superposition,[],[f1650,f95])).
fof(f7858,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X2),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),X1))))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X1)))) )),
  inference(trivial_inequality_removal,[],[f7839])).
fof(f7839,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X2),s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X0),s('type.semi_ring.semi_ring'(bool),X1))))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X1)))) )),
  inference(superposition,[],[f1650,f468])).
fof(f7718,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),X2),s(X0,X3),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(forward_demodulation,[],[f7684,f113])).
fof(f7684,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3))),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(superposition,[],[f113,f104])).
fof(f7475,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(bool,X10) != s(bool,'const.combin.K_2'(s(bool,X9),s(X11,X12))) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),'const.combin.K_1'(s(bool,X10))),s(X13,X14))) | s(bool,X9) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1566,f144])).
fof(f7474,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X1) != s(bool,'const.combin.K_2'(s(bool,X0),s(X2,X3))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X1))),s(X4,X5))) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X6),s(X7,X8)))) )),
  inference(superposition,[],[f1566,f227])).
fof(f7635,plain,(
  ( ! [X310,X304,X306,X309,X305] : (s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X304),s(bool,X305),s(bool,X306))))),s(X309,X310))) | s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X304),s(bool,X305),s(bool,X306)))) )),
  inference(forward_demodulation,[],[f7634,f113])).
fof(f7634,plain,(
  ( ! [X310,X304,X306,X309,X305] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X304),s(bool,X305),s(bool,X306))) | s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X304))),s(bool,X305))),s(bool,X306))))),s(X309,X310)))) )),
  inference(forward_demodulation,[],[f7473,f114])).
fof(f7473,plain,(
  ( ! [X308,X310,X304,X306,X309,X305,X307] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X307,X308))) != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X304),s(bool,X305),s(bool,X306))) | s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X304))),s(bool,X305))),s(bool,X306))))),s(X309,X310)))) )),
  inference(superposition,[],[f1566,f113])).
fof(f7633,plain,(
  ( ! [X302,X298,X303,X297,X299] : (s(bool,'T_0') = s(bool,happ(s(fun(X302,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X297),s(bool,X298),s(bool,X299))))),s(X302,X303))) | s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X297),s(bool,X298),s(bool,X299)))) )),
  inference(forward_demodulation,[],[f7632,f112])).
fof(f7632,plain,(
  ( ! [X302,X298,X303,X297,X299] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X297),s(bool,X298),s(bool,X299))) | s(bool,'T_0') = s(bool,happ(s(fun(X302,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X297))),s(bool,X298))),s(bool,X299))))),s(X302,X303)))) )),
  inference(forward_demodulation,[],[f7472,f114])).
fof(f7472,plain,(
  ( ! [X300,X302,X298,X301,X303,X297,X299] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X300,X301))) != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X297),s(bool,X298),s(bool,X299))) | s(bool,'T_0') = s(bool,happ(s(fun(X302,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X297))),s(bool,X298))),s(bool,X299))))),s(X302,X303)))) )),
  inference(superposition,[],[f1566,f112])).
fof(f7628,plain,(
  ( ! [X282,X279,X283,X278] : (s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),X278),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X279))))))),s(X282,X283))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X278),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X279)))))) )),
  inference(forward_demodulation,[],[f7627,f93])).
fof(f7627,plain,(
  ( ! [X282,X279,X283,X278] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X278),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X279))))) | s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X278),s('type.semi_ring.semi_ring'(bool),X279))))))),s(X282,X283)))) )),
  inference(forward_demodulation,[],[f7469,f114])).
fof(f7469,plain,(
  ( ! [X280,X282,X279,X281,X283,X278] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X280,X281))) != s(bool,happ(s(fun(bool,bool),X278),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X279))))) | s(bool,'T_0') = s(bool,happ(s(fun(X282,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X278),s('type.semi_ring.semi_ring'(bool),X279))))))),s(X282,X283)))) )),
  inference(superposition,[],[f1566,f93])).
fof(f7624,plain,(
  ( ! [X271,X267,X270] : (s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X267))))),s(X270,X271))) | s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X267)))) )),
  inference(forward_demodulation,[],[f7623,f98])).
fof(f7623,plain,(
  ( ! [X271,X267,X270,X266] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X267))) | s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X266),s('type.semi_ring.semi_ring'(bool),X267))))))),s(X270,X271)))) )),
  inference(forward_demodulation,[],[f7467,f114])).
fof(f7467,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X267))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X268,X269))) | s(bool,'T_0') = s(bool,happ(s(fun(X270,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X266),s('type.semi_ring.semi_ring'(bool),X267))))))),s(X270,X271)))) )),
  inference(superposition,[],[f1566,f98])).
fof(f7622,plain,(
  ( ! [X261,X265,X260,X264] : (s(bool,'T_0') = s(bool,happ(s(fun(X264,bool),'const.combin.K_1'(s(bool,happ(s(fun(bool,bool),X260),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X261))))))),s(X264,X265))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X260),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X261)))))) )),
  inference(forward_demodulation,[],[f7621,f94])).
fof(f7621,plain,(
  ( ! [X261,X265,X260,X264] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X260),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X261))))) | s(bool,'T_0') = s(bool,happ(s(fun(X264,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X260),s('type.semi_ring.semi_ring'(bool),X261))))))),s(X264,X265)))) )),
  inference(forward_demodulation,[],[f7466,f114])).
fof(f7466,plain,(
  ( ! [X261,X263,X265,X260,X262,X264] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X262,X263))) != s(bool,happ(s(fun(bool,bool),X260),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X261))))) | s(bool,'T_0') = s(bool,happ(s(fun(X264,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X260),s('type.semi_ring.semi_ring'(bool),X261))))))),s(X264,X265)))) )),
  inference(superposition,[],[f1566,f94])).
fof(f7616,plain,(
  ( ! [X243,X246,X247] : (s(bool,'T_0') = s(bool,happ(s(fun(X246,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X243))))),s(X246,X247))) | s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X243)))) )),
  inference(forward_demodulation,[],[f7615,f95])).
fof(f7615,plain,(
  ( ! [X243,X246,X242,X247] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X243))) | s(bool,'T_0') = s(bool,happ(s(fun(X246,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X242),s('type.semi_ring.semi_ring'(bool),X243))))))),s(X246,X247)))) )),
  inference(forward_demodulation,[],[f7463,f114])).
fof(f7463,plain,(
  ( ! [X243,X246,X244,X242,X247,X245] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X243))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X244,X245))) | s(bool,'T_0') = s(bool,happ(s(fun(X246,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X242),s('type.semi_ring.semi_ring'(bool),X243))))))),s(X246,X247)))) )),
  inference(superposition,[],[f1566,f95])).
fof(f6932,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s(bool,X149) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X150))),s(bool,X149))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X150),s(X151,X152))) | s(bool,X149) = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f601,f144])).
fof(f6931,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X148] : (s(bool,X142) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X143))),s(bool,X142))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X143),s(X144,X145))) | s(bool,X142) = s(bool,X146) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X146),s(X147,X148)))) )),
  inference(superposition,[],[f601,f227])).
fof(f6863,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X2),s(X0,X3),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(forward_demodulation,[],[f6832,f112])).
fof(f6832,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3))),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(superposition,[],[f112,f103])).
fof(f6719,plain,(
  ( ! [X368,X370,X367,X369,X371] : (s(bool,'F_0') = s(bool,happ(s(fun(X370,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X367),s(bool,X368),s(bool,X369))))),s(X370,X371))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X367))),s(bool,X368))),s(bool,X369)))) )),
  inference(superposition,[],[f5864,f113])).
fof(f6716,plain,(
  ( ! [X350,X352,X354,X351,X353] : (s(bool,X353) != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X350),s(bool,X351),s(bool,X352))) | s(bool,X354) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X350))),s(bool,X351))),s(bool,X352))) | s(bool,'T_0') = s(bool,X354) | s(bool,'T_0') = s(bool,X353)) )),
  inference(superposition,[],[f4399,f113])).
fof(f6715,plain,(
  ( ! [X346,X349,X345,X347,X348] : (s(bool,X348) != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X345),s(bool,X346),s(bool,X347))) | s(bool,X348) = s(bool,X349) | s(bool,'T_0') = s(bool,X349) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X345))),s(bool,X346))),s(bool,X347)))) )),
  inference(superposition,[],[f4399,f113])).
fof(f6713,plain,(
  ( ! [X335,X337,X339,X340,X336,X338] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X335),s(bool,X336),s(bool,X337))),s(X338,X339))) | s(bool,X340) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X335))),s(bool,X336))),s(bool,X337))) | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f4151,f113])).
fof(f6708,plain,(
  ( ! [X308,X306,X309,X305,X307] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X305),s(bool,X306),s(bool,X307))),s(X308,X309))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X305))),s(bool,X306))),s(bool,X307)))) )),
  inference(superposition,[],[f2374,f113])).
fof(f6786,plain,(
  ( ! [X285,X283,X284] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X283),s(bool,X284),s(bool,X285))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X283))),s(bool,X284))),s(bool,X285)))) )),
  inference(forward_demodulation,[],[f6704,f114])).
fof(f6704,plain,(
  ( ! [X285,X287,X283,X284,X286] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X286,X287))) != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X283),s(bool,X284),s(bool,X285))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X283))),s(bool,X284))),s(bool,X285)))) )),
  inference(superposition,[],[f1343,f113])).
fof(f6703,plain,(
  ( ! [X280,X282,X277,X279,X281,X278] : (s(bool,X280) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X277),s(bool,X278),s(bool,X279))),s(X281,X282))) | s(bool,X280) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X277))),s(bool,X278))),s(bool,X279))) | s(bool,'T_0') = s(bool,X280)) )),
  inference(superposition,[],[f1331,f113])).
fof(f6702,plain,(
  ( ! [X271,X273,X275,X276,X272,X274] : (s(bool,'const.combin.K_2'(s(bool,X274),s(X275,X276))) != s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X271),s(bool,X272),s(bool,X273))) | s(bool,X274) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X271))),s(bool,X272))),s(bool,X273))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X271))),s(bool,X272))),s(bool,X273)))) )),
  inference(superposition,[],[f1331,f113])).
fof(f6701,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X266),s(bool,X267),s(bool,X268))),s(X269,X270))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X266))),s(bool,X267))),s(bool,X268)))) )),
  inference(superposition,[],[f747,f113])).
fof(f6698,plain,(
  ( ! [X255,X253,X251,X254,X252,X250] : (s(bool,X253) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X250),s(bool,X251),s(bool,X252))),s(X254,X255))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X250))),s(bool,X251))),s(bool,X252))) | s(bool,'T_0') = s(bool,X253)) )),
  inference(superposition,[],[f250,f113])).
fof(f6694,plain,(
  ( ! [X231,X229,X227,X230,X228,X226] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X226),s(bool,X227),s(bool,X228))),s(X229,X230))) | s(bool,X231) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X226))),s(bool,X227))),s(bool,X228))) | s(bool,'T_0') = s(bool,X231)) )),
  inference(superposition,[],[f227,f113])).
fof(f6689,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,happ(s(fun(X196,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X193),s(bool,X194),s(bool,X195))))),s(X196,X197))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X193))),s(bool,X194))),s(bool,X195)))) )),
  inference(superposition,[],[f184,f113])).
fof(f6687,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X185),s(bool,X186),s(bool,X187))),s(X188,X189))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X185))),s(bool,X186))),s(bool,X187)))) )),
  inference(superposition,[],[f146,f113])).
fof(f6651,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,'const.combin.K_2'(s(bool,X44),s(X45,X46))) = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X41),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X41))),s(bool,X42))),s(bool,X43)))) )),
  inference(superposition,[],[f113,f250])).
fof(f6650,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X35),s(bool,X36),s(bool,X37))) | s(bool,X38) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X35))),s(bool,X36))),s(bool,X37))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X38),s(X39,X40)))) )),
  inference(superposition,[],[f113,f227])).
fof(f6649,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,X32) = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X29),s(bool,X30),s(bool,X31))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X29))),s(bool,X30))),s(bool,X31))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X32),s(X33,X34)))) )),
  inference(superposition,[],[f113,f227])).
fof(f6645,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X11),s(bool,X12),s(bool,X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X11))),s(bool,X12))),s(bool,X13)))) )),
  inference(superposition,[],[f113,f144])).
fof(f6644,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X7),s(bool,X8),s(bool,X9))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X7))),s(bool,X8))),s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f113,f144])).
fof(f6643,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(bool),X4),s(bool,X5),s(bool,X6))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(bool),X4))),s(bool,X5))),s(bool,X6)))) )),
  inference(superposition,[],[f113,f90])).
fof(f6774,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X5),X7),s(X5,X8),s(X5,X9))) = s(X5,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X5),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X5,X5),X6),s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8),s(X5,X9)))) )),
  inference(forward_demodulation,[],[f6544,f113])).
fof(f6544,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X5),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X5,X5),X6),s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8),s(X5,X9))) = s(X5,happ(s(fun(X5,X5),happ(s(fun(X5,fun(X5,X5)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8))),s(X5,X9)))) )),
  inference(superposition,[],[f113,f100])).
fof(f6773,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),X2),s(X0,X3),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(forward_demodulation,[],[f6543,f113])).
fof(f6543,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3))),s(X0,X4)))) )),
  inference(superposition,[],[f113,f99])).
fof(f6435,plain,(
  ( ! [X356,X358,X357,X359,X355] : (s(bool,'F_0') = s(bool,happ(s(fun(X358,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X355),s(bool,X356),s(bool,X357))))),s(X358,X359))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X355))),s(bool,X356))),s(bool,X357)))) )),
  inference(superposition,[],[f5864,f112])).
fof(f6432,plain,(
  ( ! [X341,X339,X340,X342,X338] : (s(bool,X341) != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X338),s(bool,X339),s(bool,X340))) | s(bool,X342) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X338))),s(bool,X339))),s(bool,X340))) | s(bool,'T_0') = s(bool,X342) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f4399,f112])).
fof(f6431,plain,(
  ( ! [X333,X335,X337,X334,X336] : (s(bool,X336) != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X333),s(bool,X334),s(bool,X335))) | s(bool,X336) = s(bool,X337) | s(bool,'T_0') = s(bool,X337) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X333))),s(bool,X334))),s(bool,X335)))) )),
  inference(superposition,[],[f4399,f112])).
fof(f6429,plain,(
  ( ! [X325,X327,X323,X324,X326,X328] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X323),s(bool,X324),s(bool,X325))),s(X326,X327))) | s(bool,X328) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X323))),s(bool,X324))),s(bool,X325))) | s(bool,'T_0') = s(bool,X328)) )),
  inference(superposition,[],[f4151,f112])).
fof(f6424,plain,(
  ( ! [X294,X296,X293,X295,X297] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X293),s(bool,X294),s(bool,X295))),s(X296,X297))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X293))),s(bool,X294))),s(bool,X295)))) )),
  inference(superposition,[],[f2374,f112])).
fof(f6501,plain,(
  ( ! [X271,X273,X272] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X271),s(bool,X272),s(bool,X273))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X271))),s(bool,X272))),s(bool,X273)))) )),
  inference(forward_demodulation,[],[f6420,f114])).
fof(f6420,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X274,X275))) != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X271),s(bool,X272),s(bool,X273))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X271))),s(bool,X272))),s(bool,X273)))) )),
  inference(superposition,[],[f1343,f112])).
fof(f6419,plain,(
  ( ! [X269,X265,X267,X268,X270,X266] : (s(bool,X268) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X265),s(bool,X266),s(bool,X267))),s(X269,X270))) | s(bool,X268) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X265))),s(bool,X266))),s(bool,X267))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f1331,f112])).
fof(f6418,plain,(
  ( ! [X261,X263,X259,X260,X262,X264] : (s(bool,'const.combin.K_2'(s(bool,X262),s(X263,X264))) != s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X259),s(bool,X260),s(bool,X261))) | s(bool,X262) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X259))),s(bool,X260))),s(bool,X261))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X259))),s(bool,X260))),s(bool,X261)))) )),
  inference(superposition,[],[f1331,f112])).
fof(f6417,plain,(
  ( ! [X255,X257,X254,X256,X258] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X254),s(bool,X255),s(bool,X256))),s(X257,X258))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X254))),s(bool,X255))),s(bool,X256)))) )),
  inference(superposition,[],[f747,f112])).
fof(f6414,plain,(
  ( ! [X243,X241,X239,X242,X240,X238] : (s(bool,X241) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X238),s(bool,X239),s(bool,X240))),s(X242,X243))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X238))),s(bool,X239))),s(bool,X240))) | s(bool,'T_0') = s(bool,X241)) )),
  inference(superposition,[],[f250,f112])).
fof(f6410,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X214),s(bool,X215),s(bool,X216))),s(X217,X218))) | s(bool,X219) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X214))),s(bool,X215))),s(bool,X216))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f227,f112])).
fof(f6405,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(bool,'T_0') = s(bool,happ(s(fun(X184,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X181),s(bool,X182),s(bool,X183))))),s(X184,X185))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X181))),s(bool,X182))),s(bool,X183)))) )),
  inference(superposition,[],[f184,f112])).
fof(f6403,plain,(
  ( ! [X177,X175,X173,X176,X174] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X173),s(bool,X174),s(bool,X175))),s(X176,X177))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X173))),s(bool,X174))),s(bool,X175)))) )),
  inference(superposition,[],[f146,f112])).
fof(f6369,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,'const.combin.K_2'(s(bool,X44),s(X45,X46))) = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X41),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X41))),s(bool,X42))),s(bool,X43)))) )),
  inference(superposition,[],[f112,f250])).
fof(f6368,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X35),s(bool,X36),s(bool,X37))) | s(bool,X38) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X35))),s(bool,X36))),s(bool,X37))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X38),s(X39,X40)))) )),
  inference(superposition,[],[f112,f227])).
fof(f6367,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,X32) = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X29),s(bool,X30),s(bool,X31))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X29))),s(bool,X30))),s(bool,X31))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X32),s(X33,X34)))) )),
  inference(superposition,[],[f112,f227])).
fof(f6363,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X11),s(bool,X12),s(bool,X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X11))),s(bool,X12))),s(bool,X13)))) )),
  inference(superposition,[],[f112,f144])).
fof(f6362,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X7),s(bool,X8),s(bool,X9))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X7))),s(bool,X8))),s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f112,f144])).
fof(f6361,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(bool),X4),s(bool,X5),s(bool,X6))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(bool),X4))),s(bool,X5))),s(bool,X6)))) )),
  inference(superposition,[],[f112,f90])).
fof(f6489,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X5),X7),s(X5,X8),s(X5,X9))) = s(X5,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X5),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X5,X5),X6),s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8),s(X5,X9)))) )),
  inference(forward_demodulation,[],[f6264,f112])).
fof(f6264,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X5),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X5,X5),X6),s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8),s(X5,X9))) = s(X5,happ(s(fun(X5,X5),happ(s(fun(X5,fun(X5,X5)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X5),X7))),s(X5,X8))),s(X5,X9)))) )),
  inference(superposition,[],[f112,f102])).
fof(f6488,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X2),s(X0,X3),s(X0,X4))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4)))) )),
  inference(forward_demodulation,[],[f6263,f112])).
fof(f6263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X1),s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3),s(X0,X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))),s(X0,X3))),s(X0,X4)))) )),
  inference(superposition,[],[f112,f101])).
fof(f6239,plain,(
  ( ! [X54,X50,X48,X53,X49] : (s(bool,'F_0') != s(bool,f(s(X48,sK9(X49,X50,bool,X48)))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,g(s(X48,sK9(X49,X50,bool,X48)))),s(X53,X54)))) )),
  inference(forward_demodulation,[],[f6234,f114])).
fof(f6234,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X51,X52))) != s(bool,f(s(X48,sK9(X49,X50,bool,X48)))) | s(fun(X48,bool),X49) = s(fun(X48,bool),X50) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,g(s(X48,sK9(X49,X50,bool,X48)))),s(X53,X54)))) )),
  inference(superposition,[],[f116,f2637])).
fof(f6232,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'const.combin.K_2'(s(bool,X40),s(X41,X42))) != s(bool,f(s(X37,sK9(X38,X39,bool,X37)))) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,g(s(X37,sK9(X38,X39,bool,X37))))) )),
  inference(superposition,[],[f116,f250])).
fof(f6231,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') != s(bool,f(s(X31,sK9(X32,X33,bool,X31)))) | s(fun(X31,bool),X32) = s(fun(X31,bool),X33) | s(bool,X34) = s(bool,g(s(X31,sK9(X32,X33,bool,X31)))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X34),s(X35,X36)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6230,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X28) != s(bool,f(s(X25,sK9(X26,X27,bool,X25)))) | s(fun(X25,bool),X26) = s(fun(X25,bool),X27) | s(bool,'T_0') = s(bool,g(s(X25,sK9(X26,X27,bool,X25)))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X28),s(X29,X30)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6229,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,X22) != s(bool,f(s(X19,sK9(X20,X21,bool,X19)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,g(s(X19,sK9(X20,X21,bool,X19)))),s(X23,X24)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6226,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK9(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK9(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f116,f144])).
fof(f6225,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK9(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK9(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f116,f144])).
fof(f6224,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK9(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK9(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f116,f90])).
fof(f6237,plain,(
  ( ! [X54,X50,X48,X53,X49] : (s(X50,f(s(bool,'F_0'))) != s(X50,g(s(bool,'F_0'))) | s(fun(bool,X50),X48) = s(fun(bool,X50),X49) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK9(X48,X49,X50,bool)),s(X53,X54)))) )),
  inference(forward_demodulation,[],[f6223,f114])).
fof(f6223,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(X50,f(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X51,X52))))) != s(X50,g(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X51,X52))))) | s(fun(bool,X50),X48) = s(fun(bool,X50),X49) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK9(X48,X49,X50,bool)),s(X53,X54)))) )),
  inference(superposition,[],[f116,f2637])).
fof(f6220,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(X33,f(s(bool,'T_0'))) != s(X33,g(s(bool,'T_0'))) | s(fun(bool,X33),X31) = s(fun(bool,X33),X32) | s(bool,X34) = s(bool,sK9(X31,X32,X33,bool)) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X34),s(X35,X36)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6219,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(X27,f(s(bool,X28))) != s(X27,g(s(bool,X28))) | s(fun(bool,X27),X25) = s(fun(bool,X27),X26) | s(bool,'T_0') = s(bool,sK9(X25,X26,X27,bool)) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X28),s(X29,X30)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6218,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(X21,f(s(bool,X22))) != s(X21,g(s(bool,X22))) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK9(X19,X20,X21,bool)),s(X23,X24)))) )),
  inference(superposition,[],[f116,f227])).
fof(f6215,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK9(X7,X8,X9,bool))) )),
  inference(superposition,[],[f116,f144])).
fof(f6214,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK9(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f116,f144])).
fof(f6213,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK9(X0,X1,X2,bool))) )),
  inference(superposition,[],[f116,f90])).
fof(f6021,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'F_0') = s(bool,happ(s(fun(X189,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X188))))),s(X189,X190))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X187),s('type.semi_ring.semi_ring'(bool),X188)))))) )),
  inference(superposition,[],[f5864,f92])).
fof(f6019,plain,(
  ( ! [X182,X180,X181,X179] : (s(bool,'F_0') = s(bool,happ(s(fun(X181,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X180))))),s(X181,X182))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X179),s('type.semi_ring.semi_ring'(bool),X180)))))) )),
  inference(superposition,[],[f5864,f97])).
fof(f6018,plain,(
  ( ! [X177,X175,X178,X176] : (s(bool,'F_0') = s(bool,happ(s(fun(X177,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X176))))),s(X177,X178))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X175),s('type.semi_ring.semi_ring'(bool),X176)))))) )),
  inference(superposition,[],[f5864,f98])).
fof(f6016,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'F_0') = s(bool,happ(s(fun(X169,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X168))))),s(X169,X170))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X167),s('type.semi_ring.semi_ring'(bool),X168)))))) )),
  inference(superposition,[],[f5864,f91])).
fof(f6015,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'F_0') = s(bool,happ(s(fun(X165,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X164))))),s(X165,X166))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X163),s('type.semi_ring.semi_ring'(bool),X164)))))) )),
  inference(superposition,[],[f5864,f96])).
fof(f6014,plain,(
  ( ! [X161,X159,X162,X160] : (s(bool,'F_0') = s(bool,happ(s(fun(X161,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X160))))),s(X161,X162))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X159),s('type.semi_ring.semi_ring'(bool),X160)))))) )),
  inference(superposition,[],[f5864,f95])).
fof(f5519,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,X135) = s(bool,'const.combin.K_2'(s(bool,X136),s(bool,X135))) | s(bool,'F_0') = s(bool,happ(s(fun(X137,bool),'const.combin.K_1'(s(bool,X136))),s(X137,X138))) | s(bool,X135) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f195,f144])).
fof(f5518,plain,(
  ( ! [X134,X132,X130,X128,X133,X131,X129] : (s(bool,X128) = s(bool,'const.combin.K_2'(s(bool,X129),s(bool,X128))) | s(bool,'F_0') = s(bool,happ(s(fun(X130,bool),'const.combin.K_1'(s(bool,X129))),s(X130,X131))) | s(bool,X128) = s(bool,X132) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X132),s(X133,X134)))) )),
  inference(superposition,[],[f195,f227])).
fof(f5864,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,X0))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f5863])).
fof(f5863,plain,(
  ( ! [X12,X15,X11,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X15,bool),'const.combin.K_1'(s(bool,X11))),s(X15,X16))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f5460,f2374])).
fof(f5460,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X13,X14))) | s(bool,'F_0') = s(bool,happ(s(fun(X15,bool),'const.combin.K_1'(s(bool,X11))),s(X15,X16))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f195,f144])).
fof(f4993,plain,(
  ( ! [X173,X171,X169,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,X173) = s(bool,happ(s(fun(X169,bool),'const.combin.K_1'(s(bool,X170))),s(X169,X171))) | s(bool,'T_0') = s(bool,X173) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f4399,f127])).
fof(f4992,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,X167) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X166))) | s(bool,X168) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X165),s('type.semi_ring.semi_ring'(bool),X166))))) | s(bool,'T_0') = s(bool,X168) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f4399,f92])).
fof(f4991,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,X163) != s(bool,happ(s(fun(bool,bool),X161),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X162))))) | s(bool,X164) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X161),s('type.semi_ring.semi_ring'(bool),X162))))) | s(bool,'T_0') = s(bool,X164) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f4399,f93])).
fof(f4990,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,X159) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X158))) | s(bool,X160) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X157),s('type.semi_ring.semi_ring'(bool),X158))))) | s(bool,'T_0') = s(bool,X160) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f4399,f97])).
fof(f4989,plain,(
  ( ! [X156,X154,X155,X153] : (s(bool,X155) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X154))) | s(bool,X156) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X153),s('type.semi_ring.semi_ring'(bool),X154))))) | s(bool,'T_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f4399,f98])).
fof(f4988,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,X151) != s(bool,happ(s(fun(bool,bool),X149),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X150))))) | s(bool,X152) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X149),s('type.semi_ring.semi_ring'(bool),X150))))) | s(bool,'T_0') = s(bool,X152) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f4399,f94])).
fof(f4987,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X147) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X146))) | s(bool,X148) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X145),s('type.semi_ring.semi_ring'(bool),X146))))) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f4399,f91])).
fof(f4986,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X143) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X142))) | s(bool,X144) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X141),s('type.semi_ring.semi_ring'(bool),X142))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f4399,f96])).
fof(f4985,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,X139) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X138))) | s(bool,X140) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X137),s('type.semi_ring.semi_ring'(bool),X138))))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f4399,f95])).
fof(f5000,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4)))) )),
  inference(condensation,[],[f4999])).
fof(f4999,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X58),s(X59,X60)))) )),
  inference(duplicate_literal_removal,[],[f4970])).
fof(f4970,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X58),s(X59,X60)))) )),
  inference(superposition,[],[f4399,f227])).
fof(f4958,plain,(
  ( ! [X167,X165,X166,X164,X168] : (s(bool,X165) != s(bool,X167) | s(bool,X167) = s(bool,X168) | s(bool,'T_0') = s(bool,X168) | s(bool,'T_0') = s(bool,happ(s(fun(X164,bool),'const.combin.K_1'(s(bool,X165))),s(X164,X166)))) )),
  inference(superposition,[],[f4399,f127])).
fof(f4957,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,X162) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X161))) | s(bool,X162) = s(bool,X163) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X160),s('type.semi_ring.semi_ring'(bool),X161)))))) )),
  inference(superposition,[],[f4399,f92])).
fof(f4956,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X158) != s(bool,happ(s(fun(bool,bool),X156),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X157))))) | s(bool,X158) = s(bool,X159) | s(bool,'T_0') = s(bool,X159) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X156),s('type.semi_ring.semi_ring'(bool),X157)))))) )),
  inference(superposition,[],[f4399,f93])).
fof(f4955,plain,(
  ( ! [X154,X152,X155,X153] : (s(bool,X154) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X153))) | s(bool,X154) = s(bool,X155) | s(bool,'T_0') = s(bool,X155) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X152),s('type.semi_ring.semi_ring'(bool),X153)))))) )),
  inference(superposition,[],[f4399,f97])).
fof(f4954,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,X150) != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X149))) | s(bool,X150) = s(bool,X151) | s(bool,'T_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X148),s('type.semi_ring.semi_ring'(bool),X149)))))) )),
  inference(superposition,[],[f4399,f98])).
fof(f4953,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X146) != s(bool,happ(s(fun(bool,bool),X144),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X145))))) | s(bool,X146) = s(bool,X147) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X144),s('type.semi_ring.semi_ring'(bool),X145)))))) )),
  inference(superposition,[],[f4399,f94])).
fof(f4952,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X142) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X141))) | s(bool,X142) = s(bool,X143) | s(bool,'T_0') = s(bool,X143) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X140),s('type.semi_ring.semi_ring'(bool),X141)))))) )),
  inference(superposition,[],[f4399,f91])).
fof(f4951,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,X138) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X137))) | s(bool,X138) = s(bool,X139) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X136),s('type.semi_ring.semi_ring'(bool),X137)))))) )),
  inference(superposition,[],[f4399,f96])).
fof(f4950,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,X134) != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X133))) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X132),s('type.semi_ring.semi_ring'(bool),X133)))))) )),
  inference(superposition,[],[f4399,f95])).
fof(f4616,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.combin.K_2'(s(bool,X1),s(X3,X4))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4430,f90])).
fof(f4439,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'F_0') != s(bool,X128) | s(bool,X130) = s(bool,happ(s(fun(X127,bool),'const.combin.K_1'(s(bool,X128))),s(X127,X129))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f4328,f127])).
fof(f4438,plain,(
  ( ! [X125,X126,X124] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X125))) | s(bool,X126) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X124),s('type.semi_ring.semi_ring'(bool),X125))))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f4328,f92])).
fof(f4437,plain,(
  ( ! [X123,X121,X122] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X121),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X122))))) | s(bool,X123) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X121),s('type.semi_ring.semi_ring'(bool),X122))))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f4328,f93])).
fof(f4436,plain,(
  ( ! [X118,X120,X119] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X119))) | s(bool,X120) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X118),s('type.semi_ring.semi_ring'(bool),X119))))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f4328,f97])).
fof(f4435,plain,(
  ( ! [X116,X117,X115] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X116))) | s(bool,X117) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X115),s('type.semi_ring.semi_ring'(bool),X116))))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f4328,f98])).
fof(f4434,plain,(
  ( ! [X114,X112,X113] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X112),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X113))))) | s(bool,X114) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X112),s('type.semi_ring.semi_ring'(bool),X113))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f4328,f94])).
fof(f4433,plain,(
  ( ! [X111,X109,X110] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X110))) | s(bool,X111) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X109),s('type.semi_ring.semi_ring'(bool),X110))))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f4328,f91])).
fof(f4432,plain,(
  ( ! [X107,X108,X106] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X107))) | s(bool,X108) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X106),s('type.semi_ring.semi_ring'(bool),X107))))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f4328,f96])).
fof(f4431,plain,(
  ( ! [X103,X105,X104] : (s(bool,'F_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X104))) | s(bool,X105) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X103),s('type.semi_ring.semi_ring'(bool),X104))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f4328,f95])).
fof(f4430,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') != s(bool,X99) | s(bool,X102) = s(bool,'const.combin.K_2'(s(bool,X99),s(X100,X101))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f4328,f114])).
fof(f4399,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4328,f90])).
fof(f4328,plain,(
  ( ! [X112,X115] : (s(bool,'F_0') != s(bool,X112) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f4151,f114])).
fof(f4267,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X204),s(X206,X207))) | s(bool,X208) = s(bool,happ(s(fun(X203,bool),'const.combin.K_1'(s(bool,X204))),s(X203,X205))) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f4151,f127])).
fof(f4266,plain,(
  ( ! [X198,X202,X200,X199,X201] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X199))),s(X200,X201))) | s(bool,X202) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X198),s('type.semi_ring.semi_ring'(bool),X199))))) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f4151,f92])).
fof(f4264,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X189))),s(X190,X191))) | s(bool,X192) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X188),s('type.semi_ring.semi_ring'(bool),X189))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f4151,f97])).
fof(f4263,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X184))),s(X185,X186))) | s(bool,X187) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X183),s('type.semi_ring.semi_ring'(bool),X184))))) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f4151,f98])).
fof(f4261,plain,(
  ( ! [X177,X175,X173,X176,X174] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X174))),s(X175,X176))) | s(bool,X177) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X173),s('type.semi_ring.semi_ring'(bool),X174))))) | s(bool,'T_0') = s(bool,X177)) )),
  inference(superposition,[],[f4151,f91])).
fof(f4260,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X169))),s(X170,X171))) | s(bool,X172) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X168),s('type.semi_ring.semi_ring'(bool),X169))))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f4151,f96])).
fof(f4259,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X164))),s(X165,X166))) | s(bool,X167) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X163),s('type.semi_ring.semi_ring'(bool),X164))))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f4151,f95])).
fof(f4258,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X157),s(X160,X161))) | s(bool,X162) = s(bool,'const.combin.K_2'(s(bool,X157),s(X158,X159))) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f4151,f114])).
fof(f4354,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X3),s(X4,X5))) | s(bool,X0) = s(bool,X3)) )),
  inference(condensation,[],[f4353])).
fof(f4353,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X66),s(X67,X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X70),s(X71,X72)))) )),
  inference(duplicate_literal_removal,[],[f4244])).
fof(f4244,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X66),s(X67,X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X70),s(X71,X72)))) )),
  inference(superposition,[],[f4151,f227])).
fof(f4060,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (s(bool,X206) != s(bool,'const.combin.K_2'(s(bool,X204),s(X207,X208))) | s(bool,X206) = s(bool,happ(s(fun(X203,bool),'const.combin.K_1'(s(bool,X204))),s(X203,X205))) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f1331,f127])).
fof(f4059,plain,(
  ( ! [X198,X202,X200,X199,X201] : (s(bool,X200) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X199))),s(X201,X202))) | s(bool,X200) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X198),s('type.semi_ring.semi_ring'(bool),X199))))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f1331,f92])).
fof(f4057,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,X190) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X189))),s(X191,X192))) | s(bool,X190) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X188),s('type.semi_ring.semi_ring'(bool),X189))))) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f1331,f97])).
fof(f4056,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,X185) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X184))),s(X186,X187))) | s(bool,X185) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X183),s('type.semi_ring.semi_ring'(bool),X184))))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f1331,f98])).
fof(f4054,plain,(
  ( ! [X177,X175,X173,X176,X174] : (s(bool,X175) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X174))),s(X176,X177))) | s(bool,X175) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X173),s('type.semi_ring.semi_ring'(bool),X174))))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f1331,f91])).
fof(f4053,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,X170) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X169))),s(X171,X172))) | s(bool,X170) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X168),s('type.semi_ring.semi_ring'(bool),X169))))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f1331,f96])).
fof(f4052,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,X165) != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X164))),s(X166,X167))) | s(bool,X165) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X163),s('type.semi_ring.semi_ring'(bool),X164))))) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f1331,f95])).
fof(f4051,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (s(bool,X160) != s(bool,'const.combin.K_2'(s(bool,X157),s(X161,X162))) | s(bool,X160) = s(bool,'const.combin.K_2'(s(bool,X157),s(X158,X159))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f1331,f114])).
fof(f4025,plain,(
  ( ! [X198,X202,X200,X199,X203,X201] : (s(bool,X199) != s(bool,'const.combin.K_2'(s(bool,X201),s(X202,X203))) | s(bool,X201) = s(bool,happ(s(fun(X198,bool),'const.combin.K_1'(s(bool,X199))),s(X198,X200))) | s(bool,'T_0') = s(bool,happ(s(fun(X198,bool),'const.combin.K_1'(s(bool,X199))),s(X198,X200)))) )),
  inference(superposition,[],[f1331,f127])).
fof(f4024,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X194))) != s(bool,'const.combin.K_2'(s(bool,X195),s(X196,X197))) | s(bool,X195) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X193),s('type.semi_ring.semi_ring'(bool),X194))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X193),s('type.semi_ring.semi_ring'(bool),X194)))))) )),
  inference(superposition,[],[f1331,f92])).
fof(f4023,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,'const.combin.K_2'(s(bool,X190),s(X191,X192))) != s(bool,happ(s(fun(bool,bool),X188),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X189))))) | s(bool,X190) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X188),s('type.semi_ring.semi_ring'(bool),X189))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X188),s('type.semi_ring.semi_ring'(bool),X189)))))) )),
  inference(superposition,[],[f1331,f93])).
fof(f4022,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X184))) != s(bool,'const.combin.K_2'(s(bool,X185),s(X186,X187))) | s(bool,X185) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X183),s('type.semi_ring.semi_ring'(bool),X184))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X183),s('type.semi_ring.semi_ring'(bool),X184)))))) )),
  inference(superposition,[],[f1331,f97])).
fof(f4021,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X179))) != s(bool,'const.combin.K_2'(s(bool,X180),s(X181,X182))) | s(bool,X180) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X178),s('type.semi_ring.semi_ring'(bool),X179))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X178),s('type.semi_ring.semi_ring'(bool),X179)))))) )),
  inference(superposition,[],[f1331,f98])).
fof(f4020,plain,(
  ( ! [X177,X175,X173,X176,X174] : (s(bool,'const.combin.K_2'(s(bool,X175),s(X176,X177))) != s(bool,happ(s(fun(bool,bool),X173),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X174))))) | s(bool,X175) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X173),s('type.semi_ring.semi_ring'(bool),X174))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X173),s('type.semi_ring.semi_ring'(bool),X174)))))) )),
  inference(superposition,[],[f1331,f94])).
fof(f4019,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X169))) != s(bool,'const.combin.K_2'(s(bool,X170),s(X171,X172))) | s(bool,X170) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X168),s('type.semi_ring.semi_ring'(bool),X169))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X168),s('type.semi_ring.semi_ring'(bool),X169)))))) )),
  inference(superposition,[],[f1331,f91])).
fof(f4018,plain,(
  ( ! [X167,X165,X163,X166,X164] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X164))) != s(bool,'const.combin.K_2'(s(bool,X165),s(X166,X167))) | s(bool,X165) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X163),s('type.semi_ring.semi_ring'(bool),X164))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X163),s('type.semi_ring.semi_ring'(bool),X164)))))) )),
  inference(superposition,[],[f1331,f96])).
fof(f4017,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X159))) != s(bool,'const.combin.K_2'(s(bool,X160),s(X161,X162))) | s(bool,X160) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X158),s('type.semi_ring.semi_ring'(bool),X159))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X158),s('type.semi_ring.semi_ring'(bool),X159)))))) )),
  inference(superposition,[],[f1331,f95])).
fof(f4016,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (s(bool,X152) != s(bool,'const.combin.K_2'(s(bool,X155),s(X156,X157))) | s(bool,X155) = s(bool,'const.combin.K_2'(s(bool,X152),s(X153,X154))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X152),s(X153,X154)))) )),
  inference(superposition,[],[f1331,f114])).
fof(f4147,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'const.combin.K_2'(s(bool,X46),s(X47,X48))) != s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,X45) = s(bool,X49) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46)) )),
  inference(duplicate_literal_removal,[],[f4003])).
fof(f4003,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'const.combin.K_2'(s(bool,X46),s(X47,X48))) != s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,X45) = s(bool,X49) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f1331,f250])).
fof(f4151,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f3995])).
fof(f3995,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1331,f90])).
fof(f3276,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,X135) = s(bool,'const.combin.K_2'(s(bool,X136),s(bool,X135))) | s(bool,X136) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X137,X138))) | s(bool,X135) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f2637,f144])).
fof(f3275,plain,(
  ( ! [X134,X132,X130,X128,X133,X131,X129] : (s(bool,X128) = s(bool,'const.combin.K_2'(s(bool,X129),s(bool,X128))) | s(bool,X129) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X130,X131))) | s(bool,X128) = s(bool,X132) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X132),s(X133,X134)))) )),
  inference(superposition,[],[f2637,f227])).
fof(f3105,plain,(
  ( ! [X167,X169,X166,X168] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X167))),s(X168,X169))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X166),s('type.semi_ring.semi_ring'(bool),X167)))))) )),
  inference(superposition,[],[f2374,f98])).
fof(f3131,plain,(
  ( ! [X149,X148] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X149))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X148),s('type.semi_ring.semi_ring'(bool),X149)))))) )),
  inference(forward_demodulation,[],[f3101,f114])).
fof(f3101,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X149))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X150,X151))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X148),s('type.semi_ring.semi_ring'(bool),X149)))))) )),
  inference(superposition,[],[f1343,f98])).
fof(f3100,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X145))),s(X146,X147))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X144),s('type.semi_ring.semi_ring'(bool),X145)))))) )),
  inference(superposition,[],[f747,f98])).
fof(f3097,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s(bool,X133) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X132))),s(X134,X135))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X131),s('type.semi_ring.semi_ring'(bool),X132))))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f250,f98])).
fof(f3093,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X112))),s(X113,X114))) | s(bool,X115) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X111),s('type.semi_ring.semi_ring'(bool),X112))))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f227,f98])).
fof(f3090,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') = s(bool,happ(s(fun(X97,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X96))))),s(X97,X98))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X95),s('type.semi_ring.semi_ring'(bool),X96)))))) )),
  inference(superposition,[],[f184,f98])).
fof(f3088,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X90))),s(X91,X92))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X89),s('type.semi_ring.semi_ring'(bool),X90)))))) )),
  inference(superposition,[],[f146,f98])).
fof(f3065,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X33))) = s(bool,'const.combin.K_2'(s(bool,X34),s(X35,X36))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X32),s('type.semi_ring.semi_ring'(bool),X33)))))) )),
  inference(superposition,[],[f98,f250])).
fof(f3064,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X28))) | s(bool,X29) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X27),s('type.semi_ring.semi_ring'(bool),X28))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31)))) )),
  inference(superposition,[],[f98,f227])).
fof(f3063,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X24) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X23))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X22),s('type.semi_ring.semi_ring'(bool),X23))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f98,f227])).
fof(f3059,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X8),s('type.semi_ring.semi_ring'(bool),X9)))))) )),
  inference(superposition,[],[f98,f144])).
fof(f3058,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X6))) | s(bool,X7) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X5),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f98,f144])).
fof(f3057,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.semi_ring.semi_ring'(bool),X4)))))) )),
  inference(superposition,[],[f98,f90])).
fof(f2955,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) != s(bool,'const.combin.K_2'(s(bool,X82),s(bool,X81))) | s(bool,'T_0') = s(bool,X82) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f2374,f144])).
fof(f2954,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,X76) != s(bool,'const.combin.K_2'(s(bool,X77),s(bool,X76))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X78),s(X79,X80)))) )),
  inference(superposition,[],[f2374,f227])).
fof(f2934,plain,(
  ( ! [X142,X140,X143,X141,X139] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X140),s(X142,X143))) | s(bool,'T_0') = s(bool,happ(s(fun(X139,bool),'const.combin.K_1'(s(bool,X140))),s(X139,X141)))) )),
  inference(superposition,[],[f2374,f127])).
fof(f2933,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X136))),s(X137,X138))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X135),s('type.semi_ring.semi_ring'(bool),X136)))))) )),
  inference(superposition,[],[f2374,f97])).
fof(f2931,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X128))),s(X129,X130))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X127),s('type.semi_ring.semi_ring'(bool),X128)))))) )),
  inference(superposition,[],[f2374,f92])).
fof(f2929,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X120))),s(X121,X122))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X119),s('type.semi_ring.semi_ring'(bool),X120)))))) )),
  inference(superposition,[],[f2374,f91])).
fof(f2928,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X116))),s(X117,X118))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X115),s('type.semi_ring.semi_ring'(bool),X116)))))) )),
  inference(superposition,[],[f2374,f96])).
fof(f2927,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X112))),s(X113,X114))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X111),s('type.semi_ring.semi_ring'(bool),X112)))))) )),
  inference(superposition,[],[f2374,f95])).
fof(f2771,plain,(
  ( ! [X142,X141] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X142))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X141),s('type.semi_ring.semi_ring'(bool),X142)))))) )),
  inference(forward_demodulation,[],[f2744,f114])).
fof(f2744,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X142))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X143,X144))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X141),s('type.semi_ring.semi_ring'(bool),X142)))))) )),
  inference(superposition,[],[f1343,f97])).
fof(f2743,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X138))),s(X139,X140))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X137),s('type.semi_ring.semi_ring'(bool),X138)))))) )),
  inference(superposition,[],[f747,f97])).
fof(f2740,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X126) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X125))),s(X127,X128))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X124),s('type.semi_ring.semi_ring'(bool),X125))))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f250,f97])).
fof(f2736,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X105))),s(X106,X107))) | s(bool,X108) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X104),s('type.semi_ring.semi_ring'(bool),X105))))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f227,f97])).
fof(f2733,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,happ(s(fun(X90,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X89))))),s(X90,X91))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X88),s('type.semi_ring.semi_ring'(bool),X89)))))) )),
  inference(superposition,[],[f184,f97])).
fof(f2731,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X83))),s(X84,X85))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X82),s('type.semi_ring.semi_ring'(bool),X83)))))) )),
  inference(superposition,[],[f146,f97])).
fof(f2710,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X33))) = s(bool,'const.combin.K_2'(s(bool,X34),s(X35,X36))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X32),s('type.semi_ring.semi_ring'(bool),X33)))))) )),
  inference(superposition,[],[f97,f250])).
fof(f2709,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X28))) | s(bool,X29) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X27),s('type.semi_ring.semi_ring'(bool),X28))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31)))) )),
  inference(superposition,[],[f97,f227])).
fof(f2708,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X24) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X23))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X22),s('type.semi_ring.semi_ring'(bool),X23))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f97,f227])).
fof(f2704,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X8),s('type.semi_ring.semi_ring'(bool),X9)))))) )),
  inference(superposition,[],[f97,f144])).
fof(f2703,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X6))) | s(bool,X7) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X5),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f97,f144])).
fof(f2702,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.semi_ring.semi_ring'(bool),X4)))))) )),
  inference(superposition,[],[f97,f90])).
fof(f2574,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2))) | s(bool,X0) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2)))) )),
  inference(equality_factoring,[],[f2344])).
fof(f2421,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,X136) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X138,X139))) | s(bool,'T_0') = s(bool,happ(s(fun(X135,bool),'const.combin.K_1'(s(bool,X136))),s(X135,X137)))) )),
  inference(superposition,[],[f2344,f127])).
fof(f2637,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,X0) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X3,X4)))) )),
  inference(condensation,[],[f2636])).
fof(f2636,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(duplicate_literal_removal,[],[f2401])).
fof(f2401,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(superposition,[],[f2344,f227])).
fof(f2374,plain,(
  ( ! [X4,X0,X3] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X0),s(X3,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2292,f114])).
fof(f2292,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) != s(bool,'const.combin.K_2'(s(bool,X0),s(X3,X4)))) )),
  inference(condensation,[],[f2291])).
fof(f2291,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X99,X100))) != s(bool,'const.combin.K_2'(s(bool,X96),s(X97,X98))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X96)) )),
  inference(duplicate_literal_removal,[],[f2180])).
fof(f2180,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X99,X100))) != s(bool,'const.combin.K_2'(s(bool,X96),s(X97,X98))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f1343,f250])).
fof(f2116,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (s(bool,X155) = s(bool,'const.combin.K_2'(s(bool,X153),s(X156,X157))) | s(bool,'T_0') = s(bool,happ(s(fun(X152,bool),'const.combin.K_1'(s(bool,X153))),s(X152,X154))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f250,f127])).
fof(f2114,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s(bool,X144) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X143))),s(X145,X146))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X142),s('type.semi_ring.semi_ring'(bool),X143))))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f250,f92])).
fof(f2112,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,X134) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X133))),s(X135,X136))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X132),s('type.semi_ring.semi_ring'(bool),X133))))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f250,f91])).
fof(f2111,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,X129) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X128))),s(X130,X131))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X127),s('type.semi_ring.semi_ring'(bool),X128))))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f250,f96])).
fof(f2110,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,X124) = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X123))),s(X125,X126))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X122),s('type.semi_ring.semi_ring'(bool),X123))))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f250,f95])).
fof(f2344,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f2343])).
fof(f2343,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f2089])).
fof(f2089,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f250,f90])).
fof(f2088,plain,(
  ( ! [X156,X154,X152,X157,X155,X153] : (s(bool,X153) = s(bool,'const.combin.K_2'(s(bool,X155),s(X156,X157))) | s(bool,'T_0') = s(bool,X155) | s(bool,'T_0') = s(bool,happ(s(fun(X152,bool),'const.combin.K_1'(s(bool,X153))),s(X152,X154)))) )),
  inference(superposition,[],[f250,f127])).
fof(f2087,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,'const.combin.K_2'(s(bool,X149),s(X150,X151))) = s(bool,happ(s(fun(bool,bool),X147),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X148))))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X147),s('type.semi_ring.semi_ring'(bool),X148)))))) )),
  inference(superposition,[],[f250,f93])).
fof(f2086,plain,(
  ( ! [X146,X144,X142,X145,X143] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X143))) = s(bool,'const.combin.K_2'(s(bool,X144),s(X145,X146))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X142),s('type.semi_ring.semi_ring'(bool),X143)))))) )),
  inference(superposition,[],[f250,f92])).
fof(f2085,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'const.combin.K_2'(s(bool,X139),s(X140,X141))) = s(bool,happ(s(fun(bool,bool),X137),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X138))))) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X137),s('type.semi_ring.semi_ring'(bool),X138)))))) )),
  inference(superposition,[],[f250,f94])).
fof(f2084,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X133))) = s(bool,'const.combin.K_2'(s(bool,X134),s(X135,X136))) | s(bool,'T_0') = s(bool,X134) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X132),s('type.semi_ring.semi_ring'(bool),X133)))))) )),
  inference(superposition,[],[f250,f91])).
fof(f2083,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X128))) = s(bool,'const.combin.K_2'(s(bool,X129),s(X130,X131))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X127),s('type.semi_ring.semi_ring'(bool),X128)))))) )),
  inference(superposition,[],[f250,f96])).
fof(f2082,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X123))) = s(bool,'const.combin.K_2'(s(bool,X124),s(X125,X126))) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X122),s('type.semi_ring.semi_ring'(bool),X123)))))) )),
  inference(superposition,[],[f250,f95])).
fof(f2347,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,X0),s(X4,X5))) = s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f2346])).
fof(f2346,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'const.combin.K_2'(s(bool,X46),s(X47,X48))) = s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46)) )),
  inference(duplicate_literal_removal,[],[f2073])).
fof(f2073,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'const.combin.K_2'(s(bool,X46),s(X47,X48))) = s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f250,f250])).
fof(f2056,plain,(
  ( ! [X123,X122] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X123))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X122),s('type.semi_ring.semi_ring'(bool),X123)))))) )),
  inference(forward_demodulation,[],[f2033,f114])).
fof(f2033,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X123))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X124,X125))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X122),s('type.semi_ring.semi_ring'(bool),X123)))))) )),
  inference(superposition,[],[f1343,f96])).
fof(f2032,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X119))),s(X120,X121))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X118),s('type.semi_ring.semi_ring'(bool),X119)))))) )),
  inference(superposition,[],[f747,f96])).
fof(f2027,plain,(
  ( ! [X99,X97,X95,X98,X96] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X96))),s(X97,X98))) | s(bool,X99) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X95),s('type.semi_ring.semi_ring'(bool),X96))))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f227,f96])).
fof(f2024,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(X81,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X80))))),s(X81,X82))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X79),s('type.semi_ring.semi_ring'(bool),X80)))))) )),
  inference(superposition,[],[f184,f96])).
fof(f2022,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X74))),s(X75,X76))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X73),s('type.semi_ring.semi_ring'(bool),X74)))))) )),
  inference(superposition,[],[f146,f96])).
fof(f2004,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X28))) | s(bool,X29) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X27),s('type.semi_ring.semi_ring'(bool),X28))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31)))) )),
  inference(superposition,[],[f96,f227])).
fof(f2003,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X24) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X23))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X22),s('type.semi_ring.semi_ring'(bool),X23))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f96,f227])).
fof(f1999,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X8),s('type.semi_ring.semi_ring'(bool),X9)))))) )),
  inference(superposition,[],[f96,f144])).
fof(f1998,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X6))) | s(bool,X7) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X5),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f96,f144])).
fof(f1997,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.semi_ring.semi_ring'(bool),X4)))))) )),
  inference(superposition,[],[f96,f90])).
fof(f1855,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(bool,X10) != s(bool,'const.combin.K_2'(s(bool,X9),s(X11,X12))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X10),s(X13,X14))) | s(bool,X9) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1637,f144])).
fof(f1854,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,X1) != s(bool,'const.combin.K_2'(s(bool,X0),s(X2,X3))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X1),s(X4,X5))) | s(bool,X0) = s(bool,X6) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X6),s(X7,X8)))) )),
  inference(superposition,[],[f1637,f227])).
fof(f1746,plain,(
  ( ! [X116,X117] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X117))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X116),s('type.semi_ring.semi_ring'(bool),X117)))))) )),
  inference(forward_demodulation,[],[f1727,f114])).
fof(f1727,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X117))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X118,X119))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X116),s('type.semi_ring.semi_ring'(bool),X117)))))) )),
  inference(superposition,[],[f1343,f95])).
fof(f1726,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X113))),s(X114,X115))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X112),s('type.semi_ring.semi_ring'(bool),X113)))))) )),
  inference(superposition,[],[f747,f95])).
fof(f1721,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X90))),s(X91,X92))) | s(bool,X93) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X89),s('type.semi_ring.semi_ring'(bool),X90))))) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f227,f95])).
fof(f1718,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(X75,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X74))))),s(X75,X76))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X73),s('type.semi_ring.semi_ring'(bool),X74)))))) )),
  inference(superposition,[],[f184,f95])).
fof(f1716,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X68))),s(X69,X70))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X67),s('type.semi_ring.semi_ring'(bool),X68)))))) )),
  inference(superposition,[],[f146,f95])).
fof(f1699,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X28))) | s(bool,X29) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X27),s('type.semi_ring.semi_ring'(bool),X28))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31)))) )),
  inference(superposition,[],[f95,f227])).
fof(f1698,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X24) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X23))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X22),s('type.semi_ring.semi_ring'(bool),X23))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f95,f227])).
fof(f1694,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X8),s('type.semi_ring.semi_ring'(bool),X9)))))) )),
  inference(superposition,[],[f95,f144])).
fof(f1693,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X6))) | s(bool,X7) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X5),s('type.semi_ring.semi_ring'(bool),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f95,f144])).
fof(f1692,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.semi_ring.semi_ring'(bool),X4)))))) )),
  inference(superposition,[],[f95,f90])).
fof(f1568,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,X8) != s(bool,'const.combin.K_2'(s(bool,X7),s(X9,X10))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1343,f144])).
fof(f1567,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,'const.combin.K_2'(s(bool,X0),s(X2,X3))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6)))) )),
  inference(superposition,[],[f1343,f227])).
fof(f1566,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X106) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X108,X109))) | s(bool,'T_0') = s(bool,happ(s(fun(X105,bool),'const.combin.K_1'(s(bool,X106))),s(X105,X107)))) )),
  inference(superposition,[],[f1343,f127])).
fof(f1653,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X102))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X101),s('type.semi_ring.semi_ring'(bool),X102)))))) )),
  inference(forward_demodulation,[],[f1565,f114])).
fof(f1565,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X103,X104))) != s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X102))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X101),s('type.semi_ring.semi_ring'(bool),X102)))))) )),
  inference(superposition,[],[f1343,f93])).
fof(f1652,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X98))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X97),s('type.semi_ring.semi_ring'(bool),X98)))))) )),
  inference(forward_demodulation,[],[f1564,f114])).
fof(f1564,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X98))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X99,X100))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X97),s('type.semi_ring.semi_ring'(bool),X98)))))) )),
  inference(superposition,[],[f1343,f92])).
fof(f1651,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X93),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X94))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X93),s('type.semi_ring.semi_ring'(bool),X94)))))) )),
  inference(forward_demodulation,[],[f1563,f114])).
fof(f1563,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X95,X96))) != s(bool,happ(s(fun(bool,bool),X93),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X94))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X93),s('type.semi_ring.semi_ring'(bool),X94)))))) )),
  inference(superposition,[],[f1343,f94])).
fof(f1650,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X90))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X89),s('type.semi_ring.semi_ring'(bool),X90)))))) )),
  inference(forward_demodulation,[],[f1562,f114])).
fof(f1562,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X90))) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X91,X92))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X89),s('type.semi_ring.semi_ring'(bool),X90)))))) )),
  inference(superposition,[],[f1343,f91])).
fof(f1637,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X3,X4))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) )),
  inference(condensation,[],[f1636])).
fof(f1636,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(duplicate_literal_removal,[],[f1550])).
fof(f1550,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(superposition,[],[f1343,f227])).
fof(f1331,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f227])).
fof(f1343,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1325])).
fof(f1325,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f227])).
fof(f1265,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s(bool,X95) = s(bool,'const.combin.K_2'(s(bool,X96),s(bool,X95))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X96),s(X97,X98))) | s(bool,X95) = s(bool,X99) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X99),s(X100,X101)))) )),
  inference(superposition,[],[f194,f227])).
fof(f1264,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(bool,X90) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X91))),s(bool,X90))) | s(bool,'F_0') = s(bool,X91) | s(bool,X90) = s(bool,X92) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X92),s(X93,X94)))) )),
  inference(superposition,[],[f184,f227])).
fof(f1263,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) = s(bool,'const.combin.K_2'(s(bool,X86),s(bool,X85))) | s(bool,'F_0') = s(bool,X86) | s(bool,X85) = s(bool,X87) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X87),s(X88,X89)))) )),
  inference(superposition,[],[f146,f227])).
fof(f1151,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X75) = s(bool,'const.combin.K_2'(s(bool,X76),s(bool,X75))) | s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f227,f144])).
fof(f1150,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,X69) = s(bool,'const.combin.K_2'(s(bool,X70),s(bool,X69))) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X72),s(X73,X74)))) )),
  inference(superposition,[],[f227,f227])).
fof(f1135,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X134),s(X136,X137))) | s(bool,X138) = s(bool,happ(s(fun(X133,bool),'const.combin.K_1'(s(bool,X134))),s(X133,X135))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f227,f127])).
fof(f1133,plain,(
  ( ! [X127,X125,X123,X126,X124] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X124))),s(X125,X126))) | s(bool,X127) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X123),s('type.semi_ring.semi_ring'(bool),X124))))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f227,f92])).
fof(f1131,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X114))),s(X115,X116))) | s(bool,X117) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X113),s('type.semi_ring.semi_ring'(bool),X114))))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f227,f91])).
fof(f1111,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') = s(bool,X91) | s(bool,X93) = s(bool,happ(s(fun(X90,bool),'const.combin.K_1'(s(bool,X91))),s(X90,X92))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X93),s(X94,X95)))) )),
  inference(superposition,[],[f227,f127])).
fof(f1110,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X85),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X86))))) | s(bool,X87) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X85),s('type.semi_ring.semi_ring'(bool),X86))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X87),s(X88,X89)))) )),
  inference(superposition,[],[f227,f93])).
fof(f1109,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X81))) | s(bool,X82) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X80),s('type.semi_ring.semi_ring'(bool),X81))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X82),s(X83,X84)))) )),
  inference(superposition,[],[f227,f92])).
fof(f1108,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X75),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X76))))) | s(bool,X77) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X75),s('type.semi_ring.semi_ring'(bool),X76))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X77),s(X78,X79)))) )),
  inference(superposition,[],[f227,f94])).
fof(f1107,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X71))) | s(bool,X72) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X70),s('type.semi_ring.semi_ring'(bool),X71))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X72),s(X73,X74)))) )),
  inference(superposition,[],[f227,f91])).
fof(f1089,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X10),s(X11,X12))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f227,f144])).
fof(f1088,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X5),s(X6,X7)))) )),
  inference(superposition,[],[f227,f227])).
fof(f1087,plain,(
  ( ! [X127,X125,X123,X126,X124,X122] : (s(bool,X123) = s(bool,X125) | s(bool,'T_0') = s(bool,happ(s(fun(X122,bool),'const.combin.K_1'(s(bool,X123))),s(X122,X124))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X125),s(X126,X127)))) )),
  inference(superposition,[],[f227,f127])).
fof(f1086,plain,(
  ( ! [X121,X118,X120,X119,X117] : (s(bool,X119) = s(bool,happ(s(fun(bool,bool),X117),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X118))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X117),s('type.semi_ring.semi_ring'(bool),X118))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X119),s(X120,X121)))) )),
  inference(superposition,[],[f227,f93])).
fof(f1085,plain,(
  ( ! [X116,X114,X112,X115,X113] : (s(bool,X114) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X113))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X112),s('type.semi_ring.semi_ring'(bool),X113))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X114),s(X115,X116)))) )),
  inference(superposition,[],[f227,f92])).
fof(f1084,plain,(
  ( ! [X111,X109,X107,X110,X108] : (s(bool,X109) = s(bool,happ(s(fun(bool,bool),X107),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X108))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X107),s('type.semi_ring.semi_ring'(bool),X108))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X109),s(X110,X111)))) )),
  inference(superposition,[],[f227,f94])).
fof(f1083,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,X104) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X103))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X102),s('type.semi_ring.semi_ring'(bool),X103))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X104),s(X105,X106)))) )),
  inference(superposition,[],[f227,f91])).
fof(f1497,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X3),s(X4,X5)))) )),
  inference(condensation,[],[f1496])).
fof(f1496,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X38),s(X40,X41))) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X42),s(X43,X44)))) )),
  inference(duplicate_literal_removal,[],[f1047])).
fof(f1047,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X38),s(X40,X41))) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X42),s(X43,X44)))) )),
  inference(superposition,[],[f227,f227])).
fof(f996,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X16))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X15),s('type.semi_ring.semi_ring'(bool),X16)))))) )),
  inference(superposition,[],[f94,f90])).
fof(f994,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(bool,bool),X9),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X9),s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f94,f144])).
fof(f992,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X3),s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f94,f144])).
fof(f901,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X67))),s(X68,X69))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X66),s('type.semi_ring.semi_ring'(bool),X67)))))) )),
  inference(superposition,[],[f747,f92])).
fof(f900,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X63))),s(X64,X65))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X62),s('type.semi_ring.semi_ring'(bool),X63)))))) )),
  inference(superposition,[],[f747,f91])).
fof(f777,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X66) = s(bool,'const.combin.K_2'(s(bool,X67),s(bool,X66))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X67),s(X68,X69))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f194,f90])).
fof(f747,plain,(
  ( ! [X39,X36,X40] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X36),s(X39,X40))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f194,f114])).
fof(f733,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X56) = s(bool,'const.combin.K_2'(s(bool,X57),s(bool,X56))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X57),s(X58,X59))) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f194,f144])).
fof(f674,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X16))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X15),s('type.semi_ring.semi_ring'(bool),X16)))))) )),
  inference(superposition,[],[f93,f90])).
fof(f672,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(bool,bool),X9),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X9),s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f93,f144])).
fof(f670,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X3),s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f93,f144])).
fof(f614,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X33))),s(bool,X32))) | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f184,f144])).
fof(f604,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') = s(bool,happ(s(fun(X56,bool),'const.combin.K_1'(s(bool,X54))),s(X56,X57))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),'const.combin.K_1'(s(bool,X54))),s(X53,X55)))) )),
  inference(superposition,[],[f184,f127])).
fof(f603,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X51,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X50))))),s(X51,X52))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X49),s('type.semi_ring.semi_ring'(bool),X50)))))) )),
  inference(superposition,[],[f184,f92])).
fof(f602,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),'const.combin.K_1'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X46))))),s(X47,X48))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X45),s('type.semi_ring.semi_ring'(bool),X46)))))) )),
  inference(superposition,[],[f184,f91])).
fof(f601,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(X43,bool),'const.combin.K_1'(s(bool,X40))),s(X43,X44))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X40),s(X41,X42)))) )),
  inference(superposition,[],[f184,f114])).
fof(f533,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X39))),s(X40,X41))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X38),s('type.semi_ring.semi_ring'(bool),X39)))))) )),
  inference(superposition,[],[f146,f92])).
fof(f522,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X15),s('type.semi_ring.semi_ring'(bool),X16))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X16)))) )),
  inference(superposition,[],[f92,f90])).
fof(f520,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X10))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X9),s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f92,f144])).
fof(f518,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(bool,bool),X3),s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f92,f144])).
fof(f479,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X39))),s(X40,X41))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X38),s('type.semi_ring.semi_ring'(bool),X39)))))) )),
  inference(superposition,[],[f146,f91])).
fof(f468,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X15),s('type.semi_ring.semi_ring'(bool),X16))))) | s(bool,'F_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X16)))) )),
  inference(superposition,[],[f91,f90])).
fof(f466,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X10))) | s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X9),s('type.semi_ring.semi_ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f91,f144])).
fof(f464,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),X4))) | s(bool,X5) = s(bool,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(bool),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(bool,bool),X3),s('type.semi_ring.semi_ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f91,f144])).
fof(f297,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.combin.K_2'(s(bool,X10),s(bool,X9))) | s(bool,'F_0') = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f146,f144])).
fof(f294,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f88,f144])).
fof(f251,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X24) = s(bool,happ(s(fun(X21,bool),'const.combin.K_1'(s(bool,X22))),s(X21,X23))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f144,f127])).
fof(f250,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X20) = s(bool,'const.combin.K_2'(s(bool,X17),s(X18,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f144,f114])).
fof(f240,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f144,f144])).
fof(f228,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),'const.combin.K_1'(s(bool,X29))),s(X28,X30)))) )),
  inference(superposition,[],[f144,f127])).
fof(f227,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f144,f114])).
fof(f195,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X17),s(X19,X20))) | s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),'const.combin.K_1'(s(bool,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f146,f127])).
fof(f194,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X11),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X11),s(X12,X13)))) )),
  inference(superposition,[],[f146,f114])).
fof(f184,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X5))),s(X4,X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f127,f90])).
fof(f147,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f89,f90])).
fof(f146,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f90,f114])).
fof(f144,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f90,f90])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.semi_ring.semi_ring_literal_|49||49|',aHL_EXT)).
fof(f113,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X22,X23,X24] : s(X6,'const.semi_ring.semi_ring_SRP_3'(s('type.semi_ring.semi_ring'(X6),X22),s(X6,X23),s(X6,X24))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X22))),s(X6,X23))),s(X6,X24)))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','ARITY_1_3_const.semi_ring.semi_ring_SRP')).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X22,X23,X24] : s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X22),s(X6,X23),s(X6,X24))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X22))),s(X6,X23))),s(X6,X24)))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','ARITY_1_3_const.semi_ring.semi_ring_SRM')).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),X1) != s('type.semi_ring.semi_ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f111,plain,(
  ( ! [X2,X0,X1] : (s('type.semi_ring.semi_ring'(X0),X1) = s('type.semi_ring.semi_ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X2))) | s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) != s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f106,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_3|48|_prove_ltb')).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_3|49|_prove_ltb')).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_24_prove_ltb')).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_27_prove_ltb')).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_25_prove_ltb')).
fof(f101,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_26_prove_ltb')).
fof(f100,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_23_prove_ltb')).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X18,X19] : s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRP_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_22_prove_ltb')).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_2|49|_prove_ltb')).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|49|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_2|48|_prove_ltb')).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_|49|7_prove_ltb')).
fof(f95,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_|49|8_prove_ltb')).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,happ(s(fun(X6,X6),X19),s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X18))))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_28_prove_ltb')).
fof(f91,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),X2),s('type.semi_ring.semi_ring'(X0),X1)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),X18))) = s(X6,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),X19),s('type.semi_ring.semi_ring'(X6),X18)))))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_|49|6_prove_ltb')).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.semi_ring.semi_ring_literal_|49||49|',aHL_BOOL_CASES)).
fof(f89,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.semi_ring.semi_ring_literal_|49||49|',aHL_FALSITY)).
fof(f88,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.semi_ring.semi_ring_literal_|49||49|',aHL_TRUTH)).
fof(f13080,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f87,f86,f85,f84,f83,f88,f89,f90,f91,f92,f93,f94,f95,f96,f97,f98,f99,f100,f101,f102,f103,f104,f105,f106,f111,f110,f109,f108,f107,f112,f113,f114,f127,f116,f144,f146,f147,f155,f184,f194,f195,f227,f228,f240,f250,f251,f294,f297,f464,f466,f468,f479,f518,f520,f522,f533,f601,f602,f603,f604,f614,f670,f672,f674,f708,f709,f710,f733,f747,f777,f900,f901,f992,f994,f996,f1497,f1083,f1084,f1085,f1086,f1087,f1088,f1089,f1107,f1108,f1109,f1110,f1111,f1131,f1133,f1135,f1150,f1151,f1263,f1264,f1265,f1343,f1331,f1637,f1650,f1651,f1652,f1653,f1566,f1567,f1568,f1692,f1693,f1694,f1698,f1699,f1716,f1718,f1721,f1726,f1746,f1854,f1855,f1997,f1998,f1999,f2003,f2004,f2022,f2024,f2027,f2032,f2056,f2347,f2082,f2083,f2084,f2085,f2086,f2087,f2088,f2344,f2110,f2111,f2112,f2114,f2116,f2374,f2637,f2421,f2574,f2702,f2703,f2704,f2708,f2709,f2710,f2731,f2733,f2736,f2740,f2743,f2771,f2927,f2928,f2929,f2931,f2933,f2934,f2954,f2955,f3057,f3058,f3059,f3063,f3064,f3065,f3088,f3090,f3093,f3097,f3100,f3131,f3105,f3275,f3276,f4151,f4147,f4016,f4017,f4018,f4019,f4020,f4021,f4022,f4023,f4024,f4025,f4051,f4052,f4053,f4054,f4056,f4057,f4059,f4060,f4354,f4258,f4259,f4260,f4261,f4263,f4264,f4266,f4267,f4328,f4399,f4430,f4431,f4432,f4433,f4434,f4435,f4436,f4437,f4438,f4439,f4616,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f5000,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f5864,f5518,f5519,f6014,f6015,f6016,f6018,f6019,f6021,f6213,f6214,f6215,f6218,f6219,f6220,f6237,f6224,f6225,f6226,f6229,f6230,f6231,f6232,f6239,f6488,f6489,f6361,f6362,f6363,f6367,f6368,f6369,f6403,f6405,f6410,f6414,f6417,f6418,f6419,f6501,f6424,f6429,f6431,f6432,f6435,f6773,f6774,f6643,f6644,f6645,f6649,f6650,f6651,f6687,f6689,f6694,f6698,f6701,f6702,f6703,f6786,f6708,f6713,f6715,f6716,f6719,f6863,f6931,f6932,f7616,f7622,f7624,f7628,f7633,f7635,f7474,f7475,f7718,f7858,f7851,f7852,f7853,f7854,f7938,f7939,f7940,f7899,f7900,f7901,f7972,f7965,f7966,f7967,f7968,f8097,f8193,f8194,f8755,f8778,f8890,f8891,f8892,f8849,f8850,f8851,f8965,f8966,f9187,f9188,f9189,f9190,f9191,f9192,f9193,f9194,f9196,f9197,f9198,f9199,f9200,f9234,f9235,f9236,f9237,f9238,f9239,f9240,f9241,f9243,f9244,f9286,f9287,f9288,f9290,f9291,f9293,f9294,f9295,f9296,f9324,f9325,f9326,f9509,f9522,f9658,f9671,f9676,f9678,f9679,f10096,f9722,f9723,f9724,f9727,f9728,f9730,f9738,f9739,f9742,f9746,f9872,f9885,f9987,f9975,f10750,f10749,f10933,f10931,f10804,f10805,f10806,f11007,f11012,f11013,f11014,f11386,f11391,f11392,f11393,f11665,f11666,f11667,f117,f12719,f12717,f12130,f12131,f12132,f12133,f12134,f12135,f12136,f12137,f12139,f12140,f12712,f12186,f12187,f12188,f12190,f12191,f12193,f12195,f12196,f12637,f12797,f12630,f12413,f12977,f12978,f12986,f12987,f12995,f12996,f13040,f13000])).
fof(f13079,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f87,f86,f85,f84,f83,f88,f89,f90,f91,f92,f93,f94,f95,f96,f97,f98,f99,f100,f101,f102,f103,f104,f105,f106,f111,f110,f109,f108,f107,f112,f113,f114,f127,f116,f144,f146,f147,f155,f184,f194,f195,f227,f228,f240,f250,f251,f294,f297,f464,f466,f468,f479,f518,f520,f522,f533,f601,f602,f603,f604,f614,f670,f672,f674,f708,f709,f710,f733,f747,f777,f900,f901,f992,f994,f996,f1497,f1083,f1084,f1085,f1086,f1087,f1088,f1089,f1107,f1108,f1109,f1110,f1111,f1131,f1133,f1135,f1150,f1151,f1263,f1264,f1265,f1343,f1331,f1637,f1650,f1651,f1652,f1653,f1566,f1567,f1568,f1692,f1693,f1694,f1698,f1699,f1716,f1718,f1721,f1726,f1746,f1854,f1855,f1997,f1998,f1999,f2003,f2004,f2022,f2024,f2027,f2032,f2056,f2347,f2082,f2083,f2084,f2085,f2086,f2087,f2088,f2344,f2110,f2111,f2112,f2114,f2116,f2374,f2637,f2421,f2574,f2702,f2703,f2704,f2708,f2709,f2710,f2731,f2733,f2736,f2740,f2743,f2771,f2927,f2928,f2929,f2931,f2933,f2934,f2954,f2955,f3057,f3058,f3059,f3063,f3064,f3065,f3088,f3090,f3093,f3097,f3100,f3131,f3105,f3275,f3276,f4151,f4147,f4016,f4017,f4018,f4019,f4020,f4021,f4022,f4023,f4024,f4025,f4051,f4052,f4053,f4054,f4056,f4057,f4059,f4060,f4354,f4258,f4259,f4260,f4261,f4263,f4264,f4266,f4267,f4328,f4399,f4430,f4431,f4432,f4433,f4434,f4435,f4436,f4437,f4438,f4439,f4616,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f5000,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f5864,f5518,f5519,f6014,f6015,f6016,f6018,f6019,f6021,f6213,f6214,f6215,f6218,f6219,f6220,f6237,f6224,f6225,f6226,f6229,f6230,f6231,f6232,f6239,f6488,f6489,f6361,f6362,f6363,f6367,f6368,f6369,f6403,f6405,f6410,f6414,f6417,f6418,f6419,f6501,f6424,f6429,f6431,f6432,f6435,f6773,f6774,f6643,f6644,f6645,f6649,f6650,f6651,f6687,f6689,f6694,f6698,f6701,f6702,f6703,f6786,f6708,f6713,f6715,f6716,f6719,f6863,f6931,f6932,f7616,f7622,f7624,f7628,f7633,f7635,f7474,f7475,f7718,f7858,f7851,f7852,f7853,f7854,f7938,f7939,f7940,f7899,f7900,f7901,f7972,f7965,f7966,f7967,f7968,f8097,f8193,f8194,f8755,f8778,f8890,f8891,f8892,f8849,f8850,f8851,f8965,f8966,f9187,f9188,f9189,f9190,f9191,f9192,f9193,f9194,f9196,f9197,f9198,f9199,f9200,f9234,f9235,f9236,f9237,f9238,f9239,f9240,f9241,f9243,f9244,f9286,f9287,f9288,f9290,f9291,f9293,f9294,f9295,f9296,f9324,f9325,f9326,f9509,f9522,f9658,f9671,f9676,f9678,f9679,f10096,f9722,f9723,f9724,f9727,f9728,f9730,f9738,f9739,f9742,f9746,f9872,f9885,f9987,f9975,f10750,f10749,f10933,f10931,f10804,f10805,f10806,f11007,f11012,f11013,f11014,f11386,f11391,f11392,f11393,f11665,f11666,f11667,f117,f12719,f12717,f12130,f12131,f12132,f12133,f12134,f12135,f12136,f12137,f12139,f12140,f12712,f12186,f12187,f12188,f12190,f12191,f12193,f12195,f12196,f12637,f12797,f12630,f12413,f12977,f12978,f12986,f12987,f12995,f12996,f13040])).
fof(f13041,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f13040,f120])).
fof(f13001,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f13000,f120])).
fof(f12971,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12970,f120])).
fof(f12085,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f120,f11988])).
fof(f11988,plain,(
  s(sK0,sK1) = s(sK0,sK5) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11987,f127])).
fof(f11987,plain,(
  s(sK0,sK5) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11986,f127])).
fof(f11986,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11985,f94])).
fof(f11985,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11915,f94])).
fof(f11915,plain,(
  s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f11385,f107])).
fof(f11385,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f117,f121])).
fof(f12046,plain,(
  $false | ($spl0 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f12045,f126])).
fof(f12045,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12044,f7899])).
fof(f12044,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12043,f7899])).
fof(f12043,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12042,f103])).
fof(f12042,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12041,f103])).
fof(f12041,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12040,f101])).
fof(f12040,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12039,f101])).
fof(f12039,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11916,f102])).
fof(f11916,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(superposition,[],[f102,f11385])).
fof(f12038,plain,(
  $false | ($spl0 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f12037,f126])).
fof(f12037,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12036,f7899])).
fof(f12036,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12035,f7899])).
fof(f12035,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12034,f103])).
fof(f12034,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12033,f103])).
fof(f12033,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12032,f101])).
fof(f12032,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12031,f101])).
fof(f12031,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12030,f102])).
fof(f12030,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11908,f102])).
fof(f11908,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f11385,f110])).
fof(f12012,plain,(
  $false | ($spl0 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f12011,f126])).
fof(f12011,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12010,f7899])).
fof(f12010,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12009,f7899])).
fof(f12009,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12008,f103])).
fof(f12008,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12007,f103])).
fof(f12007,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12006,f101])).
fof(f12006,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12005,f101])).
fof(f12005,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f12004,f102])).
fof(f12004,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f11912,f102])).
fof(f11912,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f11385,f110])).
fof(f11357,plain,(
  $false | ($spl0 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f11356,f126])).
fof(f11356,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11355,f7899])).
fof(f11355,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11354,f7899])).
fof(f11354,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11353,f103])).
fof(f11353,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11352,f103])).
fof(f11352,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11351,f101])).
fof(f11351,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11350,f101])).
fof(f11350,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11235,f102])).
fof(f11235,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(superposition,[],[f102,f134])).
fof(f134,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f133,f121])).
fof(f133,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f117,f123])).
fof(f11349,plain,(
  $false | ($spl0 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f11348,f126])).
fof(f11348,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11347,f7899])).
fof(f11347,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11346,f7899])).
fof(f11346,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11345,f103])).
fof(f11345,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11344,f103])).
fof(f11344,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11343,f101])).
fof(f11343,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11342,f101])).
fof(f11342,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11341,f102])).
fof(f11341,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11227,f102])).
fof(f11227,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f134,f110])).
fof(f11326,plain,(
  $false | ($spl0 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f11325,f126])).
fof(f11325,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11324,f7899])).
fof(f11324,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11323,f7899])).
fof(f11323,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11322,f103])).
fof(f11322,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11321,f103])).
fof(f11321,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11320,f101])).
fof(f11320,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11319,f101])).
fof(f11319,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11318,f102])).
fof(f11318,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f11231,f102])).
fof(f11231,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f134,f110])).
fof(f10727,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f10726,f126])).
fof(f10726,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10725,f7899])).
fof(f10725,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10724,f7899])).
fof(f10724,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10723,f103])).
fof(f10723,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10722,f103])).
fof(f10722,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10721,f101])).
fof(f10721,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10720,f101])).
fof(f10720,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10609,f102])).
fof(f10609,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(superposition,[],[f102,f135])).
fof(f135,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f134,f119])).
fof(f10719,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f10718,f126])).
fof(f10718,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10717,f7899])).
fof(f10717,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10716,f7899])).
fof(f10716,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10715,f103])).
fof(f10715,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10714,f103])).
fof(f10714,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10713,f101])).
fof(f10713,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10712,f101])).
fof(f10712,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10711,f102])).
fof(f10711,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10601,f102])).
fof(f10601,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f135,f110])).
fof(f10697,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f10696,f126])).
fof(f10696,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10695,f7899])).
fof(f10695,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10694,f7899])).
fof(f10694,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10693,f103])).
fof(f10693,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10692,f103])).
fof(f10692,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10691,f101])).
fof(f10691,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10690,f101])).
fof(f10690,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10689,f102])).
fof(f10689,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10605,f102])).
fof(f10605,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f135,f110])).
fof(f2662,plain,(
  ( ! [X96] : (s(bool,'T_0') = s(bool,X96)) ) | $spl20),
  inference(cnf_transformation,[],[f2662_D])).
fof(f2662_D,plain,(
  ( ! [X96] : (s(bool,'T_0') = s(bool,X96)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2663,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98)))) ) | $spl22),
  inference(cnf_transformation,[],[f2663_D])).
fof(f2663_D,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2661,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98))) | s(bool,'T_0') = s(bool,X96)) ) | $spl13),
  inference(subsumption_resolution,[],[f2501,f2374])).
fof(f2501,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X96),s(X99,X100))) | s(bool,'T_0') = s(bool,X96)) ) | $spl13),
  inference(superposition,[],[f1532,f2344])).
fof(f1532,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77)))) ) | $spl13),
  inference(subsumption_resolution,[],[f1224,f158])).
fof(f1224,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77)))) ) | $spl13),
  inference(superposition,[],[f158,f227])).
fof(f861,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) | $spl14),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f863,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) | $spl18),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f710,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X27,X28))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f194,f90])).
fof(f862,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) | $spl16),
  inference(cnf_transformation,[],[f862_D])).
fof(f862_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f709,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X20,X21))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X18),s(X22,X23))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f194,f144])).
fof(f708,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f194,f144])).
fof(f156,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f158,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl13),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f155,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f90])).
fof(f132,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f131])).
fof(f131,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f130,f119])).
fof(f130,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f129,f121])).
fof(f129,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl1 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f128,f123])).
fof(f128,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f118,f125])).
fof(f117,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl0),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f119,plain,(
  s(sK0,sK1) = s(sK0,sK5) | $spl2),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  s(sK0,sK1) = s(sK0,sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f83,plain,(
  s(sK0,sK1) = s(sK0,sK5) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0')))))))))),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  (s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | (s(sK0,sK1) = s(sK0,sK5) & s(sK0,sK2) = s(sK0,sK6) & s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) & s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8))) & (s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | s(sK0,sK1) != s(sK0,sK5) | s(sK0,sK2) != s(sK0,sK6) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f78])).
fof(f78,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) | (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8))) & (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) | s(X0,X1) != s(X0,X5) | s(X0,X2) != s(X0,X6) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X7) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X8)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) | (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8))) & (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) | (s(X0,X1) != s(X0,X5) | s(X0,X2) != s(X0,X6) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X7) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X8))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) <~> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) <=> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X5))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X7))),s('type.semi_ring.semi_ring'(X0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.semi_ring.semi_ring'(X0),'const.bool.ARB_0'))))))))) <=> (s(X0,X1) = s(X0,X5) & s(X0,X2) = s(X0,X6) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X7) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X8)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X7))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X8))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X9))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X10))),s('type.semi_ring.semi_ring'(X6),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X11))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X12))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X13))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X14))),s('type.semi_ring.semi_ring'(X6),'const.bool.ARB_0'))))))))) <=> (s(X6,X7) = s(X6,X11) & s(X6,X8) = s(X6,X12) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X13) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X14)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X7))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X8))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X9))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X10))),s('type.semi_ring.semi_ring'(X6),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X11))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X12))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X13))),s('type.semi_ring.semi_ring'(X6),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X14))),s('type.semi_ring.semi_ring'(X6),'const.bool.ARB_0'))))))))) <=> (s(X6,X7) = s(X6,X11) & s(X6,X8) = s(X6,X12) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X13) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X14)))),
  file('thm.semi_ring.semi_ring_literal_|49||49|','thm.semi_ring.semi_ring_literal_|49||49|')).
fof(f121,plain,(
  s(sK0,sK2) = s(sK0,sK6) | $spl4),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  s(sK0,sK2) = s(sK0,sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f84,plain,(
  s(sK0,sK2) = s(sK0,sK6) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0')))))))))),
  inference(cnf_transformation,[],[f79])).
fof(f123,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | $spl6),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f85,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK7) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0')))))))))),
  inference(cnf_transformation,[],[f79])).
fof(f125,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl8),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f86,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK8) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) = s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0')))))))))),
  inference(cnf_transformation,[],[f79])).
fof(f118,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) | $spl1),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f120,plain,(
  s(sK0,sK1) != s(sK0,sK5) | $spl3),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  s(sK0,sK1) != s(sK0,sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f122,plain,(
  s(sK0,sK2) != s(sK0,sK6) | $spl5),
  inference(cnf_transformation,[],[f122_D])).
fof(f122_D,plain,(
  s(sK0,sK2) != s(sK0,sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f124,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | $spl7),
  inference(cnf_transformation,[],[f124_D])).
fof(f124_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f126,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) | $spl9),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f87,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK7) | s(sK0,sK2) != s(sK0,sK6) | s(sK0,sK1) != s(sK0,sK5) | s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0'))))))))) != s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK5))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SR|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK7))),s('type.semi_ring.semi_ring'(sK0),'const.semi_ring.semi_ring_SRM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.semi_ring.semi_ring'(sK0),'const.bool.ARB_0')))))))))),
  inference(cnf_transformation,[],[f79])).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 86
% Passive clauses: 506
% Generated clauses: 12874
% Final active clauses: 71
% Final passive clauses: 410
% Input formulas: 26
% Initial clauses: 34
% 
% Duplicate literals: 586
% Trivial inequalities: 128
% Fw subsumption resolutions: 680
% Fw demodulations: 1637
% Bw demodulations: 6
% Condensations: 1650
% Global subsumptions: 3
% 
% Simple tautologies: 125
% Equational tautologies: 56
% Forward subsumptions: 7383
% Fw demodulations to eq. taut.: 96
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 292
% Forward superposition: 4236
% Backward superposition: 3338
% Self superposition: 188
% Equality factoring: 62
% Equality resolution: 16
% 
% Split clauses: 10
% Split components: 23
% SAT solver clauses: 549
% SAT solver unit clauses: 47
% SAT solver binary clauses: 146
% 
% Sat splits: 10
% Sat splitting refutations: 18
% 
% Memory used [KB]: 10874
% Time elapsed: 0.585 s
% ------------------------------
% ------------------------------
% Success in time 1.904 s
