Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (84350)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 196
% Passive clauses: 3349
% Generated clauses: 25941
% Final active clauses: 142
% Final passive clauses: 2723
% Input formulas: 26
% Initial clauses: 57
% 
% Duplicate literals: 903
% Trivial inequalities: 368
% Fw subsumption resolutions: 3825
% Fw demodulations: 4215
% Bw demodulations: 45
% 
% Simple tautologies: 470
% Equational tautologies: 53
% Forward subsumptions: 10846
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 1422
% Factoring: 9
% Forward superposition: 8292
% Backward superposition: 6335
% Self superposition: 155
% Equality factoring: 104
% Equality resolution: 34
% 
% Split clauses: 485
% Split components: 993
% SAT solver clauses: 1445
% SAT solver unit clauses: 28
% SAT solver binary clauses: 260
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 3
% 
% Sat splits: 485
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1436
% TWLsolver calls for satisfiability: 302
% 
% Memory used [KB]: 11641
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (84369)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 274
% Passive clauses: 1355
% Generated clauses: 26514
% Final active clauses: 121
% Final passive clauses: 321
% Input formulas: 26
% Initial clauses: 57
% 
% Duplicate literals: 688
% Trivial inequalities: 90
% Fw subsumption resolutions: 3052
% Fw demodulations: 3994
% Bw demodulations: 809
% Condensations: 894
% Global subsumptions: 9
% 
% Simple tautologies: 177
% Equational tautologies: 52
% Forward subsumptions: 12556
% Bw demodulations to eq. taut.: 8
% 
% Binary resolution: 825
% Factoring: 4
% Forward superposition: 5917
% Backward superposition: 5617
% Self superposition: 195
% Equality factoring: 74
% Equality resolution: 9
% 
% Split clauses: 393
% Split components: 787
% SAT solver clauses: 3010
% SAT solver unit clauses: 69
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 25
% 
% Sat splits: 393
% Sat splitting refutations: 179
% 
% TWLsolver clauses: 2961
% TWLsolver calls for satisfiability: 520
% 
% Memory used [KB]: 14328
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (84392)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 880
% Passive clauses: 6168
% Generated clauses: 191916
% Final active clauses: 140
% Final passive clauses: 465
% Input formulas: 26
% Initial clauses: 57
% 
% Duplicate literals: 10608
% Trivial inequalities: 1479
% Fw subsumption resolutions: 19866
% Fw demodulations: 26760
% Bw demodulations: 3430
% Condensations: 21740
% Global subsumptions: 631
% 
% Simple tautologies: 1783
% Equational tautologies: 2280
% Forward subsumptions: 99397
% Fw demodulations to eq. taut.: 63
% Bw demodulations to eq. taut.: 363
% 
% Unit resulting resolution: 9810
% Factoring: 26
% Forward superposition: 51246
% Backward superposition: 36834
% Self superposition: 917
% Equality factoring: 574
% Equality resolution: 132
% 
% Split clauses: 428
% Split components: 863
% SAT solver clauses: 8424
% SAT solver unit clauses: 65
% SAT solver binary clauses: 530
% 
% Sat splits: 428
% Sat splitting refutations: 767
% 
% Memory used [KB]: 64732
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f6572,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f99,f156,f155,f98,f158,f97,f160,f154,f195,f232,f230,f464,f502,f652,f656,f773,f817,f815,f867,f3034,f3080,f3078,f4507,f4510,f4524,f4525,f4509,f4589,f4638,f4590,f4642,f4974,f5014,f4640,f5162,f5166,f5164,f5165,f5254,f5258,f5285,f5364,f5365,f5387,f5390,f5389,f5522,f5524,f5256,f5688,f5689,f5163,f6065,f6171,f6222,f6175,f6226,f6230,f4765,f4763,f6223,f4508,f6235,f6319,f6322,f6315,f6331,f6321,f6367,f6370,f6371,f6375,f6001,f6067,f6070,f6059,f6086,f6069,f6224,f6568])).
fof(f6568,plain,(
  $false | ($spl8 | $spl11 | $spl46)),
  inference(subsumption_resolution,[],[f6533,f6069])).
fof(f6533,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl8 | $spl11 | $spl46)),
  inference(resolution,[],[f6130,f734])).
fof(f734,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl8),
  inference(resolution,[],[f229,f103])).
fof(f103,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f229,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl8),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f6130,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl11 | $spl46)),
  inference(subsumption_resolution,[],[f6096,f231])).
fof(f231,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2) | $spl11),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f6096,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl46),
  inference(resolution,[],[f6069,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ((~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))))) & (p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))))))) & (! [X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ? [X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))))) & (! [X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)))),
  inference(rectify,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ? [X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))) & (! [X3] : ((p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2] : (~s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7,X8] : (~s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.pred_set.NOT_EQUAL_SETS')).
fof(f6224,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | $spl48),
  inference(cnf_transformation,[],[f6224_D])).
fof(f6224_D,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6069,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl46),
  inference(cnf_transformation,[],[f6069_D])).
fof(f6069_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6086,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6085,f138])).
fof(f138,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f6085,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6044,f120])).
fof(f120,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.AND_CLAUSES|split|3')).
fof(f6044,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(superposition,[],[f669,f2069])).
fof(f2069,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f1520,f139])).
fof(f139,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f1520,plain,(
  ( ! [X17] : (p(s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,X17),s(fun(sK0,bool),sK2)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f868,f135])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.EQ_CLAUSES|split|2')).
fof(f868,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f822,f229])).
fof(f822,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl2),
  inference(resolution,[],[f156,f102])).
fof(f102,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X9,X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X7))),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.pred_set.INSERT_SUBSET')).
fof(f669,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X2,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(superposition,[],[f154,f110])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f79])).
fof(f6059,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f6003])).
fof(f6003,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f669,f868])).
fof(f6070,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6002,f138])).
fof(f6002,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f669,f4505])).
fof(f4505,plain,(
  ( ! [X50] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f4480])).
fof(f4480,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X50),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f136,f2069])).
fof(f136,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.EQ_CLAUSES|split|3')).
fof(f6067,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | $spl44),
  inference(cnf_transformation,[],[f6067_D])).
fof(f6067_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6001,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl4)),
  inference(resolution,[],[f669,f775])).
fof(f775,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2))))) ) | $spl4),
  inference(resolution,[],[f158,f102])).
fof(f6375,plain,(
  $false | ($spl0 | $spl2 | $spl8 | $spl29 | $spl47)),
  inference(subsumption_resolution,[],[f6374,f4642])).
fof(f6374,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(subsumption_resolution,[],[f6364,f145])).
fof(f145,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET',aHL_FALSITY)).
fof(f6364,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(superposition,[],[f6228,f152])).
fof(f152,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET',aHL_BOOL_CASES)).
fof(f6228,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(subsumption_resolution,[],[f6059,f6068])).
fof(f6068,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(cnf_transformation,[],[f6068_D])).
fof(f6068_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f6371,plain,(
  $false | ($spl0 | $spl2 | $spl8 | $spl29 | $spl47)),
  inference(subsumption_resolution,[],[f6348,f4642])).
fof(f6348,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(resolution,[],[f6228,f141])).
fof(f141,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.EQ_CLAUSES|split|0')).
fof(f6370,plain,(
  $false | ($spl0 | $spl2 | $spl8 | $spl29 | $spl47)),
  inference(subsumption_resolution,[],[f6347,f4642])).
fof(f6347,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(resolution,[],[f6228,f139])).
fof(f6367,plain,(
  $false | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(trivial_inequality_removal,[],[f6340])).
fof(f6340,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),sK2) | ($spl0 | $spl2 | $spl8 | $spl47)),
  inference(resolution,[],[f6228,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))) & ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & ~s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & ~s(fun(X6,bool),X7) = s(fun(X6,bool),X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.pred_set.PSUBSET_DEF')).
fof(f6321,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl54),
  inference(cnf_transformation,[],[f6321_D])).
fof(f6321_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6331,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6330,f138])).
fof(f6330,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6282,f120])).
fof(f6282,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(superposition,[],[f670,f2069])).
fof(f670,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,sK1,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(superposition,[],[f154,f110])).
fof(f6315,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f6237])).
fof(f6237,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f670,f868])).
fof(f6322,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f6236,f138])).
fof(f6236,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f670,f4505])).
fof(f6319,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | $spl52),
  inference(cnf_transformation,[],[f6319_D])).
fof(f6319_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6235,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK2,sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl4)),
  inference(resolution,[],[f670,f775])).
fof(f4508,plain,(
  ( ! [X6] : (~p(s(bool,X6))) ) | $spl20),
  inference(cnf_transformation,[],[f4508_D])).
fof(f4508_D,plain,(
  ( ! [X6] : (~p(s(bool,X6))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f6223,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | $spl49),
  inference(cnf_transformation,[],[f6223_D])).
fof(f6223_D,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f4763,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X2))) ) | $spl29),
  inference(subsumption_resolution,[],[f4752,f139])).
fof(f4752,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X2))) ) | $spl29),
  inference(superposition,[],[f4642,f116])).
fof(f116,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.IMP_ANTISYM_AX')).
fof(f4765,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) ) | $spl29),
  inference(subsumption_resolution,[],[f4753,f139])).
fof(f4753,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))))) ) | $spl29),
  inference(superposition,[],[f4642,f116])).
fof(f6230,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl4 | $spl47)),
  inference(subsumption_resolution,[],[f6001,f6068])).
fof(f6226,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl50),
  inference(cnf_transformation,[],[f6226_D])).
fof(f6226_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f6175,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl47)),
  inference(resolution,[],[f6068,f669])).
fof(f6222,plain,(
  $false | ($spl0 | $spl2 | $spl8 | $spl24 | $spl47)),
  inference(subsumption_resolution,[],[f6074,f6068])).
fof(f6074,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2 | $spl8 | $spl24)),
  inference(subsumption_resolution,[],[f6073,f143])).
fof(f143,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : (p(s(bool,'F_0')) => p(s(bool,X8)))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.FALSITY')).
fof(f6073,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl2 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f6059,f4638])).
fof(f6171,plain,(
  $false | ($spl0 | $spl8 | $spl24 | $spl46)),
  inference(subsumption_resolution,[],[f6136,f6069])).
fof(f6136,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl8 | $spl24 | $spl46)),
  inference(resolution,[],[f6128,f734])).
fof(f6128,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl24 | $spl46)),
  inference(subsumption_resolution,[],[f6127,f145])).
fof(f6127,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl24 | $spl46)),
  inference(forward_demodulation,[],[f6092,f4638])).
fof(f6092,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl46)),
  inference(resolution,[],[f6069,f667])).
fof(f667,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X0,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X0,sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(superposition,[],[f154,f109])).
fof(f6065,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl4 | $spl24)),
  inference(subsumption_resolution,[],[f6064,f120])).
fof(f6064,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK5(sK1,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl4 | $spl24)),
  inference(forward_demodulation,[],[f6001,f4638])).
fof(f5163,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(cnf_transformation,[],[f5163_D])).
fof(f5163_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f5689,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | ($spl2 | $spl8 | $spl31)),
  inference(subsumption_resolution,[],[f5678,f139])).
fof(f5678,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | ($spl2 | $spl8 | $spl31)),
  inference(superposition,[],[f5395,f116])).
fof(f5395,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl2 | $spl8 | $spl31)),
  inference(resolution,[],[f5385,f1522])).
fof(f1522,plain,(
  ( ! [X19] : (p(s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,X19),s(fun(sK0,bool),sK2)))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f868,f138])).
fof(f5385,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl31),
  inference(trivial_inequality_removal,[],[f5381])).
fof(f5381,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl31),
  inference(superposition,[],[f5013,f141])).
fof(f5013,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl31),
  inference(cnf_transformation,[],[f5013_D])).
fof(f5013_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f5688,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X2))) ) | ($spl2 | $spl8 | $spl31)),
  inference(subsumption_resolution,[],[f5677,f139])).
fof(f5677,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X2))) ) | ($spl2 | $spl8 | $spl31)),
  inference(superposition,[],[f5395,f116])).
fof(f5256,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | $spl36),
  inference(cnf_transformation,[],[f5256_D])).
fof(f5256_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5524,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(cnf_transformation,[],[f5524_D])).
fof(f5524_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5522,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f5440,f868])).
fof(f5440,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f887,f774])).
fof(f774,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl4),
  inference(resolution,[],[f158,f103])).
fof(f887,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,sK1,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X3))))) ) | $spl6),
  inference(superposition,[],[f159,f110])).
fof(f159,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | $spl6),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5389,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl41),
  inference(cnf_transformation,[],[f5389_D])).
fof(f5389_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f5390,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl31),
  inference(subsumption_resolution,[],[f5377,f139])).
fof(f5377,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl31),
  inference(superposition,[],[f5013,f116])).
fof(f5387,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,X2))) ) | $spl31),
  inference(subsumption_resolution,[],[f5376,f139])).
fof(f5376,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ~p(s(bool,X2))) ) | $spl31),
  inference(superposition,[],[f5013,f116])).
fof(f5365,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f5284,f5165])).
fof(f5284,plain,(
  ~p(s(bool,'T_0')) | $spl34),
  inference(resolution,[],[f5165,f147])).
fof(f147,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.NOT_CLAUSES|split|1')).
fof(f5364,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f5283,f5165])).
fof(f5283,plain,(
  ~p(s(bool,'T_0')) | $spl34),
  inference(resolution,[],[f5165,f149])).
fof(f149,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.bool.NOT_CLAUSES|split|2')).
fof(f5285,plain,(
  $false | $spl34),
  inference(resolution,[],[f5165,f145])).
fof(f5258,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl38),
  inference(cnf_transformation,[],[f5258_D])).
fof(f5258_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5254,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f5190,f868])).
fof(f5190,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f886,f774])).
fof(f886,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5(sK1,X2,sK0)),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X2))))) ) | $spl6),
  inference(superposition,[],[f159,f110])).
fof(f5165,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl34),
  inference(cnf_transformation,[],[f5165_D])).
fof(f5165_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5164,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl32),
  inference(cnf_transformation,[],[f5164_D])).
fof(f5164_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5166,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | p(s(bool,X1))) ) | ($spl8 | $spl30)),
  inference(subsumption_resolution,[],[f5149,f135])).
fof(f5149,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | p(s(bool,X1))) ) | ($spl8 | $spl30)),
  inference(superposition,[],[f5107,f113])).
fof(f113,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f63])).
fof(f5107,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl30)),
  inference(subsumption_resolution,[],[f5042,f150])).
fof(f150,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET',aHL_TRUTH)).
fof(f5042,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl30)),
  inference(superposition,[],[f1606,f5014])).
fof(f1606,plain,(
  ( ! [X18] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK1)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,X18),s(fun(sK0,bool),sK2)))) ) | $spl8),
  inference(resolution,[],[f734,f134])).
fof(f134,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f5162,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | ($spl8 | $spl30)),
  inference(subsumption_resolution,[],[f5148,f135])).
fof(f5148,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))))) ) | ($spl8 | $spl30)),
  inference(superposition,[],[f5107,f113])).
fof(f4640,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | $spl26),
  inference(cnf_transformation,[],[f4640_D])).
fof(f4640_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5014,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl30),
  inference(cnf_transformation,[],[f5014_D])).
fof(f5014_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4974,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f2142,f775])).
fof(f2142,plain,(
  ( ! [X16] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X16)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X16,sK1,sK0)),s(fun(sK0,bool),sK1)))) ) | $spl6),
  inference(resolution,[],[f1896,f139])).
fof(f1896,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X1,sK1,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X1))))) ) | $spl6),
  inference(resolution,[],[f874,f104])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f874,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X2))))) ) | $spl6),
  inference(resolution,[],[f159,f103])).
fof(f4642,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl29),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f4590,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl4),
  inference(resolution,[],[f775,f2026])).
fof(f2026,plain,(
  ( ! [X9] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X9)))) | s(bool,'T_0') != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X9)))) ) | $spl4),
  inference(resolution,[],[f1938,f138])).
fof(f1938,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X0))))) ) | $spl4),
  inference(resolution,[],[f702,f106])).
fof(f106,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f702,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X2))))) ) | $spl4),
  inference(resolution,[],[f421,f103])).
fof(f421,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK2)))) | $spl4),
  inference(resolution,[],[f158,f100])).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f4638,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl24),
  inference(cnf_transformation,[],[f4638_D])).
fof(f4638_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4589,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2))) | $spl4),
  inference(resolution,[],[f775,f2024])).
fof(f2024,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),X7)))) | s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X7)))) ) | $spl4),
  inference(resolution,[],[f1938,f135])).
fof(f4509,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))) ) | $spl22),
  inference(cnf_transformation,[],[f4509_D])).
fof(f4509_D,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4525,plain,(
  ( ! [X37,X36] : (~p(s(bool,X37)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X36),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X36),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4470,f134])).
fof(f4470,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,X37) | ~p(s(bool,X37)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X36),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X36),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f116,f2069])).
fof(f4524,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4469,f134])).
fof(f4469,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,X35) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X34),s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f116,f2069])).
fof(f4510,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | ~p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4432,f134])).
fof(f4432,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK1))) | ~p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X7),s(fun(sK0,bool),sK2))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f2069,f116])).
fof(f4507,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X6))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4431,f134])).
fof(f4431,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X5),s(fun(sK0,bool),sK2)))) | ~p(s(bool,X6))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f2069,f116])).
fof(f3078,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | $spl16),
  inference(cnf_transformation,[],[f3078_D])).
fof(f3078_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3080,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK1)))) | $spl18),
  inference(cnf_transformation,[],[f3080_D])).
fof(f3080_D,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK1)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3034,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))) | ($spl2 | $spl6 | $spl8)),
  inference(resolution,[],[f2699,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f2699,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK1)))) | ($spl2 | $spl6 | $spl8)),
  inference(resolution,[],[f875,f2588])).
fof(f2588,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f2557])).
fof(f2557,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(resolution,[],[f1513,f105])).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f1513,plain,(
  ( ! [X8] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X8,sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X8))))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f868,f104])).
fof(f875,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),X3))),s(fun(sK0,bool),sK1))))) ) | $spl6),
  inference(resolution,[],[f159,f102])).
fof(f867,plain,(
  $false | ($spl2 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f866,f145])).
fof(f866,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl4 | $spl7)),
  inference(forward_demodulation,[],[f821,f169])).
fof(f169,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl7),
  inference(resolution,[],[f160,f135])).
fof(f821,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ($spl2 | $spl4)),
  inference(resolution,[],[f156,f158])).
fof(f815,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl12),
  inference(cnf_transformation,[],[f815_D])).
fof(f815_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f817,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | $spl14),
  inference(cnf_transformation,[],[f817_D])).
fof(f817_D,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f773,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))) | $spl4),
  inference(resolution,[],[f158,f108])).
fof(f656,plain,(
  $false | ($spl4 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f655,f145])).
fof(f655,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl7 | $spl10)),
  inference(forward_demodulation,[],[f654,f169])).
fof(f654,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | ($spl4 | $spl10)),
  inference(forward_demodulation,[],[f421,f232])).
fof(f652,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f469,f232])).
fof(f469,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2) | $spl0),
  inference(resolution,[],[f154,f107])).
fof(f502,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f468,f230])).
fof(f468,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(resolution,[],[f154,f106])).
fof(f464,plain,(
  $false | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f422,f230])).
fof(f422,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl4),
  inference(resolution,[],[f158,f101])).
fof(f101,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f230,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl9),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f232,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | $spl10),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f195,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(resolution,[],[f155,f108])).
fof(f154,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f160,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | $spl7),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f97,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3),s(fun(sK0,bool),sK1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))))) & (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f68])).
fof(f68,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X4),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f67])).
fof(f67,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2)))))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <~> ? [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),X1))),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ? [X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X7))),s(fun(X6,bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ? [X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X7))),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.PSUBSET_INSERT_SUBSET','thm.pred_set.PSUBSET_INSERT_SUBSET')).
fof(f158,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | $spl4),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f98,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f69])).
fof(f155,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f156,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))))) ) | $spl2),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f99,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) )),
  inference(cnf_transformation,[],[f69])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 160
% Passive clauses: 2255
% Generated clauses: 6474
% Final active clauses: 119
% Final passive clauses: 1792
% Input formulas: 26
% Initial clauses: 56
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 201
% Trivial inequalities: 105
% Fw subsumption resolutions: 643
% Fw demodulations: 632
% Bw demodulations: 26
% 
% Simple tautologies: 289
% Equational tautologies: 15
% Forward subsumptions: 2317
% 
% Binary resolution: 1423
% Factoring: 5
% Forward superposition: 2422
% Backward superposition: 904
% Equality resolution: 2
% 
% Split clauses: 34
% Split components: 70
% SAT solver clauses: 50
% SAT solver unit clauses: 3
% SAT solver binary clauses: 6
% 
% Sat splits: 34
% Sat splitting refutations: 15
% 
% Memory used [KB]: 3198
% Time elapsed: 0.142 s
% ------------------------------
% ------------------------------
% Success in time 11.93 s
