Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.realax
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.realax
% SZS output start Proof for thm.realax
fof(f5344,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5338,f28,f40,f38,f27,f42,f26,f39,f44,f5013,f4071])).
fof(f4071,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f4070,f40])).
fof(f4070,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f4069,f42])).
fof(f4069,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))))))) | ($spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f4066])).
fof(f4066,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))))))) | ($spl4 | $spl6)),
  inference(superposition,[],[f3225,f42])).
fof(f3225,plain,(
  ( ! [X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f3224,f29])).
fof(f29,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.realax.real_ABS_REP_CLASS|split|1',aHL_TRUTH)).
fof(f3224,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f3219,f252])).
fof(f252,plain,(
  s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) | ($spl4 | $spl6)),
  inference(superposition,[],[f237,f77])).
fof(f77,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f76,f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.realax.real_ABS_REP_CLASS|split|1',aHL_FALSITY)).
fof(f76,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) | ($spl4 | $spl6)),
  inference(superposition,[],[f74,f32])).
fof(f32,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.realax.real_ABS_REP_CLASS|split|1',aHL_BOOL_CASES)).
fof(f74,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f44,f42])).
fof(f237,plain,(
  ( ! [X0] : (s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0)))) ) | $spl4),
  inference(superposition,[],[f36,f42])).
fof(f36,plain,(
  ( ! [X0,X1] : (s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) )),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1] : s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X10),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X11))) = s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X10))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X11)))),
  file('thm.realax.real_ABS_REP_CLASS|split|1','ARITY_1_2_const.realax.treal_eq')).
fof(f3219,plain,(
  ( ! [X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))))) ) | $spl4),
  inference(superposition,[],[f3065,f42])).
fof(f3065,plain,(
  ( ! [X2,X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))))) )),
  inference(forward_demodulation,[],[f33,f36])).
fof(f33,plain,(
  ( ! [X2,X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : ((! [X2] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) & (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0)))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0))))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f23])).
fof(f23,plain,(
  ! [X0] : ((! [X2] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) & (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | ? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))))),
  inference(rectify,[],[f22])).
fof(f22,plain,(
  ! [X0] : ((! [X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) & (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | ? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))))),
  inference(nnf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0] : (? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) <=> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : (? [X9] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X9))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X9)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X7) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X9)))) <=> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X7) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X7))))))),
  file('thm.realax.real_ABS_REP_CLASS|split|1','thm.realax.real_bijections|split|1')).
fof(f5013,plain,(
  $false | ($spl0 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f5012,f29])).
fof(f5012,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f5011,f252])).
fof(f5011,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f5009])).
fof(f5009,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) | ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) | ($spl0 | $spl4)),
  inference(superposition,[],[f4079,f42])).
fof(f4079,plain,(
  ( ! [X1] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) | ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))) ) | $spl0),
  inference(forward_demodulation,[],[f38,f36])).
fof(f44,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) | $spl6),
  inference(cnf_transformation,[],[f44_D])).
fof(f44_D,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f39,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | $spl2),
  inference(cnf_transformation,[],[f39_D])).
fof(f39_D,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f26,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))))),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ((p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0)))))) & (! [X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f20])).
fof(f20,plain,(
  ? [X0] : ((? [X2] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X2)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) & (! [X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))))),
  inference(rectify,[],[f19])).
fof(f19,plain,(
  ? [X0] : ((? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) & (! [X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ? [X0] : (? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) <~> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ~! [X0] : (? [X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1)))) <=> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (? [X7] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7)))) <=> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (? [X7] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7)))) & s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X7)))) <=> s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X6))))))),
  file('thm.realax.real_ABS_REP_CLASS|split|1','thm.realax.real_ABS_REP_CLASS|split|1')).
fof(f42,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) | $spl4),
  inference(cnf_transformation,[],[f42_D])).
fof(f42_D,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f27,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK1)))),
  inference(cnf_transformation,[],[f21])).
fof(f38,plain,(
  ( ! [X1] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))) ) | $spl0),
  inference(cnf_transformation,[],[f38_D])).
fof(f38_D,plain,(
  ( ! [X1] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f40,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | $spl3),
  inference(cnf_transformation,[],[f40_D])).
fof(f40_D,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f28,plain,(
  ( ! [X1] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0))))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),X1))))) )),
  inference(cnf_transformation,[],[f21])).
fof(f5338,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f5337,f4767])).
fof(f4767,plain,(
  p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0))))) | $spl2),
  inference(trivial_inequality_removal,[],[f4759])).
fof(f4759,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0))))) | $spl2),
  inference(superposition,[],[f2234,f39])).
fof(f2234,plain,(
  ( ! [X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0)))))) )),
  inference(forward_demodulation,[],[f34,f36])).
fof(f34,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0)))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0))))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0)))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f5337,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0))))) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f5331])).
fof(f5331,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) | ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0))))) | ($spl0 | $spl2)),
  inference(superposition,[],[f4079,f4766])).
fof(f4766,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)))) | $spl2),
  inference(trivial_inequality_removal,[],[f4760])).
fof(f4760,plain,(
  s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),sK0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(sK0)))) | $spl2),
  inference(superposition,[],[f35,f39])).
fof(f35,plain,(
  ( ! [X0] : (s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) != s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.real_REP_CLASS_1'(s('type.realax.real','const.realax.real_ABS_CLASS_1'(s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0))))) | s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),X0) = s(fun('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),bool),'const.realax.treal_eq_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),sK2(X0))))) )),
  inference(cnf_transformation,[],[f24])).
% SZS output end Proof for thm.realax
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 53
% Passive clauses: 1089
% Generated clauses: 5161
% Final active clauses: 42
% Final passive clauses: 895
% Input formulas: 8
% Initial clauses: 12
% 
% Duplicate literals: 497
% Trivial inequalities: 56
% Fw subsumption resolutions: 438
% Fw demodulations: 99
% Bw demodulations: 3
% 
% Simple tautologies: 114
% Equational tautologies: 83
% Forward subsumptions: 2008
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 26
% Forward superposition: 1944
% Backward superposition: 1578
% Self superposition: 234
% Equality factoring: 210
% Equality resolution: 3
% 
% Split clauses: 313
% Split components: 627
% SAT solver clauses: 823
% SAT solver unit clauses: 238
% SAT solver binary clauses: 251
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 3
% 
% Sat splits: 313
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 820
% TWLsolver calls for satisfiability: 46
% 
% Memory used [KB]: 2814
% Time elapsed: 0.203 s
% ------------------------------
% ------------------------------
% Success in time 0.208 s
