Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (111593)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 690
% Passive clauses: 9170
% Generated clauses: 19413
% Final active clauses: 687
% Final passive clauses: 8042
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 362
% Trivial inequalities: 67
% Fw subsumption resolutions: 1672
% Fw demodulations: 619
% Bw demodulations: 3
% 
% Simple tautologies: 263
% Equational tautologies: 65
% Forward subsumptions: 6944
% 
% Binary resolution: 10440
% Factoring: 1645
% Forward superposition: 1885
% Backward superposition: 2254
% Self superposition: 36
% Equality factoring: 28
% Equality resolution: 22
% 
% Split clauses: 375
% Split components: 752
% SAT solver clauses: 681
% SAT solver unit clauses: 8
% SAT solver binary clauses: 196
% 
% Sat splits: 375
% 
% TWLsolver clauses: 681
% TWLsolver calls for satisfiability: 418
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (111608)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 396
% Passive clauses: 2771
% Generated clauses: 12520
% Final active clauses: 382
% Final passive clauses: 2327
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 342
% Trivial inequalities: 20
% Fw subsumption resolutions: 717
% Fw demodulations: 511
% Bw demodulations: 59
% Condensations: 424
% Global subsumptions: 22
% 
% Simple tautologies: 278
% Equational tautologies: 17
% Forward subsumptions: 4105
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 2857
% Factoring: 62
% Forward superposition: 2205
% Backward superposition: 2206
% Self superposition: 38
% Equality factoring: 28
% Equality resolution: 7
% 
% Split clauses: 108
% Split components: 216
% SAT solver clauses: 3507
% SAT solver unit clauses: 232
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 108
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 3419
% TWLsolver calls for satisfiability: 1957
% 
% Memory used [KB]: 15607
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (111620)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2208
% Passive clauses: 17395
% Generated clauses: 115299
% Final active clauses: 1220
% Final passive clauses: 4499
% Input formulas: 59
% Initial clauses: 229
% Discarded non-redundant clauses: 13645
% 
% Duplicate literals: 2447
% Trivial inequalities: 108
% Fw subsumption resolutions: 2654
% Fw demodulations: 19903
% Bw demodulations: 1055
% Condensations: 2612
% Global subsumptions: 112
% 
% Simple tautologies: 239
% Equational tautologies: 1839
% Forward subsumptions: 65935
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 257
% 
% Unit resulting resolution: 27981
% Factoring: 4
% Forward superposition: 31517
% Backward superposition: 25409
% Self superposition: 136
% Equality factoring: 52
% Equality resolution: 6
% 
% Split clauses: 177
% Split components: 359
% SAT solver clauses: 18525
% SAT solver unit clauses: 1840
% SAT solver binary clauses: 5530
% 
% Sat splits: 177
% Sat splitting refutations: 662
% 
% Memory used [KB]: 105925
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (112040)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2630
% Passive clauses: 38381
% Generated clauses: 255789
% Final active clauses: 2091
% Final passive clauses: 25658
% Input formulas: 59
% Initial clauses: 213
% 
% Selected by SInE selection: 51
% SInE iterations: 1
% 
% Duplicate literals: 6666
% Trivial inequalities: 250
% Fw subsumption resolutions: 33075
% Fw demodulations: 5054
% Bw demodulations: 75
% 
% Simple tautologies: 4515
% Equational tautologies: 403
% Forward subsumptions: 162579
% 
% Binary resolution: 196178
% Factoring: 393
% Forward superposition: 7159
% Backward superposition: 3889
% Self superposition: 5
% Equality resolution: 31
% 
% Split clauses: 9060
% Split components: 19790
% SAT solver clauses: 9749
% SAT solver unit clauses: 28
% SAT solver binary clauses: 689
% 
% Sat splits: 9060
% Sat splitting refutations: 116
% 
% Memory used [KB]: 63581
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (112130)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 326
% Passive clauses: 3415
% Generated clauses: 15175
% Final active clauses: 318
% Final passive clauses: 2986
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 506
% Trivial inequalities: 70
% Fw subsumption resolutions: 1286
% Fw demodulations: 100
% Bw demodulations: 52
% Global subsumptions: 10
% 
% Simple tautologies: 373
% Equational tautologies: 61
% Forward subsumptions: 9101
% 
% Unit resulting resolution: 1563
% Factoring: 8
% Forward superposition: 4387
% Backward superposition: 6766
% Self superposition: 116
% Equality factoring: 6
% Equality resolution: 10
% 
% Split clauses: 18
% Split components: 37
% SAT solver clauses: 3509
% SAT solver unit clauses: 258
% SAT solver binary clauses: 308
% 
% Sat splits: 18
% Sat splitting refutations: 64
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (112150)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 946
% Passive clauses: 8402
% Generated clauses: 27642
% Final active clauses: 412
% Final passive clauses: 4097
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 1063
% Trivial inequalities: 86
% Fw subsumption resolutions: 1697
% Bw subsumption resolutions: 933
% Fw demodulations: 1752
% Bw demodulations: 1284
% Condensations: 1708
% Global subsumptions: 7
% 
% Simple tautologies: 1570
% Equational tautologies: 56
% Forward subsumptions: 10959
% Backward subsumptions: 1808
% 
% Binary resolution: 6785
% Factoring: 20
% Forward superposition: 5716
% Backward superposition: 3162
% Self superposition: 39
% Equality factoring: 12
% Equality resolution: 13
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 8680
% SAT solver unit clauses: 573
% SAT solver binary clauses: 2854
% 
% Sat splits: 3
% Sat splitting refutations: 265
% 
% Memory used [KB]: 15607
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% (112208)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 769
% Passive clauses: 5314
% Generated clauses: 19200
% Final active clauses: 429
% Final passive clauses: 2361
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 191
% Trivial inequalities: 19
% Fw demodulations: 2013
% Bw demodulations: 1086
% Global subsumptions: 351
% 
% Simple tautologies: 481
% Equational tautologies: 59
% Forward subsumptions: 9101
% Backward subsumptions: 296
% 
% Binary resolution: 5397
% Unit resulting resolution: 6332
% Factoring: 17
% Forward superposition: 1016
% Backward superposition: 854
% Equality resolution: 3
% 
% Split clauses: 127
% Split components: 264
% SAT solver clauses: 8800
% SAT solver unit clauses: 115
% SAT solver binary clauses: 4449
% 
% Sat splits: 127
% Sat splitting refutations: 369
% 
% Memory used [KB]: 13816
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f27512,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f27415])).
fof(f27415,plain,(
  s(bool,'T_0') != s(bool,'T_0')),
  inference(resolution,[],[f27398,f289])).
fof(f289,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.bool.EQ_CLAUSES|split|1')).
fof(f27398,plain,(
  ~p(s(bool,'T_0'))),
  inference(duplicate_literal_removal,[],[f27397])).
fof(f27397,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f27396,f1546])).
fof(f1546,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X20)))) )),
  inference(resolution,[],[f892,f288])).
fof(f288,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.bool.EQ_CLAUSES|split|0')).
fof(f892,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X3))))) )),
  inference(forward_demodulation,[],[f827,f815])).
fof(f815,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))),
  inference(resolution,[],[f271,f288])).
fof(f271,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(bool,X2)))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(bool,X2)))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X1,bool),X4),s(bool,X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X10)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X7,bool),X10),s(bool,X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X10)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X7,bool),X10),s(bool,X8)))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.quantHeuristics.GUESS_RULES_DISJ|split|5')).
fof(f827,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))),s(bool,X3))))) )),
  inference(resolution,[],[f271,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X25,X26] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X25))),s(bool,X26)))) <=> (p(s(bool,X25)) | p(s(bool,X26))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','|bslash||slash|_HL_0')).
fof(f27396,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f27257,f26686])).
fof(f26686,plain,(
  ( ! [X639,X640,X642,X641,X643] : (s(X640,happ(s(fun(X641,X640),happ(s(fun(bool,fun(X641,X640)),X642),s(bool,'T_0'))),s(X641,X643))) = s(X640,happ(s(fun(sK2,X640),'_dst_x0x1x0_3'(s(fun(bool,fun(X641,X640)),X642),s(fun(sK2,bool),sK5),s(X641,X643))),s(sK2,X639)))) )),
  inference(superposition,[],[f628,f16701])).
fof(f16701,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X20)))) )),
  inference(resolution,[],[f16688,f288])).
fof(f16688,plain,(
  ( ! [X157] : (p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X157))))) )),
  inference(global_subsumption,[],[f16687,f888])).
fof(f888,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f817,f815])).
fof(f817,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))),
  inference(resolution,[],[f271,f291])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.bool.EQ_CLAUSES|split|3')).
fof(f16687,plain,(
  ( ! [X157] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X157))))) )),
  inference(forward_demodulation,[],[f16686,f11612])).
fof(f11612,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X7)))))) )),
  inference(backward_demodulation,[],[f11393,f3793])).
fof(f3793,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X7)))))) )),
  inference(resolution,[],[f944,f288])).
fof(f944,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X0))))))) )),
  inference(resolution,[],[f272,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7,X6] : ! [X9,X10] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X10)))) <=> ! [X19] : (~p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X19)))) => ? [X20] : ~p(s(bool,happ(s(fun(X7,bool),X10),s(X7,happ(s(fun(X6,X7),X9),s(X6,X20))))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f272,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))))))),
  inference(cnf_transformation,[],[f182])).
fof(f11393,plain,(
  s(bool,'F_0') = s(bool,sK3)),
  inference(resolution,[],[f11239,f273])).
fof(f273,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','thm.bool.TRUTH')).
fof(f11239,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK3)),
  inference(backward_demodulation,[],[f10989,f4280])).
fof(f4280,plain,(
  s(bool,'F_0') = s(bool,sK3) | ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'T_0'))),s(sK2,sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'T_0')),sK4,sK1,sK2)))))),
  inference(resolution,[],[f1008,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f1008,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK3)),
  inference(superposition,[],[f272,f280])).
fof(f280,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5',aHL_BOOL_CASES)).
fof(f10989,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,happ(s(fun(X51,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(X51,bool),X52),s(bool,'T_0'))),s(X51,X53)))) )),
  inference(superposition,[],[f1748,f628])).
fof(f1748,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X8))),s(bool,'T_0')))) )),
  inference(resolution,[],[f893,f288])).
fof(f893,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f828,f815])).
fof(f828,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))))) )),
  inference(resolution,[],[f271,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f227])).
fof(f16686,plain,(
  ( ! [X157] : (s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,sK14(X157,sK5,sK4,sK1,sK2)))))) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X157))))) )),
  inference(forward_demodulation,[],[f16623,f628])).
fof(f16623,plain,(
  ( ! [X157] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,sK14(X157,sK5,sK4,sK1,sK2)))))))),s(bool,'F_0'))) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X157))))) )),
  inference(resolution,[],[f2329,f814])).
fof(f814,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,sK14(X0,sK5,sK4,sK1,sK2))))))) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X0))))) )),
  inference(resolution,[],[f271,f408])).
fof(f408,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f2329,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X11))),s(bool,'F_0')))) )),
  inference(resolution,[],[f1068,f292])).
fof(f292,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f1068,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f958,f948])).
fof(f948,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3)))))),
  inference(resolution,[],[f272,f292])).
fof(f958,plain,(
  ( ! [X4] : (p(s(bool,X4)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))))))))) )),
  inference(resolution,[],[f272,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f628,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X27,X28,X29,X30] : ! [X31,X32,X33,X34] : s(X27,happ(s(fun(X29,X27),'_dst_x0x1x0_3'(s(fun(X28,fun(X30,X27)),X31),s(fun(X29,X28),X32),s(X30,X33))),s(X29,X34))) = s(X27,happ(s(fun(X30,X27),happ(s(fun(X28,fun(X30,X27)),X31),s(X28,happ(s(fun(X29,X28),X32),s(X29,X34))))),s(X30,X33)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|5','_dst_x0x1x0_ax')).
fof(f27257,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))),s(sK2,sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')),X0,X1,sK2)))))) )),
  inference(resolution,[],[f26857,f409])).
fof(f26857,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK2),X1),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')))))) | ~p(s(bool,'T_0'))) )),
  inference(forward_demodulation,[],[f26835,f1546])).
fof(f26835,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'F_0')))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK2),X1),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))))))) )),
  inference(backward_demodulation,[],[f26686,f11980])).
fof(f11980,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK2),X1),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')))))) | ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))),s(sK2,happ(s(fun(X0,sK2),X1),s(X0,sK14(sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')),sK4,sK1,sK2),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')),X1,X0,sK2)))))))) )),
  inference(forward_demodulation,[],[f11783,f11393])).
fof(f11783,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0'))),s(sK2,happ(s(fun(X0,sK2),X1),s(X0,sK14(sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')),sK4,sK1,sK2),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,'F_0')),X1,X0,sK2))))))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK2),X1),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))))))) )),
  inference(backward_demodulation,[],[f11393,f6610])).
fof(f6610,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))),s(sK2,happ(s(fun(X0,sK2),X1),s(X0,sK14(sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3)),sK4,sK1,sK2),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3)),X1,X0,sK2))))))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK2),X1),s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))))))) )),
  inference(resolution,[],[f943,f408])).
fof(f943,plain,(
  ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3))),s(sK2,sK15('_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(sK2,bool),sK5),s(bool,sK3)),sK4,sK1,sK2)))))),
  inference(resolution,[],[f272,f409])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 296
% Passive clauses: 3107
% Generated clauses: 26612
% Final active clauses: 239
% Final passive clauses: 2360
% Input formulas: 59
% Initial clauses: 358
% 
% Duplicate literals: 1555
% Trivial inequalities: 108
% Fw demodulations: 3298
% Bw demodulations: 562
% Condensations: 2316
% Global subsumptions: 12
% 
% Simple tautologies: 1789
% Equational tautologies: 8
% Forward subsumptions: 14871
% Bw demodulations to eq. taut.: 33
% 
% Binary resolution: 8051
% Factoring: 6
% Forward superposition: 8797
% Backward superposition: 2167
% Self superposition: 3
% Equality resolution: 9
% 
% SAT solver clauses: 3165
% SAT solver unit clauses: 176
% SAT solver binary clauses: 615
% SAT solver learnt clauses: 24
% SAT solver learnt literals: 40
% 
% TWLsolver clauses: 3119
% TWLsolver calls for satisfiability: 65469
% 
% Memory used [KB]: 14072
% Time elapsed: 0.858 s
% ------------------------------
% ------------------------------
% Success in time 19.918 s
